[1358118108] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358118116] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358118134] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1358118134] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1358118188] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358118229] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1358118274] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358118396] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1358118414] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=90 ms
[1358118560] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358118656] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358118686] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358118697] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1358118849] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1358118949] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358119029] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358119029] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1358119045] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1358119178] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358119197] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358119259] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358119274] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358119377] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358119466] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1358119484] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1358119544] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358119557] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=92 ms
[1358119640] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358119659] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358119734] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=109 ms
[1358119754] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1358119770] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1358119846] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1358119929] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358120007] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=148 ms
[1358120007] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358120018] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=106 ms
[1358120100] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1358120194] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358120220] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358120244] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358120298] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=4 ms
[1358120298] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358120307] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1358120467] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=132 ms
[1358120477] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358120530] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358120647] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358120812] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1358120911] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1358120929] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358120998] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358121312] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=103 ms
[1358121643] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1358121676] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1358121895] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358121913] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358121968] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1358121973] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358122008] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=116 ms
[1358122183] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358122256] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1358122263] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358122267] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358122299] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1358122556] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358122595] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=162 ms
[1358122601] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1358122603] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358122664] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358122664] HOST ALERT: cassius.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358122665] SERVICE ALERT: sda-test.sensors.elex.be;XRAS;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358122666] SERVICE ALERT: king-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358122685] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=110 ms
[1358122695] SERVICE ALERT: pintail2-test.sensors.elex.be;MEMORY;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358122703] HOST ALERT: dlog2-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358122703] HOST ALERT: dlog3-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358122703] HOST ALERT: pintail-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358122717] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358122783] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=180 ms
[1358122823] HOST ALERT: cassius.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 12.24 ms
[1358122844] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358122863] HOST ALERT: dlog2-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358122863] HOST ALERT: dlog3-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358122863] HOST ALERT: pintail-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358122883] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=19 ms
[1358122883] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358122893] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358122933] SERVICE ALERT: pintail2-test.sensors.elex.be;MEMORY;OK;SOFT;2;Memory OK - 90.2% (22306756 kB) free
[1358122944] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;OK;SOFT;2;FTP OK - 0.021 second response time on port 21 [220 Welcome to the Melexis Data Exchange FTP service.]
[1358122949] SERVICE ALERT: sda-test.sensors.elex.be;XRAS;OK;SOFT;2;TCP OK - 0.018 second response time on port 6933
[1358122954] SERVICE ALERT: king-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1358122970] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=102 ms
[1358122994] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358123049] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358123070] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358123173] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358123257] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=100 ms
[1358123287] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1358123306] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1358123338] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358123463] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1358123593] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1358123605] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=14 ms
[1358123625] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=8 ms
[1358123645] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358123764] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358123883] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358123893] SERVICE ALERT: esb-b-test.kuching.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=10.02% system=4.61% iowait=77.96% idle=7.41%
[1358123907] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (55.9037900874636)
[1358123933] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358123953] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1358124053] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358124183] SERVICE ALERT: esb-b-test.kuching.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.00% idle=100.00%
[1358124238] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=92 ms
[1358124333] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1358124526] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358124573] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1358124665] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358124836] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358124848] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1358124868] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=161 ms
[1358124975] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1358124985] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358125015] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358125036] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358125184] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358125197] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358125217] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358125296] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=130 ms
[1358125298] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358125326] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358125506] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1358125589] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358125590] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358125854] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358125937] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358126055] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358126144] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1358126195] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 22.49, 18.89, 13.47
[1358126224] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358126294] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.60% system=0.00% iowait=80.20% idle=19.20%
[1358126344] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358126434] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358126479] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 17.62, 19.43, 15.26
[1358126487] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1358126566] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358126584] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.20% system=0.40% iowait=7.40% idle=92.00%
[1358126635] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358126688] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358126717] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358126768] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 14.49, 17.32, 15.52
[1358126770] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358126854] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358126918] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358126936] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358126976] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358127058] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358127058] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358127108] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1358127221] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358127264] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358127344] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358127414] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358127459] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 6.05, 9.10, 12.46
[1358127514] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358127554] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358127636] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358127694] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358127736] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358127749] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358127834] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1358127958] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358128018] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.9037900874636)
[1358128097] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358128124] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=46 ms
[1358128149] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 12.75, 14.54, 13.41
[1358128155] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358128244] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1358128326] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358128380] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358128437] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 12.55, 12.59, 12.80
[1358128438] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358128543] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=48 ms
[1358128584] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358128605] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358128666] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358128724] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1358128764] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358128784] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 12.67, 12.88, 13.02
[1358128788] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358128881] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=14 ms
[1358128934] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358128954] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=134 ms
[1358129053] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358129073] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1358129077] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358129132] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 5.90, 9.05, 11.35
[1358129132] SERVICE FLAPPING ALERT: dvorak-test.colo.elex.be;LOAD;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1358129174] SERVICE ALERT: pintaildb-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 5.46, 3.83, 3.05
[1358129221] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358129224] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1358129245] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358129331] HOST ALERT: tristram-test.tess.elex.be;DOWN;SOFT;1;Server down!
[1358129351] HOST ALERT: tristram-test.tess.elex.be;UP;SOFT;2;Server alive
[1358129364] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358129464] SERVICE ALERT: pintaildb-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 2.88, 3.49, 3.09
[1358129511] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=59 ms
[1358129571] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358129594] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358129713] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358129861] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1358129861] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358129881] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358130003] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358130062] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=126 ms
[1358130211] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358130211] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358130344] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358130353] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358130501] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358130501] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1358130629] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358130643] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1358130791] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1358130791] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=46 ms
[1358130916] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358130931] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358130993] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358131016] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 20.0% (3293284 kB) free
[1358131046] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=108 ms
[1358131141] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358131141] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1358131221] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358131267] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358131283] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358131311] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;SOFT;2;Memory OK - 20.5% (3376284 kB) free
[1358131333] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358131554] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358131571] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358131631] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=94 ms
[1358131634] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358131841] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=95 ms
[1358131854] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358131855] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358131861] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1358131971] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=80 ms
[1358132142] SERVICE ALERT: king-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358132186] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.9037900874636)
[1358132204] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1358132205] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358132244] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358132376] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1358132477] SERVICE ALERT: king-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1358132597] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358132604] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358132630] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358132734] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1358132887] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358132894] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1358132918] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358133018] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=113 ms
[1358133085] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1358133177] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1358133234] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358133255] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1358133311] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1358133370] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1358133526] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1358133676] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358133687] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=59 ms
[1358133845] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358133848] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1358133914] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=12 ms
[1358133996] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358134056] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1358134148] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358134277] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358134359] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358134586] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1358134677] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1358134787] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358134795] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358134882] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358134896] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1358134977] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=71 ms
[1358135080] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358135084] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1358135094] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1358135167] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=6 ms
[1358135256] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;WARNING;SOFT;1;WARNING - QueueSize=15
[1358135296] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1358135327] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=13 ms
[1358135368] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1358135374] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358135377] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1358135414] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358135517] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358135545] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358135640] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1358135675] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1358135718] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1358135768] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358135858] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358135942] SERVICE FLAPPING ALERT: dvorak-test.colo.elex.be;LOAD;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358135983] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358135991] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358136021] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358136053] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1358136093] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358136183] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358136222] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358136288] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358136318] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1358136338] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1358136388] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1358136471] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358136518] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358136624] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358136638] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358136665] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=105 ms
[1358136768] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358136914] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1358136928] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358136998] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=227 ms
[1358137013] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1358137031] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=153 ms
[1358137101] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358137111] SERVICE ALERT: murray-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 25.37 ms
[1358137111] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358137210] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=190 ms
[1358137244] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1358137265] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358137289] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=20 ms
[1358137322] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;2;P1.PL WARNING - io read latency=42 ms
[1358137382] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1358137384] SERVICE ALERT: murray-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 16.41 ms
[1358137400] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358137536] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=67 ms
[1358137558] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358137568] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=76 ms
[1358137613] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=6 ms
[1358137828] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1358138378] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=381 ms
[1358138414] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=123 ms
[1358138478] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358138668] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1358138702] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=101 ms
[1358138770] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358138959] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1358138993] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1358139058] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=58 ms
[1358139189] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358139273] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358139357] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358139406] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1358139407] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358139475] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1358139540] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358139618] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358139628] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358139708] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358139729] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358139777] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1358139790] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358139790] SERVICE ALERT: murray-test.sensors.elex.be;SSHBACK;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358139805] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=1021 ms
[1358139827] HOST ALERT: contactmonitoring-test.sensors.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1358139827] HOST ALERT: murray-test.sensors.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1358139831] SERVICE ALERT: dss-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358139836] SERVICE ALERT: pintail-fb-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358139836] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1358139837] SERVICE ALERT: dss-test.sensors.elex.be;IOSTAT;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358139845] HOST ALERT: cassius.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358139846] SERVICE ALERT: dlog1-test.sensors.elex.be;CFENGINE;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358139850] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=161 ms
[1358139852] SERVICE ALERT: postprocess-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 12.98 ms
[1358139852] SERVICE ALERT: pintail1-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 9.36 ms
[1358139853] SERVICE ALERT: pintail-fb-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 83%, RTA = 12.77 ms
[1358139856] HOST ALERT: contactmonitoring-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358139856] HOST ALERT: murray-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358139856] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1358139865] HOST ALERT: cassius.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.13 ms
[1358139948] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1358139955] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358140043] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=60 ms
[1358140043] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358140099] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_4.3p2 Debian-9etch3 (protocol 2.0)
[1358140100] SERVICE ALERT: murray-test.sensors.elex.be;SSHBACK;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358140100] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=125 ms
[1358140101] SERVICE ALERT: dss-test.sensors.elex.be;IOSTAT;OK;SOFT;2;SDA:rsec/s=,wsec/s=,%util=:SDB:rsec/s=0,wsec/s=0,%util=0
[1358140101] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=343489.00 MB (72.96%)
[1358140116] SERVICE ALERT: dlog1-test.sensors.elex.be;CFENGINE;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1358140120] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358140126] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358140140] SERVICE ALERT: dss-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1358140146] SERVICE ALERT: postprocess-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.92 ms
[1358140146] SERVICE ALERT: pintail-fb-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 11.61 ms
[1358140146] SERVICE ALERT: pintail1-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 11.58 ms
[1358140146] SERVICE ALERT: pintail-fb-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358140169] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358140275] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358140356] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358140422] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1358140423] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1358140437] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358140542] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358140555] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=57 ms
[1358140651] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358140707] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358140707] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1358140805] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358140867] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358140875] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358140875] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=4 ms
[1358140973] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358141089] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358141131] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358141206] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=107 ms
[1358141211] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1358141211] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358141427] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1358141473] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=174 ms
[1358141536] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1358141539] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358141811] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358141856] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358141858] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1358142036] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358142139] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358142246] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358142247] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358142359] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1358142420] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358142569] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358142629] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358142689] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1358142690] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358142747] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358142959] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358142963] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358143018] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358143020] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1358143082] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1358143331] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358143340] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358143368] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358143418] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358143424] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=72 ms
[1358143548] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358143741] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358143761] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=74 ms
[1358143888] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358143928] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358144049] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358144068] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1358144170] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1358144178] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1358144198] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1358144213] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358144339] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1358144412] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1358144452] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358144471] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=141 ms
[1358144678] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358144738] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1358144759] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358144760] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1358144808] SERVICE ALERT: arwen.printers.erfurt.elex.be;PRINTER_STATUS;OK;HARD;3;OK,C-KIT DCC2725 C is at 59% OK,C-KIT DCC2725 M is at 42% OK,C-KIT DCC2725 Y is at 38% OK,C-KIT DCC2725 K is at 80% OK,Waste Toner Box is FULL
[1358144962] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358145028] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1358145053] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=73 ms
[1358145108] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358145219] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358145342] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=185 ms
[1358145391] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358145490] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358145512] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358145692] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=46 ms
[1358145778] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358145953] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1358146048] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1358146053] SERVICE ALERT: openlink-test.sofia.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=2.20% system=2.81% iowait=78.56% idle=%
[1358146091] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=122 ms
[1358146243] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358146348] SERVICE ALERT: openlink-test.sofia.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.20% idle=%
[1358146382] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1358146388] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358146482] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358146728] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1358146738] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358146825] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358147007] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358147026] HOST ALERT: hma-lp.sofia.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 47.82 ms
[1358147031] SERVICE ALERT: hma-lp.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 47.35 ms
[1358147089] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358147175] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358147299] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358147358] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358147375] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358147467] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358147570] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358147648] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358147651] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1358147665] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1358147735] HOST ALERT: hma-lp.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 49.19 ms
[1358147735] HOST FLAPPING ALERT: hma-lp.sofia.elex.be;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1358147777] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=99 ms
[1358147879] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;WARNING;SOFT;1;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1358147935] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358148009] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358148127] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=73 ms
[1358148167] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358148168] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;WARNING;SOFT;2;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1358148298] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358148299] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358148362] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1358148400] SERVICE ALERT: ortmelei.apps.elex.be;TNS;CRITICAL;SOFT;1;Connection refused
[1358148456] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1358148458] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;WARNING;HARD;3;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1358148519] SERVICE ALERT: ortmelei.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;CRITICAL: SQL query failed!
[1358148585] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1358148587] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1358148621] SERVICE ALERT: ortmelei.apps.elex.be;TNSLSNR;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1358148621] SERVICE ALERT: ortmelei.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1358148635] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1358148648] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358148686] SERVICE ALERT: ortmelei.apps.elex.be;TNS;CRITICAL;SOFT;2;Connection refused
[1358148746] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1358148809] SERVICE ALERT: ortmelei.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;2;CRITICAL: SQL query failed!
[1358148819] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1358148910] SERVICE ALERT: ortmelei.apps.elex.be;TNSLSNR;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1358148911] SERVICE ALERT: ortmelei.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1358148915] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358148927] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358148935] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=176 ms
[1358148969] SERVICE ALERT: ortmelei.apps.elex.be;TNS;CRITICAL;HARD;3;Connection refused
[1358148997] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358149036] SERVICE ALERT: ortmelei.apps.elex.be;DISK;OK;HARD;3;OK - all disk
[1358149086] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 27.17, 13.01, 8.80
[1358149095] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1358149109] SERVICE ALERT: ortmelei.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1358149129] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358149131] SERVICE ALERT: ortmelei.apps.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 10.67, 4.71, 1.99
[1358149208] SERVICE ALERT: ortmelei.apps.elex.be;TNSLSNR;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1358149219] SERVICE ALERT: ortmelei.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1358149244] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358149246] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=190 ms
[1358149249] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=106 ms
[1358149313] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358149398] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 34.94, 26.61, 15.94
[1358149443] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358149445] SERVICE ALERT: ortmelei.apps.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 11.78, 9.31, 4.82
[1358149543] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358149553] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL CRITICAL - io read latency=132 ms
[1358149556] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=90 ms
[1358149584] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358149600] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1358149693] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 38.90, 34.11, 22.09
[1358149755] SERVICE ALERT: ortmelei.apps.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 11.86, 10.85, 6.74
[1358149777] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358149789] SERVICE ALERT: ortmelei.apps.elex.be;DISK;WARNING;SOFT;1;[ /opt/oracle ] -> Size: 392G Used: 297G Avail: 76G Free: (20%)
[1358149860] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1358149871] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL CRITICAL - io read latency=154 ms
[1358149911] SERVICE ALERT: esb-a-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 3.39, 5.14, 3.50
[1358149933] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358149949] SERVICE ALERT: bipmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=6.41% system=0.40% iowait=92.99% idle=%
[1358150010] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=99 ms
[1358150132] SERVICE FLAPPING ALERT: esb-b-test.kuching.elex.be;CFAGENT;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358150174] SERVICE ALERT: ortmelei.apps.elex.be;DISK;WARNING;SOFT;2;[ /opt/oracle ] -> Size: 392G Used: 308G Avail: 65G Free: (17%)
[1358150244] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1358150292] SERVICE ALERT: esb-a-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 6.27, 4.47, 3.58
[1358150322] SERVICE ALERT: bipmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.20% system=0.40% iowait=12.97% idle=%
[1358150422] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358150457] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1358150493] SERVICE ALERT: ortmelei.apps.elex.be;DISK;WARNING;HARD;3;[ /opt/oracle ] -> Size: 392G Used: 318G Avail: 55G Free: (14%)
[1358150542] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=47 ms
[1358150542] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 63.04 ms
[1358150576] SERVICE ALERT: esb-a-test.colo.elex.be;LOAD;OK;SOFT;3;OK - load average: 2.74, 3.64, 3.50
[1358150752] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358150756] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358150772] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358150815] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 54%, RTA = 87.52 ms
[1358150832] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=7 ms
[1358150935] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1358150983] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358151004] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;HARD;3;P1.PL WARNING - io read latency=76 ms
[1358151045] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358151054] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1358151069] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1358151072] SERVICE ALERT: riverbed.colo.elex.be;RIVERBED_STATUS;WARNING;SOFT;1;Steelhead 2050 (2050L): Healthy, optimisation service: runninghealth_status=Healthy, service_status=running,
[1358151187] SERVICE ALERT: ortmelei.apps.elex.be;DISK;CRITICAL;HARD;3;[ /opt/oracle ] -> Size: 392G Used: 348G Avail: 25G Free: (6%)
[1358151284] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358151316] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358151358] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1358151365] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=37 ms
[1358151375] SERVICE ALERT: riverbed.colo.elex.be;RIVERBED_STATUS;OK;SOFT;2;Steelhead 2050 (2050L): Healthy, optimisation service: runninghealth_status=Healthy, service_status=running,
[1358151406] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358151533] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 73.60 ms
[1358151605] SERVICE ALERT: ortmelei.apps.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 1.54, 5.48, 8.08
[1358151605] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358151634] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358151645] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=114 ms
[1358151695] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1358151777] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358151816] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 73%, RTA = 59.02 ms
[1358151924] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358151925] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Unable to load user properties file /usr/share/apache-servicemix-4.4.1-fuse-07-11/etc/users.properties
[1358151935] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=85 ms
[1358151954] SERVICE ALERT: ortmelei.apps.elex.be;LOAD;OK;HARD;3;OK - load average: 1.42, 3.01, 6.14
[1358151983] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1358151984] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358151985] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358151987] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358151997] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152007] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152008] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152014] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152015] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152015] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152016] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152016] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152017] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152018] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152024] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152024] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152024] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152033] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152037] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152054] SERVICE ALERT: ortmelei.apps.elex.be;EM_ALERTS;OK;HARD;3;OK: no outstanding alert(s) in EM
[1358152056] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1358152064] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1358152073] HOST FLAPPING ALERT: hma-lp.sofia.elex.be;STOPPED; Host appears to have stopped flapping (0.0% change < 5.0% threshold)
[1358152085] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152088] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152094] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152094] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152094] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152095] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152128] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152128] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152129] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152134] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152134] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152134] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152134] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152134] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152134] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152135] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152135] SERVICE ALERT: ortmelei.apps.elex.be;TNSLSNR;OK;HARD;3;PROCS OK: 1 process with UID = 700 (oracle), args 'tnslsnr'
[1358152136] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152137] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152143] SERVICE ALERT: ortmelei.apps.elex.be;ARCHIVELOG;OK;HARD;3;OK - ARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.3.0 - 64bit Production)
[1358152144] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152144] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152147] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152153] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358152154] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152154] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152154] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152156] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152166] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152166] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152167] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152174] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152176] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152183] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152184] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152184] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152187] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152187] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152206] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152207] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152215] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1358152226] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=69 ms
[1358152234] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152235] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152246] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152246] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152248] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152263] SERVICE ALERT: ortmelei.apps.elex.be;TNS;OK;HARD;3;TCP OK - 0.000 second response time on port 1521
[1358152265] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152274] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152277] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152277] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152277] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152278] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152283] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection reset
[1358152283] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152287] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152294] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152303] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection reset
[1358152304] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152304] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152305] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152308] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152309] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152309] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152309] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152310] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152310] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152310] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152313] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection reset
[1358152313] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152314] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152315] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152315] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152315] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152315] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152315] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152316] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152326] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152326] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358152328] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152328] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152333] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358152356] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=81 ms
[1358152374] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152384] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152384] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152387] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152388] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152388] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152393] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152403] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152404] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152413] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152413] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152425] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152426] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152434] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152434] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152434] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152434] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152434] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152443] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152444] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152444] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152445] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358152445] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358152455] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152455] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152455] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152456] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152463] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152463] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152464] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152475] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152475] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152483] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152483] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152484] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152484] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152484] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152504] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152504] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152523] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152534] SERVICE ALERT: esb-a-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.69% system=0.30% iowait=82.24% idle=16.77%
[1358152535] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152535] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152543] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152544] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152544] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152553] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152557] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152562] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152570] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152570] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152574] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152577] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152577] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152577] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152580] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152580] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1358152596] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358152599] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152600] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152600] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152600] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152602] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152604] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152610] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152614] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152615] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152615] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152616] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152618] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152618] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152618] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152618] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152619] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152620] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152620] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152642] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 14.78, 10.36, 16.13
[1358152644] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152644] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1358152644] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152645] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152657] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1358152704] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152718] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152718] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152718] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152719] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152719] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152722] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152723] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152723] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152723] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152733] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152736] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152745] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=101 ms
[1358152752] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152756] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152757] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152762] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152764] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152764] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152765] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152765] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152775] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152775] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152776] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358152792] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152792] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152793] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152798] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358152799] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152799] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152802] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152814] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152814] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152819] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152820] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152820] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152821] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152821] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152837] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358152839] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152839] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152862] SERVICE ALERT: esb-a-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=7.75% system=3.17% iowait=0.11% idle=88.97%
[1358152872] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152873] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152876] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152886] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152892] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152892] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152903] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152903] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152904] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152904] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152907] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152923] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1358152932] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152933] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152946] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152952] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152953] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152953] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152963] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358152973] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358152982] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1358152983] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358152983] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1358153005] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.69, 6.61, 12.56
[1358153025] HOST ALERT: nki-laptop.sofia.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 45.66 ms
[1358153096] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358153102] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1358153103] SERVICE ALERT: nki-laptop.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 46.26 ms
[1358153132] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=86 ms
[1358153186] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358153203] HOST ALERT: dlog1-test.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358153222] HOST ALERT: dlog1-test.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.90 ms
[1358153234] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;Connection refused
[1358153256] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1358153264] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358153335] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=219 ms
[1358153359] SERVICE ALERT: ewaf-test.colo.elex.be;MEMORY;OK;HARD;3;Memory OK - 65.6% (2665236 kB) free
[1358153413] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153413] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153413] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153413] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153413] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153414] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153415] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153415] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153416] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1358153416] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153416] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153416] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153433] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153433] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153433] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153436] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153436] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153437] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153437] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153438] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 29.80, 12.41, 4.85
[1358153445] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1358153452] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=101 ms
[1358153453] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153453] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153453] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153457] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153458] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153472] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153472] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153472] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358153472] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153472] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153472] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153492] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153493] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153497] SERVICE ALERT: postprocess-test.sensors.elex.be;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358153516] HOST ALERT: sda32-test.sensors.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1358153516] HOST ALERT: postprocess-test.sensors.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1358153516] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153516] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153518] SERVICE ALERT: king-test.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358153523] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 0.000 second response time on port 61616
[1358153524] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153532] HOST ALERT: postprocess-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358153533] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153534] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153536] SERVICE ALERT: dlog1-test.sensors.elex.be;DLD;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 2193 (dlog), args '/usr/bin/dld'
[1358153553] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153556] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153558] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153559] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1358153562] HOST ALERT: sda32-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358153577] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153578] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153583] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153584] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153585] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153586] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153596] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153596] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153596] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153604] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153606] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153606] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358153608] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153612] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153623] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153623] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153626] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1358153626] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153626] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153632] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358153642] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153643] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153643] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153644] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153644] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153644] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153645] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153645] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153645] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153649] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153649] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358153649] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153663] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153664] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358153785] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 0.97, 4.69, 3.79
[1358153794] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1358153833] SERVICE ALERT: postprocess-test.sensors.elex.be;HTTP;OK;SOFT;2;HTTP OK - HTTP/1.1 302 Found - 0.036 second response time
[1358153843] SERVICE ALERT: king-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.45 ms
[1358153862] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=86 ms
[1358153882] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153882] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153883] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153883] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153883] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153884] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153885] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153885] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153886] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153886] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153886] SERVICE ALERT: dlog1-test.sensors.elex.be;DLD;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 2193 (dlog), args '/usr/bin/dld'
[1358153886] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153893] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153904] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153904] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153907] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153907] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153908] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153908] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358153913] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Unable to load user properties file /usr/share/apache-servicemix-4.4.1-fuse-07-11/etc/users.properties
[1358153932] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358153948] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358153968] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1358153968] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;OK;SOFT;2;FTP OK - 0.021 second response time on port 21 [220 Welcome to the Melexis Data Exchange FTP service.]
[1358154043] SERVICE ALERT: esb-a-test.kuching.elex.be;MEMORY;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358154043] SERVICE ALERT: esb-a-test.kuching.elex.be;DISK;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358154094] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358154098] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;3;OK - load average: 0.04, 1.83, 2.80
[1358154113] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1358154118] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154132] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154133] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154133] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154134] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358154134] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358154135] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154136] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358154136] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358154136] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358154145] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154145] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154145] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154145] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154152] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154152] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358154167] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358154173] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358154182] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358154214] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1358154215] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1358154298] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358154317] SERVICE ALERT: esb-b-test.kuching.elex.be;CPU_STATS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358154318] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=122 ms
[1358154342] SERVICE ALERT: esb-a-test.kuching.elex.be;MEMORY;OK;SOFT;2;Memory OK - 86.8% (1789052 kB) free
[1358154342] SERVICE ALERT: esb-a-test.kuching.elex.be;DISK;OK;SOFT;2;OK - all disk
[1358154346] SERVICE ALERT: ftp-test.kuching.elex.be;FTP_FETCH;UNKNOWN;HARD;3;**ePN /usr/lib/nagios/plugins/check_ftp.pl: "ERROR: Cannot conect to 10.60.112.53".
[1358154386] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358154413] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358154457] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1358154472] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358154585] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358154612] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358154613] SERVICE ALERT: esb-b-test.kuching.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.20% system=0.80% iowait=1.60% idle=97.39%
[1358154706] SERVICE ALERT: ftp-test.kuching.elex.be;FTP_FETCH;WARNING;HARD;3;WARNING: server says: Failed to open file.
[1358154707] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=10 ms
[1358154907] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1358154934] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358154962] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358155053] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=42 ms
[1358155085] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358155202] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358155224] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1358155252] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=95 ms
[1358155343] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358155375] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358155494] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1358155517] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=72 ms
[1358155536] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1358155555] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358155637] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358155678] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358155724] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358155726] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 20.0% (3290296 kB) free
[1358155813] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358155842] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1358155846] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358155852] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358155866] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1358155868] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358155925] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358155955] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358156012] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358156013] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;SOFT;2;Memory OK - 20.3% (3346632 kB) free
[1358156104] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358156142] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358156158] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358156215] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=14 ms
[1358156235] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1358156362] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358156395] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1358156446] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=90 ms
[1358156485] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 59.82 ms
[1358156512] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;UNKNOWN;SOFT;1;UNKOWN: error reading value
[1358156566] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358156584] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=100 ms
[1358156652] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358156707] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 19.5% (3209368 kB) free
[1358156755] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;OK;SOFT;2;OK: 19.47
[1358156758] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 60.22 ms
[1358156798] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358156842] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358156858] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358156879] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358156896] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358156996] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;SOFT;2;Memory OK - 20.1% (3308720 kB) free
[1358157002] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358157093] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=73 ms
[1358157106] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358157123] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358157132] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358157182] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358157227] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358157326] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358157386] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358157434] SERVICE ALERT: hma-lp.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 47.46 ms
[1358157444] HOST ALERT: hma-lp.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 49.30 ms
[1358157448] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=39 ms
[1358157455] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358157534] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 77.95 ms
[1358157541] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358157579] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=93 ms
[1358157654] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358157709] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358157778] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358157784] SERVICE ALERT: king-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 9.15 ms
[1358157786] SERVICE ALERT: pintail1-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358157786] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358157787] SERVICE ALERT: ftp-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358157797] SERVICE ALERT: sda32-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358157798] SERVICE ALERT: pintail2-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358157805] HOST ALERT: dlog1-test.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358157814] HOST ALERT: dlog1-test.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.11 ms
[1358157834] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358157877] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=82 ms
[1358157890] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 59.32 ms
[1358157911] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358157914] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;UNKNOWN;SOFT;1;UNKOWN: error reading value
[1358157949] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358157949] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358158054] SERVICE ALERT: king-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 10.77 ms
[1358158057] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358158057] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1358158067] SERVICE ALERT: pintail1-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358158067] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358158068] SERVICE ALERT: ftp-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358158077] SERVICE ALERT: sda32-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.4 (protocol 2.0)
[1358158077] SERVICE ALERT: pintail2-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358158119] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358158155] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;OK;SOFT;2;OK: 19.59
[1358158164] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 57.85 ms
[1358158198] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1358158226] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1358158238] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358158238] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1358158344] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358158396] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358158468] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358158548] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358158575] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=105 ms
[1358158588] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=100 ms
[1358158674] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;OK;SOFT;2;TCP OK - 0.009 second response time on port 6913
[1358158690] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1358158757] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358158757] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358158797] SERVICE ALERT: ftp-test.kuching.elex.be;INODES;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358158807] SERVICE ALERT: esb-a-test.kuching.elex.be;AT;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358158834] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=85 ms
[1358158850] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358158864] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 58.31 ms
[1358158868] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=99 ms
[1358158937] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1358158977] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358159047] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358159088] SERVICE ALERT: ftp-test.kuching.elex.be;INODES;OK;SOFT;2;OK - all disk
[1358159097] SERVICE ALERT: esb-a-test.kuching.elex.be;AT;OK;SOFT;2;PROCS OK: 1 process with UID = 1 (daemon), args '/usr/sbin/atd'
[1358159115] SERVICE ALERT: esb-a-test.kuching.elex.be;IOSTAT;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358159124] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358159148] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358159148] HOST FLAPPING ALERT: temp1.sofia.elex.be;STARTED; Host appears to have started flapping (27.4% change > 20.0% threshold)
[1358159156] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=76 ms
[1358159214] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 60.81 ms
[1358159284] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358159324] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358159364] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1358159404] SERVICE ALERT: esb-a-test.kuching.elex.be;IOSTAT;OK;SOFT;2;SDA:rsec/s=0,wsec/s=0,%util=0:SDB:rsec/s=0,wsec/s=0,%util=0
[1358159474] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358159574] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1358159614] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=85 ms
[1358159650] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358159756] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=88 ms
[1358159806] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358159864] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358159904] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358159938] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1358160054] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1358160095] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358160096] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358160197] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 19.5% (3216144 kB) free
[1358160198] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358160218] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1358160387] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358160397] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1358160397] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1358160487] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;2;Memory WARNING - 19.2% (3160464 kB) free
[1358160564] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1358160779] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;HARD;3;Memory WARNING - 19.7% (3250644 kB) free
[1358160849] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358161033] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 59.30 ms
[1358161153] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358161187] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;HARD;3;Memory OK - 20.6% (3396560 kB) free
[1358161205] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358161259] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=135 ms
[1358161305] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=120 ms
[1358161338] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 58.24 ms
[1358161440] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1358161496] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358161547] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358161603] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358161614] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 58%, RTA = 62.48 ms
[1358161707] HOST FLAPPING ALERT: temp1.sofia.elex.be;STOPPED; Host appears to have stopped flapping (0.0% change < 5.0% threshold)
[1358161733] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1358161791] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1358161949] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358162144] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358162146] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358162238] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358162417] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1358162435] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358162437] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358162438] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358162597] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1358162703] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358162726] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358162767] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358162784] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1358162794] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1358162888] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358162906] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=54 ms
[1358162936] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=108 ms
[1358162969] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1358163014] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 58.03 ms
[1358163074] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358163076] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358163144] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358163194] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=13 ms
[1358163226] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1358163238] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358163373] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358163425] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358163435] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358163516] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1358163533] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358163653] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=81 ms
[1358163709] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358163727] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1358163877] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1358164003] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1358164046] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358164075] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358164167] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358164325] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 20.0% (3291500 kB) free
[1358164325] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1358164364] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358164406] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358164423] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 81%, RTA = 58.84 ms
[1358164443] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358164619] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358164647] SERVICE ALERT: eva.printers.tess.elex.be;PRINTER_STATUS;OK;HARD;3;OK,Toner is FULL
[1358164676] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;2;Memory WARNING - 20.0% (3293448 kB) free
[1358164711] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1358164752] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358164752] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358164756] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 61%, RTA = 62.07 ms
[1358164791] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358164917] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=112 ms
[1358164926] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=166 ms
[1358164966] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;SOFT;3;Memory OK - 20.2% (3327052 kB) free
[1358165078] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358165210] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1358165218] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358165457] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358165459] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358165469] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 58.98 ms
[1358165568] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1358165576] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358165667] HOST ALERT: dlog1-test.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358165669] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 19.7% (3243504 kB) free
[1358165677] SERVICE ALERT: sda-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358165678] SERVICE ALERT: postgres-stage.sensors.elex.be;POSTGRES_TCP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358165679] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358165681] SERVICE ALERT: mona.printers.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 9.77 ms
[1358165686] HOST ALERT: cassius.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358165696] HOST ALERT: dlog1-test.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.24 ms
[1358165706] HOST ALERT: cassius.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.17 ms
[1358165769] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358165776] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358165786] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1358165786] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 65.00 ms
[1358165840] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358165847] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1358165847] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358165848] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358165848] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358165848] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358165848] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358165849] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358165849] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358165849] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358165849] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358165849] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358165850] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358165857] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1358165857] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358165858] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1358165858] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358165858] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1358165866] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358165867] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358165867] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358165867] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358165867] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358165868] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358165889] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358165896] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1358165899] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358165937] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1358165937] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358165937] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358165939] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358165940] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358165946] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358165976] SERVICE ALERT: mona.printers.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 11.04 ms
[1358165981] SERVICE ALERT: sda-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze1 (protocol 2.0)
[1358165982] SERVICE ALERT: postgres-stage.sensors.elex.be;POSTGRES_TCP;OK;SOFT;2;TCP OK - 0.010 second response time on port 5432
[1358165987] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;OK;SOFT;2;FTP OK - 0.020 second response time on port 21 [220 Welcome to the Melexis Data Exchange FTP service.]
[1358165988] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;SOFT;2;Memory OK - 20.2% (3325400 kB) free
[1358166086] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1358166086] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 58%, RTA = 61.84 ms
[1358166096] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1358166100] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358166157] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166167] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166167] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166167] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166167] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166167] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166168] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166168] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166169] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166169] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166169] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166170] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166170] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166172] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166172] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166176] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166176] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166176] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166188] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166188] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166189] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166189] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166189] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 0.65, 4.81, 2.93
[1358166189] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166190] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166211] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166217] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166257] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166257] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166257] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166258] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358166260] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1358166262] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358166509] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.11, 1.67, 2.07
[1358166518] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358166578] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=117 ms
[1358166648] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1358166747] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=43 ms
[1358166820] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358166929] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1358166940] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=109 ms
[1358167037] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358167078] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=73 ms
[1358167201] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1358167206] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 58.45 ms
[1358167246] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1358167299] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358167397] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1358167516] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 73%, RTA = 60.90 ms
[1358167519] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358167576] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358167620] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358167679] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358167896] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358167900] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358167944] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358168008] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1358168260] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358168278] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1358168278] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358168326] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358168381] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1358168438] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1358168558] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1358168579] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358168646] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1358168697] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358168706] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358168752] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358169012] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358169016] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358169096] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 59.56 ms
[1358169132] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1358169330] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1358169347] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=49 ms
[1358169410] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 75%, RTA = 64.73 ms
[1358169489] SERVICE ALERT: viiperbi-test.apps.elex.be;DISK;OK;HARD;3;OK - all disk
[1358169519] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1358169658] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358169709] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1358169711] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358169817] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358169867] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358169869] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358169880] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358169992] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358170126] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1358170142] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 58.20 ms
[1358170160] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358170166] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1358170168] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=107 ms
[1358170176] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;UNKNOWN;SOFT;1;UNKOWN: error reading value
[1358170178] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358170229] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=59 ms
[1358170291] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1358170416] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358170417] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 54%, RTA = 66.67 ms
[1358170417] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;OK;SOFT;2;OK: 19.98
[1358170456] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1358170460] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358170510] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358170516] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=4 ms
[1358170618] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (55.2721088435374)
[1358170668] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358170798] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358170806] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1358170816] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358170968] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358171000] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1358171116] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358171119] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=105 ms
[1358171131] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358171136] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358171286] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1358171318] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358171439] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=83 ms
[1358171440] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358171448] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1358171496] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358171638] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358171831] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358171875] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358171882] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1358171884] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1358172071] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=69 ms
[1358172202] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358172213] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358172261] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358172264] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1358172381] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;2;P1.PL WARNING - io read latency=57 ms
[1358172498] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358172555] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358172555] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1358172674] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=21 ms
[1358172795] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358172853] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=102 ms
[1358172906] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358173144] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358173193] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358173203] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1358173227] HOST ALERT: tanja.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358173228] HOST ALERT: stella.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358173232] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 9.32 ms
[1358173232] HOST ALERT: cassius.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358173232] SERVICE ALERT: sda-test.sensors.elex.be;XRAS;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358173233] HOST ALERT: dlog2-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358173234] SERVICE ALERT: dss-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 9.41 ms
[1358173234] SERVICE ALERT: cassius.sensors.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 9.37 ms
[1358173234] SERVICE ALERT: king-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358173234] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358173246] HOST ALERT: king-test.sensors.elex.be;UNREACHABLE;SOFT;1;(Host Check Timed Out)
[1358173246] HOST ALERT: lechuck-test.sensors.elex.be;UNREACHABLE;SOFT;1;(Host Check Timed Out)
[1358173246] HOST ALERT: murray-test.sensors.elex.be;UNREACHABLE;SOFT;1;(Host Check Timed Out)
[1358173246] HOST ALERT: postgres-stage.sensors.elex.be;UNREACHABLE;SOFT;1;(Host Check Timed Out)
[1358173251] HOST ALERT: dlog1-test.sensors.elex.be;UNREACHABLE;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358173251] SERVICE ALERT: dlog-pp-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358173252] SERVICE ALERT: murray-test.sensors.elex.be;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1358173253] HOST ALERT: pintail-fb-test.sensors.elex.be;UNREACHABLE;SOFT;1;(Host Check Timed Out)
[1358173253] HOST ALERT: dlog3-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358173262] HOST ALERT: tanja.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 22.81 ms
[1358173271] HOST ALERT: lechuck-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358173271] HOST ALERT: postgres-stage.sensors.elex.be;UP;SOFT;2;Server alive
[1358173271] HOST ALERT: murray-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358173271] HOST ALERT: pintail-fb-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358173271] HOST ALERT: pintail-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358173271] HOST ALERT: dlog1-test.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.24 ms
[1358173271] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358173271] HOST ALERT: king-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358173281] HOST ALERT: cassius.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.17 ms
[1358173321] HOST ALERT: dlog2-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358173321] HOST ALERT: dlog3-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358173321] HOST ALERT: pintail-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358173341] HOST ALERT: stella.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.28 ms
[1358173435] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358173501] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.28 ms
[1358173511] SERVICE ALERT: dss-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.33 ms
[1358173522] SERVICE ALERT: sda-test.sensors.elex.be;XRAS;OK;SOFT;2;TCP OK - 0.009 second response time on port 6933
[1358173523] SERVICE ALERT: king-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1358173523] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;OK;SOFT;2;TCP OK - 0.009 second response time on port 6445
[1358173536] SERVICE ALERT: dlog-pp-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze2 (protocol 2.0)
[1358173554] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358173571] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358173572] SERVICE ALERT: cassius.sensors.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 9.38 ms
[1358173584] SERVICE ALERT: murray-test.sensors.elex.be;VSERVERS_MOUNTS;CRITICAL;SOFT;1;Total mounts missing: 2
[1358173602] SERVICE ALERT: murray-test.sensors.elex.be;SSH;OK;HARD;1;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358173624] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - HOST CPU Unknown error
[1358173624] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358173731] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1358173853] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=90 ms
[1358173881] SERVICE ALERT: murray-test.sensors.elex.be;VSERVERS_MOUNTS;CRITICAL;SOFT;2;Total mounts missing: 2
[1358173906] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358173922] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL CRITICAL - HOST CPU Unknown error
[1358173922] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1358174083] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358174146] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358174201] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=86 ms
[1358174206] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358174211] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1358174281] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358174314] SERVICE ALERT: murray-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358174375] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358174455] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358174494] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=58 ms
[1358174502] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=8 ms
[1358174502] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358174515] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 19.3% (3175840 kB) free
[1358174576] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358174581] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358174612] SERVICE ALERT: murray-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1358174691] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358174731] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 101.70, 40.70, 18.52
[1358174774] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1358174842] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;2;Memory WARNING - 19.8% (3259704 kB) free
[1358174862] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 102.75 ms
[1358174881] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1358174901] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1358174901] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1358174975] SERVICE ALERT: pintaildb-test.colo.elex.be;CFAGENT;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358175051] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 143.27, 113.12, 57.40
[1358175161] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 63.08 ms
[1358175162] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;SOFT;3;Memory OK - 20.8% (3432572 kB) free
[1358175244] SERVICE ALERT: pintaildb-test.colo.elex.be;CFAGENT;OK;SOFT;2;OK: Connection to cfengine master OK
[1358175273] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358175401] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1358175401] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (29.2% change >= 20.0% threshold)
[1358175401] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1358175424] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358175484] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358175544] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1358175544] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (29.2% change >= 20.0% threshold)
[1358175544] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1358175557] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358175557] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (64.9% change >= 20.0% threshold)
[1358175668] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358175679] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1358175688] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358175739] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358175755] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358175755] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358175935] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358176087] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358176100] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358176108] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358176115] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358176148] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1358176148] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1358176305] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358176309] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358176320] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358176409] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358176568] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 62.63 ms
[1358176575] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1358176578] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1358176606] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 9.43, 7.72, 18.57
[1358176676] SERVICE ALERT: ftp-test.erfurt.elex.be;CPU_STATS;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358176725] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176725] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176725] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176727] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176727] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176727] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176729] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358176737] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1358176747] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176747] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176747] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176747] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 58.03 ms
[1358176748] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176748] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176748] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176756] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1358176778] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358176789] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1358176790] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358176796] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 19.7% (3249760 kB) free
[1358176806] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358176806] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358176806] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358176808] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1358176827] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358176827] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1358176828] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358176828] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358176829] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1358176830] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1358176830] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358176830] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358176830] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1358176830] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358176831] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358176836] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358176867] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358176962] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 11.10, 9.15, 4.18
[1358176965] SERVICE ALERT: ftp-test.erfurt.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.20% system=0.00% iowait=0.00% idle=99.80%
[1358177057] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177057] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177057] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177059] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177059] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177060] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177076] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=97 ms
[1358177077] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358177086] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177087] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177087] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177088] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177088] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177088] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177096] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177118] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177127] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177127] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177136] SERVICE ALERT: atlan.erfurt.elex.be;MEMORY;OK;SOFT;2;Memory OK - 20.3% (3345792 kB) free
[1358177138] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177138] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177138] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177139] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177158] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177158] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177159] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177159] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177160] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177165] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177165] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177165] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177166] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177166] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177167] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177175] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358177199] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358177228] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1358177287] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.06, 3.05, 2.93
[1358177377] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.38, 4.84, 10.72
[1358177397] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1358177397] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358177556] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358177735] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=183 ms
[1358177786] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358177878] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=58 ms
[1358178099] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358178115] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1358178306] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358178375] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358178465] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1358178567] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358178596] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1358178668] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358178805] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358178849] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1358178856] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358178889] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1358179090] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;1;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358179149] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358179263] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=92 ms
[1358179382] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;2;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358179430] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358179441] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=94 ms
[1358179553] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1358179672] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358179719] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358179730] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=17 ms
[1358179792] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358179849] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358179859] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358180009] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1358180051] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=113 ms
[1358180081] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358180149] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358180192] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358180359] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1358180403] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=83 ms
[1358180410] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1358180489] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=105 ms
[1358180535] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358180629] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358180753] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=101 ms
[1358180773] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1358180779] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358180826] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358180849] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358180889] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358180919] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1358181061] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358181069] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (53.2312925170068)
[1358181103] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1358181124] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358181140] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358181189] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358181352] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358181402] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1358181430] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1358181480] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358181642] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1358181722] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358181733] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1358181771] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358181971] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1358181999] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358182019] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358182039] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1358182122] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=96 ms
[1358182261] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=82 ms
[1358182289] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358182310] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=8 ms
[1358182343] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1358182410] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358182443] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358182549] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358182608] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1358182689] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358182698] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1358182776] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1358182809] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1358182985] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358183006] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1358183106] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1358183265] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358183308] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1358183310] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1358183332] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358183397] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1358183427] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358183556] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358183617] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1358183620] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358183658] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358183686] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358183718] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358183776] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=67 ms
[1358183919] HOST ALERT: hma-lp.sofia.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1358183945] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358183977] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358184045] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358184068] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=12 ms
[1358184325] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358184335] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358184618] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358184705] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358184798] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358184908] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1358184995] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358185155] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358185259] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358185345] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358185440] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1358185547] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358185638] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358185666] SERVICE ALERT: openlink-test.sofia.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=2.00% system=1.20% iowait=70.54% idle=%
[1358185727] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358185729] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1358185928] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1358185957] SERVICE ALERT: openlink-test.sofia.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.00% idle=%
[1358186017] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358186045] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358186274] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358186341] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1358186401] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358186425] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358186492] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358186565] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358186656] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358186718] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358186753] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=104 ms
[1358186762] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358186843] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358186951] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1358187044] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358187061] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=101 ms
[1358187071] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358187141] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358187206] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358187242] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1358187266] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358187351] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358187364] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358187391] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358187481] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1358187498] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358187557] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358187681] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358187772] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358187787] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1358188141] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1358188403] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=121 ms
[1358188416] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358188643] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1358188643] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1358188692] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=107 ms
[1358188706] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358188734] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358188823] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358188985] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=65 ms
[1358189024] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358189113] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358189152] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358189342] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358189351] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358189381] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358189447] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1358189465] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358189645] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1358189673] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358189813] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1358189825] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358190006] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1358190076] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358190087] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=150 ms
[1358190113] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=88 ms
[1358190163] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358190178] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358190186] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358190363] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358190369] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358190385] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358190405] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=74 ms
[1358190414] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (55.9037900874636)
[1358190469] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1358190484] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358190656] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=76 ms
[1358190678] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1358190697] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358190766] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1358190947] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1358191025] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358191028] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358191233] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=44 ms
[1358191316] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358191397] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358191525] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=8 ms
[1358191573] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358191688] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358191864] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1358191907] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358192026] SERVICE ALERT: dss-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358192046] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358192087] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358192133] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358192163] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=76 ms
[1358192175] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358192204] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358192313] SERVICE ALERT: dss-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1358192336] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358192358] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1358192383] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1358192426] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358192465] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1358192497] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1358192514] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358192514] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1358192633] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1358192708] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358192733] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358192755] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=101 ms
[1358193004] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358193026] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358193107] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1358193217] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358193345] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358193398] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1358193424] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358193447] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1358193527] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358193573] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358193694] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358193725] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358193743] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358193756] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358193818] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1358193863] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1358193986] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358194025] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358194036] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1358194046] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358194114] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358194207] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358194264] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358194374] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358194503] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358194558] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1358194573] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358194674] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358194743] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1358194755] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358194823] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358194853] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358194854] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1358194854] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1358194869] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1358195025] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358195046] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358195117] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358195146] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358195163] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1358195206] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358195315] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358195403] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358195443] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1358195448] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358195504] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358195514] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1358195634] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1358195673] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358195696] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358195739] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=124 ms
[1358195819] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1358195867] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358195927] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1358195965] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1358196033] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=58 ms
[1358196045] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1358196113] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=60 ms
[1358196163] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358196216] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;HARD;3;CRITICAL: Unable to establish connection with cfengine master server
[1358196216] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1358196325] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358196336] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358196383] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=121 ms
[1358196404] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=96 ms
[1358196497] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358196515] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358196616] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358196664] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358196687] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358196736] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1358196785] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358196803] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1358196885] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1358196941] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1358197011] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358197032] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358197137] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358197147] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1358197157] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=90 ms
[1358197189] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358197220] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1358197388] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358197429] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 24.38, 16.10, 9.60
[1358197477] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1358197510] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;OK;SOFT;3;OK: Connection to cfengine master OK
[1358197680] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358197718] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 19.67, 19.13, 12.74
[1358197777] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358197842] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=109 ms
[1358197850] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358197860] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=100 ms
[1358198003] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 22.24, 20.19, 14.79
[1358198029] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358198117] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358198133] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=93 ms
[1358198147] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358198317] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358198408] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1358198421] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=12 ms
[1358198559] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=86 ms
[1358198687] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (55.2721088435374)
[1358198762] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1358198770] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358198909] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358199017] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358199053] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358199061] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=78 ms
[1358199260] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358199308] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358199308] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358199357] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=157 ms
[1358199554] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358199608] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358199608] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358199707] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358199752] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 5.41, 8.56, 11.12
[1358199842] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1358199888] HOST ALERT: mona.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358199888] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358199897] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1358199899] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358199907] HOST ALERT: dlog1-test.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358199909] SERVICE ALERT: lechuck-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358199919] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358199920] SERVICE ALERT: postgres-stage.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358199921] SERVICE ALERT: stella.printers.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358199921] SERVICE ALERT: postprocess-test.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358199927] HOST ALERT: cassius.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358199928] SERVICE ALERT: mona.printers.sensors.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 9.81 ms
[1358199932] HOST ALERT: dlog2-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358199932] HOST ALERT: dlog3-test.sensors.elex.be;UNREACHABLE;HARD;1;(Host Check Timed Out)
[1358199937] HOST ALERT: dlog1-test.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 14.09 ms
[1358199938] SERVICE ALERT: dss-test.sensors.elex.be;CFENGINE;CRITICAL;SOFT;1;(Service Check Timed Out)
[1358199947] HOST ALERT: cassius.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.20 ms
[1358199987] HOST ALERT: dlog2-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358199987] HOST ALERT: dlog3-test.sensors.elex.be;DOWN;HARD;1;(Host Check Timed Out)
[1358200059] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=9 ms
[1358200087] HOST ALERT: mona.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.58 ms
[1358200108] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358200172] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;OK;SOFT;2;TCP OK - 0.009 second response time on port 6445
[1358200177] SERVICE ALERT: dss-test.sensors.elex.be;CFENGINE;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1358200188] SERVICE ALERT: postgres-stage.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 12.53 ms
[1358200188] SERVICE ALERT: stella.printers.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 13.50 ms
[1358200188] SERVICE ALERT: postprocess-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 10.85 ms
[1358200191] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358200192] SERVICE ALERT: lechuck-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1358200193] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1358200198] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;OK;SOFT;2;TCP OK - 0.009 second response time on port 6933
[1358200249] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358200253] SERVICE ALERT: mona.printers.sensors.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 14.71 ms
[1358200397] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358200409] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=107 ms
[1358200598] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358200609] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1358200757] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1358200807] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358200827] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358200917] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=95 ms
[1358200997] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358201098] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1358201119] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1358201207] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1358201281] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358201389] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358201389] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1358201558] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358201631] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358201739] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358201769] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1358201808] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358201917] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358201999] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358202029] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358202058] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358202099] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358202222] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=43 ms
[1358202290] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358202349] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358202354] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358202375] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358202512] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1358202601] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358202647] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1358202667] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1358202800] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=104 ms
[1358202900] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1358202937] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=68 ms
[1358202959] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1358202961] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358202997] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358203048] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1358203091] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=94 ms
[1358203171] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358203188] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358203254] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=6 ms
[1358203287] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358203309] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358203377] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1358203398] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358203458] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358203501] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358203621] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358203637] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1358203661] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358203689] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358203730] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358203797] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358203888] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=130 ms
[1358203917] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=7 ms
[1358203951] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358203981] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1358204017] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358204050] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1358204253] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=85 ms
[1358204377] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms