[1360278017] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=122 ms
[1360278037] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360278038] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360278384] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360278446] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360278604] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=1.20% system=6.71% iowait=77.38% idle=14.71%
[1360278748] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360278766] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360278839] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=88 ms
[1360278914] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;SOFT;2;CPU IOWAIT WARNING : user=0.70% system=1.60% iowait=71.23% idle=26.47%
[1360278996] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 19.25, 12.79, 9.42
[1360279067] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360279155] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360279194] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360279232] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;HARD;3;CPU IOWAIT WARNING : user=27.17% system=1.30% iowait=70.23% idle=1.30%
[1360279316] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 12.30, 12.07, 10.21
[1360279384] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1360279508] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360279527] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360279611] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;HARD;3;CPU IOWAIT CRITICAL : user=6.49% system=1.10% iowait=82.92% idle=9.49%
[1360279611] SERVICE FLAPPING ALERT: orpmelei-test.apps.elex.be;CPU_STATS;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1360279754] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360279846] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=67 ms
[1360279896] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1360279984] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;HARD;3;CPU STATISTICS OK : user=16.18% system=0.50% iowait=35.46% idle=47.85%
[1360280070] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360280137] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=177 ms
[1360280165] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1360280216] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1360280454] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1360280535] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;LOAD;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1360280539] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1360280644] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360280645] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1360280666] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360280774] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=69 ms
[1360280915] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360280958] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360280959] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360280980] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=102 ms
[1360281235] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1360281295] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1360281337] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360281527] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1360281651] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360281724] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360281743] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1360281867] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360281979] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360282039] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360282078] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360282188] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360282292] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360282500] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=8 ms
[1360282676] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360282786] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360282796] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360282856] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1360282920] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360282990] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1360283100] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360283108] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360283170] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1360283208] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360283238] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1360283311] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1360283422] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1360283489] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1360283519] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360283610] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=42 ms
[1360283676] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1360283686] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1360283709] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360283728] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360283898] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360283932] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1360283988] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1360284003] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1360284026] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360284040] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360284217] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=78 ms
[1360284240] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360284316] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=68 ms
[1360284366] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360284380] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=99 ms
[1360284533] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=90 ms
[1360284556] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360284682] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360284700] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360284876] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1360284912] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1360284912] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1360284999] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1360285499] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360285740] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360285791] SERVICE ALERT: dvorak-test.colo.elex.be;DISK;OK;HARD;3;OK - all disk
[1360285818] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360285830] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1360285859] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 20.31, 14.99, 9.72
[1360285923] SERVICE ALERT: lenor-test.sofia.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 18.82, 12.35, 8.12
[1360286061] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=67 ms
[1360286083] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360286129] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1360286140] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360286148] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360286178] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 25.54, 23.49, 15.20
[1360286241] SERVICE ALERT: lenor-test.sofia.elex.be;LOAD;OK;SOFT;2;OK - load average: 3.29, 7.17, 7.22
[1360286380] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360286401] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360286461] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1360286469] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1360286499] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 17.93, 21.48, 16.92
[1360286757] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=88 ms
[1360286849] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1360287069] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1360287218] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360287229] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360287278] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360287314] SERVICE FLAPPING ALERT: orpmelei-test.apps.elex.be;CPU_STATS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360287389] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=102 ms
[1360287410] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1360287537] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360287541] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360287547] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360287599] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360287768] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1360287852] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360287898] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360287914] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360287917] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1360288010] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 6.67, 8.06, 11.45
[1360288147] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360288167] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360288217] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360288233] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360288233] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360288257] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360288468] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1360288487] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1360288540] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360288541] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter System Status;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360288551] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1360288557] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=64 ms
[1360288580] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360288689] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360288788] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1360288827] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1360288910] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter System Status;OK;SOFT;2;OK System-State (255) Normal
[1360288939] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1360289022] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360289070] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360289217] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=44 ms
[1360289229] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=115 ms
[1360289344] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360289381] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360289450] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360289531] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360289549] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=106 ms
[1360289767] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360289811] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360289870] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360290047] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360290089] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1360290129] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360290247] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360290360] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360290387] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=12 ms
[1360290560] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=136 ms
[1360290767] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360290838] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360290879] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1360291087] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360291159] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360291211] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360291261] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360291462] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=103 ms
[1360291530] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=67 ms
[1360291537] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1360291579] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=6 ms
[1360291642] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1360291760] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (55.2721088435374)
[1360291777] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360291847] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1360291857] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360291959] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1360291991] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360292011] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1360292012] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=56 ms
[1360292157] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=115 ms
[1360292170] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1360292281] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1360292312] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.49.32.16'
[1360292334] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=5 ms
[1360292393] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=107 ms
[1360292407] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (54.7862001943635)
[1360292541] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360292618] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1360292618] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1360292704] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.49.32.16'
[1360292772] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360292871] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=89 ms
[1360293048] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360293090] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1360293101] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360293152] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=43 ms
[1360293248] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360293368] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=74 ms
[1360293423] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360293478] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1360293540] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360293643] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360293708] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1360293781] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360293881] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360293890] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360293993] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360294128] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360294220] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360294348] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1360294528] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1360294574] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1360294742] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1360294912] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360295062] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360295080] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360295232] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360295375] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1360295398] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=73 ms
[1360295608] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1360295671] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360295718] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=99 ms
[1360295753] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1360295860] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360295928] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360295991] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=87 ms
[1360296123] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1360296148] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1360296210] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360296230] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.9037900874636)
[1360296359] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360296368] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1360296461] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=105 ms
[1360296503] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360296528] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360296551] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (53.2312925170068)
[1360296678] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1360296820] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360296841] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1360297002] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360297221] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1360297288] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=122 ms
[1360297288] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1360297504] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1360297541] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360297608] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1360297608] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360297752] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360297758] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360297860] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=121 ms
[1360297928] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360297980] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360298068] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360298074] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360298241] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360298299] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360298453] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1360298620] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=65 ms
[1360298640] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360298773] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360298960] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360298999] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1360299059] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360299088] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=50 ms
[1360299372] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360299465] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1360299465] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1360299833] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360300135] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360300187] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360300249] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1360300496] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360300537] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360300602] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1360300622] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1360300843] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (55.9037900874636)
[1360300954] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1360300975] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360301078] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360301217] SERVICE ALERT: starcraft-test.colo.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 17.9% (2943848 kB) free
[1360301383] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1360301452] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=98 ms
[1360301585] SERVICE ALERT: starcraft-test.colo.elex.be;MEMORY;OK;SOFT;2;Memory OK - 40.0% (6583092 kB) free
[1360301607] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360301784] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360301787] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1360301927] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1360302013] SERVICE ALERT: lechuck-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 13.08 ms
[1360302013] SERVICE ALERT: ftp-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 14.76 ms
[1360302014] SERVICE ALERT: tanja.printers.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 27.88 ms
[1360302103] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360302165] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360302246] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1360302312] SERVICE ALERT: lechuck-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.27 ms
[1360302312] SERVICE ALERT: ftp-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.29 ms
[1360302312] SERVICE ALERT: tanja.printers.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.47 ms
[1360302366] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360302424] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360302484] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=97 ms
[1360302565] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360302627] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=137 ms
[1360302687] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1360302749] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360302813] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360302885] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360302950] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360303074] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360303195] SERVICE ALERT: hma-lp.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 49.52 ms
[1360303202] HOST ALERT: hma-lp.sofia.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 49.17 ms
[1360303278] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1360303282] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1360303332] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360303602] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360303613] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360303624] SERVICE ALERT: batch-test.tess.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:batch-test: Can't stat /var/lib/cfengine2/masterfiles/files/cronjobs/templog-batch-test.tess.elex.be in copy
[1360303635] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360303692] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=12 ms
[1360303929] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360303943] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1360304263] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1360304452] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360304489] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360304643] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1360304689] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360304772] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1360304776] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360304812] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360305007] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360305093] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360305132] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1360305328] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1360305512] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360305706] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360305774] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360305832] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360305852] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1360305902] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360306025] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360306094] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1360306148] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1360306154] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360306222] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360306225] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360306344] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1360306412] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=68 ms
[1360306477] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360306552] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1360306664] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1360306784] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1360306794] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360306832] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;CRITICAL: SQL query failed!
[1360306847] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360306847] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360306911] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360306916] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360306929] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1360306935] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360306969] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360306978] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360307042] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360307042] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;OK;SOFT;2;PROCS OK: 1 process with UID = 59998 (applmgr), args 'tnslsnr'
[1360307155] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360307215] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;OK;SOFT;2;OK: no outstanding alert(s) in EM
[1360307236] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;OK;SOFT;2;OK - 'Process_transaction_interface' Completed Normal      08/02/2013 07:56
[1360307237] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360307237] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;OK;SOFT;2;OK - ARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.3.0 - 64bit Production)
[1360307265] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360307266] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;OK;SOFT;2;PROCS OK: 16 processes with UID = 59998 (applmgr), args 'Apache'
[1360307266] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1360307279] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360307300] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;OK;SOFT;2;PROCS OK: 1 process with UID = 59998 (applmgr), args 'FNDCRM'
[1360307300] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1360307365] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360307475] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360307548] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1360307585] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360307621] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1360307681] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1360307696] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360307746] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;WARNING;SOFT;1;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1360307868] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1360307905] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (55.2721088435374)
[1360307928] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360307967] SERVICE ALERT: mail-test.colo.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name %s was not found - nrpe
[1360308018] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360308061] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1360308062] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;WARNING;SOFT;2;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1360308067] SERVICE ALERT: itbi-test.colo.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360308160] SERVICE ALERT: crowd-xpeqt-test.colo.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name %s was not found - nrpe
[1360308180] SERVICE ALERT: nagios-test.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360308188] SERVICE ALERT: tchai-test.admin.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360308248] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360308250] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360308278] SERVICE ALERT: pfus2-test.colo.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360308286] SERVICE ALERT: mail-test.colo.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360308378] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=125 ms
[1360308378] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360308379] SERVICE ALERT: restoreoraclexe-test.colo.elex.be;CFAGENT;WARNING;HARD;3;WARNING: cfengine:restoreoraclexe-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/restoreoraclexe-test.colo.elex.be in copy
[1360308385] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 62.19 ms
[1360308387] SERVICE ALERT: itbi-test.colo.elex.be;NRPE;OK;SOFT;2;PROCS OK: 4 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360308397] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360308479] SERVICE ALERT: crowd-xpeqt-test.colo.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360308487] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360308497] SERVICE ALERT: orpmelei-test.apps.elex.be;ARCHIVELOG;WARNING;HARD;3;NRPE: Unable to read output
[1360308501] SERVICE ALERT: nagios-test.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360308506] SERVICE ALERT: tchai-test.admin.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360308560] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1360308595] SERVICE ALERT: pfus2-test.colo.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360308679] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360308697] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1360308698] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 86%, RTA = 63.78 ms
[1360308699] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360308710] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360308727] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 66.29 ms
[1360308816] SERVICE ALERT: cassius.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 9.66 ms
[1360308868] SERVICE ALERT: orpmelei-test.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - NOARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.1.0 - 64bit Production)
[1360308886] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1360308996] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360309015] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1360309065] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 72%, RTA = 69.34 ms
[1360309105] SERVICE ALERT: cassius.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 11.03 ms
[1360309151] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360309187] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360309256] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=103 ms
[1360309460] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;OK;SOFT;2;PROCS OK: 17 processes with UID = 59998 (applmgr), args 'Apache'
[1360309498] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;OK;SOFT;2;PROCS OK: 1 process with UID = 59998 (applmgr), args 'FNDCRM'
[1360309576] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=108 ms
[1360309609] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;SOFT;1;Connection refused
[1360309636] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;OK;HARD;3;OK: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Synchronization Successful. HRAgentB - Synchronization Successful. HRAgentC - Synchronization Successful.
[1360309687] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1360309755] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360309855] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1360309857] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360309897] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360309918] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=73 ms
[1360309949] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;SOFT;2;Connection refused
[1360309999] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;SOFT;1;CRITICAL - 'WIP_Move_Transaction_Manager' Pending Normal	      08/02/2013 08:41
[1360310028] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1360310077] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360310088] SERVICE ALERT: bipmelei-test.apps.elex.be;EM_ALERTS;WARNING;HARD;3;NRPE: Unable to read output
[1360310088] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360310088] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360310165] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;SOFT;1;CRITICAL - 'Manager:_Lot_Move_Transactions' Pending Normal	      08/02/2013 08:45
[1360310178] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1360310181] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360310219] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360310247] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360310269] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;HARD;3;Connection refused
[1360310275] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;SOFT;1;CRITICAL - 'Process_transaction_interface' Pending Normal	      08/02/2013 08:46
[1360310298] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1360310319] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;SOFT;2;CRITICAL - 'WIP_Move_Transaction_Manager' has not been scheduled for 5 mins
[1360310329] SERVICE ALERT: bipmelei-test.apps.elex.be;NRPE;UNKNOWN;SOFT;1;User name nrpe was not found
[1360310349] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1360310409] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360310530] SERVICE ALERT: bipmelei-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360310549] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;SOFT;2;CRITICAL - 'Manager:_Lot_Move_Transactions' has not been scheduled for 5 mins
[1360310566] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=58 ms
[1360310567] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360310607] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360310614] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360310629] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360310657] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;SOFT;2;CRITICAL - 'Process_transaction_interface' has not been scheduled for 5 mins
[1360310707] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;HARD;3;CRITICAL - 'WIP_Move_Transaction_Manager' has not been scheduled for 5 mins
[1360310715] SERVICE ALERT: bipmelei-test.apps.elex.be;NRPE;OK;SOFT;2;OK - 3 processes running with UID = 18018 (nrpe), args /usr/sbin/nrpe
[1360310727] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=130 ms
[1360310787] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360310855] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=47 ms
[1360310897] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;HARD;3;CRITICAL - 'Manager:_Lot_Move_Transactions' has not been scheduled for 5 mins
[1360310935] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360310945] SERVICE ALERT: bipmelei-test.apps.elex.be;EM_ALERTS;WARNING;HARD;3;NRPE: Unable to read output
[1360310964] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360310996] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;HARD;3;CRITICAL - 'Process_transaction_interface' has not been scheduled for 5 mins
[1360311037] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360311075] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=5 ms
[1360311205] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360311257] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360311286] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360311310] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=153 ms
[1360311345] SERVICE ALERT: bipmelei-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360311345] SERVICE FLAPPING ALERT: bipmelei-test.apps.elex.be;EM_ALERTS;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1360311396] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360311614] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360311714] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1360311736] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1360311776] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1360312077] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360312116] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360312385] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360312397] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1360312409] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360312435] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360312489] SERVICE ALERT: heidi.printers.bevaix.elex.be;PRINTER_STATUS;CRITICAL;SOFT;1;CRITICAL,Toner noir is at CRITICAL Level OK,Toner usag 1 is FULL OK,Toner cyan is FULL OK,Toner magenta is FULL OK,Toner jaune is FULL
[1360312544] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360312647] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=41 ms
[1360312695] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 1560 bytes in 0.015 seconds
[1360312705] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360312720] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1360312726] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360312758] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;OK;HARD;3;PROCS OK: 50 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360312767] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;OK;HARD;3;PROCS OK: 1 process with UID = 59998 (applmgr), args 'tnslsnr'
[1360312808] SERVICE ALERT: heidi.printers.bevaix.elex.be;PRINTER_STATUS;OK;SOFT;2;OK,Toner noir is FULL OK,Toner usag 1 is FULL OK,Toner cyan is FULL OK,Toner magenta is FULL OK,Toner jaune is FULL
[1360312864] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1360312867] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360312918] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;OK;HARD;3;PROCS OK: 11 processes with UID = 59998 (applmgr), args 'Apache'
[1360312966] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;OK;HARD;3;PROCS OK: 1 process with UID = 59998 (applmgr), args 'FNDCRM'
[1360312974] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=11 ms
[1360313047] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1360313066] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;OK;HARD;3;OK - 'WIP_Move_Transaction_Manager' Completed Normal      08/02/2013 09:37
[1360313188] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360313247] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;OK;HARD;3;OK - 'Manager:_Lot_Move_Transactions' Completed Normal      08/02/2013 09:37
[1360313341] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;OK;HARD;3;OK - 'Process_transaction_interface' Completed Normal      08/02/2013 09:37
[1360313508] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360313514] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=101 ms
[1360313599] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360313618] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360313668] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360313759] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360313866] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360313916] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360313948] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1360313965] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360314026] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1360314150] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360314309] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360314416] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1360314538] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360314585] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1360314644] SERVICE ALERT: dlog1-test.sensors.elex.be;CPU_STATS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1360314696] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360314707] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=90 ms
[1360314799] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360314816] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360314879] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360314946] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360314985] SERVICE ALERT: dlog1-test.sensors.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.00% idle=100.00%
[1360315027] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=103 ms
[1360315038] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1360315047] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360315094] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 78.61 ms
[1360315134] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360315137] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360315144] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360315269] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360315271] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1360315357] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=4 ms
[1360315374] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1360315386] SERVICE ALERT: bipmelei-test.apps.elex.be;ARCHIVELOG;WARNING;HARD;3;NRPE: Unable to read output
[1360315414] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 54%, RTA = 71.83 ms
[1360315434] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeListener;CRITICAL;SOFT;1;TestScapeListener: Stopped
[1360315468] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1360315470] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360315474] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1360315556] SERVICE ALERT: tscape-test.colo.elex.be;DISK_C_NT;WARNING;HARD;3;C:\ - total: 19.99 Gb - used: 17.59 Gb (88%) - free 2.39 Gb (12%)
[1360315604] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360315607] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1360315738] SERVICE ALERT: bipmelei-test.apps.elex.be;EM_ALERTS;WARNING;HARD;3;NRPE: Unable to read output
[1360315758] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeListener;CRITICAL;SOFT;2;TestScapeListener: Stopped
[1360315758] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=13 ms
[1360315758] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1360315761] SERVICE ALERT: bipmelei-test.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360315820] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360315849] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360315864] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360315928] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1360315935] SERVICE ALERT: openft-test.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360316087] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeListener;OK;SOFT;3;TestScapeListener: Started
[1360316135] SERVICE ALERT: bipmelei-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360316149] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360316155] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360316184] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1360316194] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360316264] SERVICE ALERT: openft-test.elex.be;NRPE;UNKNOWN;HARD;2;check_procs: User name was not found - nrpe
[1360316304] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360316317] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360316467] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeListener;CRITICAL;SOFT;1;TestScapeListener: Stopped
[1360316468] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360316506] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1360316535] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=104 ms
[1360316616] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360316636] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1360316700] SERVICE ALERT: pmelei-test.apps.elex.be;NRPE;UNKNOWN;SOFT;1;User name nrpe was not found
[1360316795] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeListener;CRITICAL;SOFT;2;TestScapeListener: Stopped
[1360316854] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360316865] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360316909] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (55.9037900874636)
[1360316964] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360317005] SERVICE ALERT: openft-test.elex.be;NRPE;OK;HARD;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360317025] SERVICE ALERT: pmelei-test.apps.elex.be;NRPE;OK;SOFT;2;OK - 3 processes running with UID = 18018 (nrpe), args /usr/sbin/nrpe
[1360317127] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeListener;CRITICAL;HARD;3;TestScapeListener: Stopped
[1360317188] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1360317246] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeLoader;CRITICAL;SOFT;1;TestScapeLoader: Stopped
[1360317314] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=169 ms
[1360317357] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=186 ms
[1360317517] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeListener;OK;HARD;3;TestScapeListener: Started
[1360317576] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeLoader;CRITICAL;SOFT;2;TestScapeLoader: Stopped
[1360317687] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=14 ms
[1360317688] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360317729] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1360317749] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=101 ms
[1360317806] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360317857] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360317964] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeLoader;CRITICAL;HARD;3;TestScapeLoader: Stopped
[1360318064] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360318077] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360318096] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1360318135] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360318178] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360318213] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 74.49 ms
[1360318285] SERVICE ALERT: batch-test.tess.elex.be;CFAGENT;WARNING;SOFT;1;WARNING: cfengine:batch-test: Can't stat /var/lib/cfengine2/masterfiles/files/cronjobs/templog-batch-test.tess.elex.be in copy
[1360318464] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1360318526] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 73%, RTA = 64.14 ms
[1360318615] SERVICE ALERT: batch-test.tess.elex.be;CFAGENT;WARNING;SOFT;2;WARNING: cfengine:batch-test: Can't stat /var/lib/cfengine2/masterfiles/files/cronjobs/templog-batch-test.tess.elex.be in copy
[1360318803] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360318944] SERVICE ALERT: batch-test.tess.elex.be;CFAGENT;WARNING;HARD;3;WARNING: cfengine:batch-test: Can't stat /var/lib/cfengine2/masterfiles/files/cronjobs/templog-batch-test.tess.elex.be in copy
[1360319123] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1360319224] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360319295] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360319517] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360319539] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=83 ms
[1360319613] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360319844] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360319865] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=60 ms
[1360319934] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1360320017] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=118 ms
[1360320266] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360320314] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1360320405] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=65 ms
[1360320487] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360320493] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360320603] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360320819] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360320823] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360320844] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 62.78 ms
[1360320933] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1360321018] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1360321037] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1360321037] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1360321037] SERVICE FLAPPING ALERT: ldap-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1360321037] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (29.7% change >= 20.0% threshold)
[1360321051] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1360321051] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1360321153] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=91 ms
[1360321192] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360321316] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=105 ms
[1360321377] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360321408] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 54%, RTA = 74.85 ms
[1360321483] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1360321494] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=81 ms
[1360321638] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360321653] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 69.46 ms
[1360321699] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1360321699] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1360321764] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=105 ms
[1360321764] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1360321893] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360322087] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1360322173] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360322295] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360322453] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=131 ms
[1360322493] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1360322625] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360322737] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360322783] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1360322858] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1360322883] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1360322944] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360322944] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1360323056] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1360323095] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360323166] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360323387] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360323415] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1360323493] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360323563] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360323564] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360323564] SERVICE ALERT: murray-test.sensors.elex.be;SSHBACK;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360323584] SERVICE ALERT: lechuck-test.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360323587] SERVICE ALERT: dlog1-test.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360323595] SERVICE ALERT: ftp-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 28.86 ms
[1360323646] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360323768] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360323785] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360323873] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=96 ms
[1360323875] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_4.3p2 Debian-9etch3 (protocol 2.0)
[1360323884] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;OK;SOFT;2;TCP OK - 0.011 second response time on port 6913
[1360323884] SERVICE ALERT: murray-test.sensors.elex.be;SSHBACK;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1360323887] SERVICE ALERT: lechuck-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.24 ms
[1360323893] SERVICE ALERT: dlog1-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 11.09 ms
[1360323893] SERVICE ALERT: ftp-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.38 ms
[1360323976] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360324088] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360324108] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1360324175] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360324193] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1360324353] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360324393] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360324433] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1360324503] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360324573] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360324647] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360324673] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360324706] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1360324725] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360324743] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 60.65 ms
[1360324832] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1360324832] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1360324896] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360324956] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1360325026] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360325090] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=96 ms
[1360325258] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360325291] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360325358] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (55.9037900874636)
[1360325457] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 63.72 ms
[1360325476] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360325476] HOST FLAPPING ALERT: temp1.sofia.elex.be;STARTED; Host appears to have started flapping (20.9% change > 20.0% threshold)
[1360325477] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeLoader;OK;HARD;3;TestScapeLoader: Started
[1360325484] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360325569] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360325576] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 65.92 ms
[1360325577] SERVICE ALERT: tscape-test.colo.elex.be;TestScapeAgent;OK;HARD;3;TestScapeAgent: Started
[1360325646] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360325686] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360325782] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 71.74 ms
[1360325867] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360325887] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360325938] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360325956] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360325966] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360326008] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360326097] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 70%, RTA = 79.01 ms
[1360326200] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360326262] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360326278] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360326368] SERVICE ALERT: dvorak-test.colo.elex.be;DISK;WARNING;SOFT;1;[ /vservers ] -> Size: 124G Used: 98G Avail: 27G Free: (20%)
[1360326396] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360326428] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360326449] SERVICE ALERT: bipmelei-test.apps.elex.be;NRPE;UNKNOWN;SOFT;1;User name nrpe was not found
[1360326620] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=90 ms
[1360326706] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360326737] SERVICE ALERT: dvorak-test.colo.elex.be;DISK;WARNING;SOFT;2;[ /vservers ] -> Size: 124G Used: 99G Avail: 26G Free: (20%)
[1360326766] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=119 ms
[1360326797] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360326818] SERVICE ALERT: bipmelei-test.apps.elex.be;NRPE;OK;SOFT;2;OK - 3 processes running with UID = 18018 (nrpe), args /usr/sbin/nrpe
[1360326848] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1360326987] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360327050] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360327079] SERVICE ALERT: dvorak-test.colo.elex.be;DISK;WARNING;HARD;3;[ /vservers ] -> Size: 124G Used: 99G Avail: 26G Free: (20%)
[1360327099] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360327109] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=131 ms
[1360327138] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360327179] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360327427] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360327500] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1360327519] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=118 ms
[1360327649] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360327697] SERVICE ALERT: openft-test.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360327749] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=115 ms
[1360327809] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=49 ms
[1360327850] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360327892] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360327991] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360328027] SERVICE ALERT: openft-test.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360328097] SERVICE ALERT: fili-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360328100] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1360328166] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1360328250] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360328267] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1360328474] SERVICE ALERT: fili-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1360328484] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=114 ms
[1360328561] HOST FLAPPING ALERT: temp1.sofia.elex.be;STOPPED; Host appears to have stopped flapping (0.0% change < 5.0% threshold)
[1360328644] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360328654] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360328685] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360328733] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360328800] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 75.47 ms
[1360328890] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1360329015] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360329130] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360329163] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=59 ms
[1360329302] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1360329462] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360329472] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.5150631681244)
[1360329564] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=12 ms
[1360329646] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=93 ms
[1360329690] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=53 ms
[1360329791] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360329803] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 88.64 ms
[1360329825] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360329872] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360329976] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=164 ms
[1360330003] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1360330050] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=5 ms
[1360330141] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360330141] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 66.39 ms
[1360330173] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1360330225] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360330262] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1360330330] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360330531] SERVICE ALERT: cmdb-test.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Cfengine is not running for more than 1day!
[1360330586] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360330620] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360330684] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=185 ms
[1360330733] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360330815] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=106 ms
[1360330871] SERVICE ALERT: cmdb-test.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Cfengine is not running for more than 1day!
[1360330922] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360331061] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360331074] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1360331193] SERVICE ALERT: cmdb-test.elex.be;CFAGENT;CRITICAL;HARD;3;CRITICAL: Cfengine is not running for more than 1day!
[1360331204] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=77 ms
[1360331325] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360331480] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=170 ms
[1360331492] SERVICE ALERT: oragrid2-test.apps.elex.be;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.32.241.189)
[1360331504] SERVICE ALERT: oragrid2-test.apps.elex.be;INODES;CRITICAL;SOFT;1;Connection refused or timed out
[1360331504] HOST ALERT: oragrid2-test.apps.elex.be;DOWN;SOFT;1;Server down!
[1360331540] SERVICE ALERT: oragrid2-test.apps.elex.be;NRPE;CRITICAL;HARD;1;Connection refused or timed out
[1360331610] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1360331630] HOST ALERT: oragrid2-test.apps.elex.be;UP;SOFT;2;Server alive
[1360331661] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360331830] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=125 ms
[1360331831] SERVICE ALERT: oragrid2-test.apps.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 4.53 ms
[1360331832] SERVICE ALERT: oragrid2-test.apps.elex.be;TNS;CRITICAL;SOFT;1;Connection refused
[1360331930] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=48 ms
[1360331946] SERVICE ALERT: oragrid2-test.apps.elex.be;NRPE;OK;HARD;1;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360332049] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=94 ms
[1360332063] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=113 ms
[1360332063] SERVICE ALERT: oragrid2-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360332168] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360332238] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1360332238] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 61.23 ms
[1360332239] SERVICE ALERT: oragrid2-test.apps.elex.be;TNS;CRITICAL;SOFT;2;Connection refused
[1360332245] SERVICE ALERT: oragrid2-test.apps.elex.be;INODES;OK;SOFT;2;OK - all disk
[1360332281] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360332323] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=13 ms
[1360332469] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1360332470] SERVICE ALERT: oragrid2-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360332488] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1360332488] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1360332590] SERVICE ALERT: oragrid2-test.apps.elex.be;TNS;CRITICAL;HARD;3;Connection refused
[1360332659] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360332702] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360332818] SERVICE ALERT: oragrid2-test.apps.elex.be;TNSLSNR;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360332968] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360332988] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1360333023] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360333039] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 65.92 ms
[1360333122] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360333288] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360333311] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1360333360] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 75%, RTA = 67.48 ms
[1360333448] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360333773] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (55.9037900874636)
[1360333793] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360334000] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360334062] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360334084] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=185 ms
[1360334121] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360334168] SERVICE ALERT: viiper-test.apps.elex.be;SYSLOG;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1360334263] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360334289] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360334341] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360334398] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360334411] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1360334418] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360334438] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360334458] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 82.31 ms
[1360334492] SERVICE ALERT: viiper-test.apps.elex.be;SYSLOG;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'syslogd'
[1360334511] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360334739] SERVICE ALERT: irmina.printers.erfurt.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360334739] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360334758] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360334759] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360334784] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360334838] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360334889] SERVICE ALERT: irmina.printers.erfurt.elex.be;PRINTER_STATUS;OK;HARD;3;(No output returned from plugin)
[1360334900] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1360334930] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360335059] SERVICE ALERT: irmina.printers.erfurt.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1360335081] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=50 ms
[1360335108] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360335120] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=88 ms
[1360335169] SERVICE ALERT: dss-test.colo.elex.be;DSS_POSTGRES;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 2737 (syncmgr), args 'postgres'
[1360335180] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1360335218] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360335229] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1360335248] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360335288] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1360335318] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360335458] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360335509] SERVICE ALERT: dss-test.colo.elex.be;DSS_POSTGRES;OK;SOFT;2;PROCS OK: 10 processes with UID = 2737 (syncmgr), args 'postgres'
[1360335520] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360335558] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360335581] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1360335601] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360335734] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360335776] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 67.95 ms
[1360335776] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1360335865] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360335865] HOST FLAPPING ALERT: temp1.sofia.elex.be;STARTED; Host appears to have started flapping (21.7% change > 20.0% threshold)
[1360335938] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1360335940] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=59 ms
[1360335964] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=94 ms
[1360336024] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1360336114] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 61.61 ms
[1360336120] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1360336125] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 67.73 ms
[1360336134] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=59 ms
[1360336184] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1360336275] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=102 ms
[1360336337] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=103 ms
[1360336359] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1360336377] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1360336387] SERVICE ALERT: confluence-xpeqt-test.colo.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360336474] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=18 ms
[1360336538] SERVICE ALERT: crowd-test.colo.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360336539] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 66.41 ms
[1360336584] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1360336607] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1360336695] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360336725] SERVICE ALERT: confluence-xpeqt-test.colo.elex.be;NRPE;OK;SOFT;2;PROCS OK: 4 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360336734] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1360336864] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 64.76 ms
[1360336865] SERVICE ALERT: crowd-test.colo.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360337028] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1360337137] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1360337159] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=100 ms
[1360337196] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 54%, RTA = 69.71 ms
[1360337247] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360337415] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360337508] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360337548] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1360337595] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360337758] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;OK;SOFT;2;HTTP OK HTTP/1.1 200 OK - 1803 bytes in 0.055 seconds
[1360337825] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1360337857] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1360337947] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360337990] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360338034] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360338050] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360338154] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360338249] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360338285] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360338336] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360338394] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360338410] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360338535] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=103 ms
[1360338605] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360338634] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360338658] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1360338844] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 64.49 ms
[1360339066] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=153 ms
[1360339076] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360339234] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360339337] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1360339467] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360339514] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1360339535] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1360339606] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;CRITICAL;SOFT;1;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 41% CRITICAL,Waste Toner Box is at CRITICAL Level
[1360339624] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360339664] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360339678] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360339759] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=111 ms
[1360339874] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360339898] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1360339924] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;CRITICAL;SOFT;2;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 41% CRITICAL,Waste Toner Box is at CRITICAL Level
[1360339984] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360339997] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360340075] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=7 ms
[1360340155] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1360340189] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360340244] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;CRITICAL;HARD;3;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 41% CRITICAL,Waste Toner Box is at CRITICAL Level
[1360340277] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360340536] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360340597] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360340738] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1360340855] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360340986] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360341060] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1360341138] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360341308] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360341329] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360341357] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1360341380] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1360341459] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360341609] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360341647] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360341670] HOST FLAPPING ALERT: temp1.sofia.elex.be;STOPPED; Host appears to have stopped flapping (0.0% change < 5.0% threshold)
[1360341674] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=84 ms
[1360341926] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=101 ms
[1360341966] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=50 ms
[1360341997] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1360342048] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360342065] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360342070] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360342137] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1360342235] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1360342248] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=103 ms
[1360342347] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1360342368] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1360342388] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360342447] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360342575] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1360342627] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=81 ms
[1360342707] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1360342769] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360342964] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;OK;HARD;3;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 41% OK,Waste Toner Box is FULL
[1360342976] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1360343316] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1360343456] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=105 ms
[1360343559] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=47 ms
[1360343579] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=103 ms
[1360343595] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360343647] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1360343836] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1360343879] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=4 ms
[1360343919] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360343959] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1360343984] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1360344099] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360344295] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360344314] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360344321] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1360344340] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=112 ms
[1360344424] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360344615] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360344634] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=87 ms
[1360344659] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1360344679] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360344719] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=76 ms
[1360344955] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1360344977] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1360345000] HOST ALERT: hma-lp.sofia.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1360345000] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360345179] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360345335] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1360345357] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360345379] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360345504] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360345525] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360345678] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360345700] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.9037900874636)
[1360345716] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360345760] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360345824] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1360345861] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=119 ms
[1360346034] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=84 ms
[1360346058] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1360346080] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1360346242] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=89 ms
[1360346346] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=12 ms
[1360346376] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360346677] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1360346753] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1360346790] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1360347057] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1360347108] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360347160] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=42 ms
[1360347340] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360347379] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360347481] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=21 ms
[1360347518] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=99 ms
[1360347660] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360347661] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360347840] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360347867] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360347988] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360348187] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360348508] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1360348517] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360348549] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=91 ms
[1360348801] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360348837] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360348870] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1360348980] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360349121] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360349187] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1360349299] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360349558] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360349597] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360349679] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1360349877] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360349917] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1360350029] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1360350263] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (55.9037900874636)
[1360350312] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1360350334] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1360350416] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360350475] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360350496] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360350655] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360350794] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1360350815] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360351116] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1360351116] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360351122] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1360351138] SERVICE ALERT: sda32-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360351143] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360351143] SERVICE ALERT: murray-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360351144] SERVICE ALERT: sda-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360351159] SERVICE ALERT: sda32-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 21.86 ms
[1360351218] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360351413] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360351435] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360351442] SERVICE ALERT: sda32-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.4 (protocol 2.0)
[1360351452] SERVICE ALERT: sda32-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 10.63 ms
[1360351453] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;OK;SOFT;2;TCP OK - 0.009 second response time on port 6445
[1360351456] SERVICE ALERT: murray-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1360351458] SERVICE ALERT: sda-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze1 (protocol 2.0)
[1360351538] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360351538] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (55.2721088435374)
[1360351538] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1360351539] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360351736] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360351736] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;UNKNOWN;SOFT;1;UNKOWN: error reading value
[1360351743] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360351756] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360351766] SERVICE ALERT: postgres-stage.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360351832] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=120 ms
[1360351857] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1360351873] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1360351935] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1360352054] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=11 ms
[1360352055] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;OK;SOFT;2;OK: 15.5
[1360352056] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;OK;SOFT;2;TCP OK - 0.012 second response time on port 6933
[1360352062] SERVICE ALERT: postgres-stage.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 10.46 ms
[1360352155] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=115 ms
[1360352178] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1360352255] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1360352256] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360352475] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1360352575] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360352575] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=82 ms
[1360352666] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360352855] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1360352937] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360352955] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1360352983] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.5150631681244)
[1360353062] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360353122] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=51 ms
[1360353234] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360353385] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360353442] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=17 ms
[1360353555] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1360353568] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360353774] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360353943] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1360353948] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360354095] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360354272] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1360354322] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1360354533] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360354616] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360354642] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360354652] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1360354683] SERVICE ALERT: elaine-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360354685] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360354855] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1360354855] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360354902] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360354933] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (53.2312925170068)
[1360354988] SERVICE ALERT: elaine-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1360354992] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360355002] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1360355175] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360355182] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=103 ms
[1360355222] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1360355313] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360355494] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=90 ms
[1360355542] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1360355562] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1360355584] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=1.10% system=1.00% iowait=97.90% idle=0.00%
[1360355614] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1360355626] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1360355784] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360355903] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=48.85% system=2.20% iowait=36.66% idle=12.29%
[1360355922] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360356093] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360356102] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360356242] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360356254] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360356322] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1360356412] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1360356482] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360356663] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=9.81% system=4.20% iowait=82.38% idle=3.60%
[1360356694] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360356804] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360356834] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360356984] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=9.60% system=0.70% iowait=40.20% idle=49.50%
[1360357022] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=45 ms
[1360357023] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360357152] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.5150631681244)
[1360357358] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=261 ms
[1360357389] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1360357429] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=393 ms
[1360357606] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1360357618] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=252 ms
[1360357686] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1360357756] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360357829] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1360357922] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=8 ms
[1360357936] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360358063] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=224 ms
[1360358116] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360358136] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360358148] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360358386] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=13 ms
[1360358436] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360358456] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1360358497] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360358589] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360358696] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360358776] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1360358820] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1360358896] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=41 ms
[1360358909] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=104 ms
[1360358911] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360358926] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=55 ms
[1360359016] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=162 ms
[1360359201] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360359217] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1360359227] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360359236] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1360359246] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1360359250] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360359276] SERVICE ALERT: mainrouter-test.kiev.elex.be;AT;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1360359309] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=491 ms
[1360359336] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=319 ms
[1360359496] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=58 ms
[1360359519] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1360359578] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360359587] SERVICE ALERT: mainrouter-test.kiev.elex.be;AT;OK;SOFT;2;PROCS OK: 1 process with UID = 1 (daemon), args '/usr/sbin/atd'
[1360359630] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1360359716] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1360359816] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1360359840] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1360359900] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (55.2721088435374)
[1360359946] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=45 ms
[1360359988] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=100 ms
[1360360006] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360360096] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360360186] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;1;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360360266] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=12 ms
[1360360308] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360360317] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1360360499] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;2;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360360600] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1360360630] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1360360638] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1360360680] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360360844] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360360862] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360360914] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=149 ms
[1360360975] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1360361051] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360361065] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1360361071] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360361152] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360361221] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360361254] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1360361385] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360361391] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360361475] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360361575] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=85 ms
[1360361771] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360362092] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360362194] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360362356] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360362413] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1360362514] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360362524] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=107 ms
[1360362676] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360362715] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360362715] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1360362845] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360362894] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360363036] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1360363057] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1360363096] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360363163] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360363211] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360363323] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360363377] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1360363416] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360363531] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1360363554] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360363645] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360363696] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;HARD;3;CRITICAL: Unable to establish connection with cfengine master server
[1360363696] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1360363762] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360363795] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360363873] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1360363923] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=91 ms
[1360363966] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1360364014] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360364077] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360364116] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360364193] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1360364244] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360364334] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms