[1364252448] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364252449] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 13.78, 6.02, 4.28
[1364252482] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364252482] SERVICE ALERT: janus-test.colo.elex.be;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364252503] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=99 ms
[1364252553] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364252558] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=0 ms
[1364252592] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 18.68, 11.55, 7.13
[1364252599] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364252621] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364252644] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1364252658] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.16 ms
[1364252690] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364252704] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1364252758] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=0.95 %
[1364252905] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364252926] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1364252928] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 18.84, 12.84, 7.70
[1364252954] SERVICE ALERT: janus-test.colo.elex.be;HTTP;OK;SOFT;2;HTTP OK HTTP/1.1 200 OK - 8265 bytes in 0.334 seconds
[1364253025] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364253033] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364253046] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1364253068] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 4.34, 5.45, 5.78
[1364253086] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1364253105] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364253132] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.18 ms
[1364253155] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364253172] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364253176] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364253186] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1364253216] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;3;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364253322] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2198.00 KBps
[1364253362] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1364253365] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 6.78, 11.71, 9.72
[1364253438] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364253535] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1364253536] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1364253565] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 86%, RTA = 2.32 ms
[1364253602] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1364253603] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364253622] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=14.09% system=3.10% iowait=81.82% idle=1.00%
[1364253655] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364253674] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1364253837] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1364253844] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364253854] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 10.45, 16.46, 13.51
[1364253866] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364253925] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;UNKNOWN;HARD;3;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364253934] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364253962] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=1 ms
[1364253962] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi IO read;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1364253965] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1364254003] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254052] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;SOFT;2;CPU IOWAIT WARNING : user=17.07% system=4.49% iowait=74.85% idle=3.59%
[1364254077] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254092] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364254142] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254242] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2343.00 KBps
[1364254273] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1364254302] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364254335] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1364254343] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 7.76, 10.27, 11.56
[1364254364] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364254435] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254447] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=90 ms
[1364254482] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;3;CPU STATISTICS OK : user=17.80% system=11.70% iowait=58.20% idle=12.30%
[1364254508] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254565] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1364254572] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364254573] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254814] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.88 %
[1364254814] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254859] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254875] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=102 ms
[1364254898] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=0 ms
[1364254954] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;3;P1.PL OK - overall status=green
[1364254984] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364254992] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364255002] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1364255025] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364255036] SERVICE ALERT: gambit-test.colo.elex.be;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1364255133] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1364255212] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.56 %
[1364255255] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364255273] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364255303] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1364255310] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;OK;HARD;3;OK - load average: 5.23, 6.23, 8.54
[1364255310] SERVICE FLAPPING ALERT: atlan.erfurt.elex.be;LOAD;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1364255352] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364255365] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1364255424] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364255462] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1364255477] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1364255523] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 2.21 ms
[1364255552] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364255562] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364255652] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=3491.00 KBps
[1364255783] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364255785] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1364255817] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364255823] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364255858] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1364255885] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1364255906] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1364255908] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364255982] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1364255985] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256052] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1364256112] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1364256225] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256249] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256253] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364256268] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256275] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1364256287] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 10.15, 10.60, 9.66
[1364256332] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1364256345] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256368] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364256527] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364256540] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1364256599] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=76 ms
[1364256657] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.20 %
[1364256715] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364256737] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364256737] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256746] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256766] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1364256786] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 5.83, 8.87, 9.12
[1364256826] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256848] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364256969] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1364256985] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1364257036] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1364257107] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364257191] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364257197] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=1894.00 KBps
[1364257208] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364257227] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364257268] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364257292] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=0 ms
[1364257330] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1364257375] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364257400] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1364257476] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364257515] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364257535] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1364257561] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364257605] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=0.81 %
[1364257655] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 63.66 ms
[1364257675] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.88 %
[1364257736] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364257756] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=116 ms
[1364257769] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1364257815] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1364257839] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364257907] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1364258027] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1364258146] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258169] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364258170] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258189] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1364258206] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258227] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258247] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258272] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258307] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364258345] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1364258455] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1364258566] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2702.00 KBps
[1364258586] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258597] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364258619] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1364258627] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364258646] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364258668] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=2 ms
[1364258676] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364258735] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1364258748] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364258805] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364258885] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1364258895] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364258965] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.16 %
[1364259055] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364259077] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364259078] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364259096] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364259165] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1364259176] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364259179] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1364259185] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364259196] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364259315] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1364259316] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1364259355] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364259355] HOST FLAPPING ALERT: temp1.sofia.elex.be;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1364259370] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=88 ms
[1364259526] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364259566] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364259598] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1364259605] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364259615] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 54.56 ms
[1364259625] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364259648] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1364259668] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1364259690] SERVICE ALERT: lenor-test.sofia.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 38.04, 15.11, 7.48
[1364259747] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1364259748] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1364259806] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1364259849] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (59.5966958211856)
[1364259915] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=3250.00 KBps
[1364260000] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364260037] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364260045] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=1 ms
[1364260058] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364260080] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=6.20% system=1.90% iowait=85.70% idle=6.20%
[1364260086] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364260138] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1364260157] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364260176] SERVICE ALERT: lenor-test.sofia.elex.be;LOAD;OK;SOFT;2;OK - load average: 14.35, 13.07, 9.77
[1364260291] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364260350] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364260436] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364260521] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 10.88, 10.58, 8.43
[1364260525] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364260537] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=1.70% system=6.40% iowait=63.20% idle=28.70%
[1364260577] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=148 ms
[1364260579] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1364260608] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364260647] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1364260718] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1364260778] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364260898] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1364260949] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 12.08, 12.49, 10.21
[1364261010] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1364261038] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364261045] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL CRITICAL - Server version unavailable at 'https://10.32.241.183:443/sdk/vimService.wsdl' at /usr/share/perl5/VMware/VICommon.pm line 544.
[1364261065] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1364261075] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1364261087] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;HARD;3;P1.PL OK - mem usage=2.88 %
[1364261152] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1364261268] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364261328] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364261376] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 18.74, 18.46, 13.70
[1364261425] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=1 ms
[1364261435] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=50 ms
[1364261468] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364261486] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364261495] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364261505] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1364261515] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=3.99% system=1.00% iowait=79.04% idle=15.97%
[1364261525] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 78%, RTA = 1.58 ms
[1364261548] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1364261587] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364261639] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1364261691] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1364261815] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364261845] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.19 %
[1364261865] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=1601.00 KBps
[1364261866] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 7.05, 11.23, 12.57
[1364261936] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;2;CPU IOWAIT CRITICAL : user=0.40% system=0.90% iowait=98.70% idle=0.00%
[1364261977] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364261989] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364262018] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.19 ms
[1364262038] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=145 ms
[1364262119] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1364262125] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1364262245] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364262275] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;UNKNOWN;SOFT;1;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 Connect failed: connect: Connection timed out: Connection timed out
[1364262365] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;CRITICAL;HARD;3;CPU IOWAIT CRITICAL : user=2.90% system=7.40% iowait=82.90% idle=6.80%
[1364262387] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364262408] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1364262415] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364262446] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364262449] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.18 ms
[1364262465] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364262478] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=12 ms
[1364262508] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364262553] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1364262661] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364262716] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 10.72, 6.41, 5.08
[1364262766] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.01 %
[1364262837] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.39, 5.46, 8.10
[1364262838] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364262845] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364262862] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1364262879] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364262880] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364262883] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 86%, RTA = 1.22 ms
[1364262898] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364262916] HOST ALERT: gambit-test.colo.elex.be;UP;SOFT;2;Server alive
[1364262955] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=102 ms
[1364262965] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1364262986] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;HARD;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364263025] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364263095] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1364263147] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 4.63, 5.44, 5.06
[1364263209] HOST FLAPPING ALERT: temp1.sofia.elex.be;STOPPED; Host appears to have stopped flapping (0.0% change < 5.0% threshold)
[1364263209] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263215] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1364263310] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263356] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;1;P1.PL OK - net usage=2583.00 KBps
[1364263356] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi Network usage;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1364263369] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263376] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.15 ms
[1364263385] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=166 ms
[1364263388] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1364263405] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=105 ms
[1364263431] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263456] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364263528] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263531] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364263616] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364263705] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364263766] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (47.8377065111759)
[1364263766] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1364263797] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263846] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.20 ms
[1364263859] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263860] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=69 ms
[1364263860] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364263867] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364263880] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364263880] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1364263885] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;OK;HARD;3;CPU STATISTICS OK : user=4.10% system=2.30% iowait=54.80% idle=38.80%
[1364263928] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364263958] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364264010] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1364264045] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;Server down!
[1364264128] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1364264149] SERVICE ALERT: prb423.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1364264176] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=0.85 %
[1364264176] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_CPU;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1364264185] HOST ALERT: gambit-test.colo.elex.be;UP;SOFT;2;Server alive
[1364264265] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364264297] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364264307] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1364264355] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364264356] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364264357] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364264357] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1364264375] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=2.20% system=4.30% iowait=74.20% idle=19.30%
[1364264397] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364264506] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364264529] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364264555] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1364264715] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - Server version unavailable at 'https://10.32.241.183:443/sdk/vimService.wsdl' at /usr/share/perl5/VMware/VICommon.pm line 544.
[1364264715] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=111 ms
[1364264719] SERVICE FLAPPING ALERT: atlan.erfurt.elex.be;LOAD;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1364264738] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1364264788] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364264789] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364264789] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364264795] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364264806] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=4.50% system=2.40% iowait=65.63% idle=27.47%
[1364264845] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1364264895] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364264967] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364264985] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1364265098] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=0 ms
[1364265098] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1364265138] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1364265165] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1364265209] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364265269] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364265328] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=151 ms
[1364265346] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1364265346] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.88 %
[1364265346] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364265447] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1364265475] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364265639] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364265647] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364265758] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1364265769] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364265770] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364265770] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1364265807] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364265875] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364265888] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364265899] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1364265919] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364265960] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364266025] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1364266178] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364266187] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364266199] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364266199] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364266200] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364266237] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364266248] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1364266250] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364266265] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1364266298] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364266325] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364266360] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364266387] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1364266572] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364266572] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1364266597] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1364266668] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364266675] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1364266678] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364266686] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364266695] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1364266758] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364266790] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1364266825] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=101 ms
[1364266885] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=114 ms
[1364267040] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1364267041] SERVICE ALERT: pintaildb-test.colo.elex.be;DISK;WARNING;HARD;3;[ / ] -> Size: 14G Used: 12G Avail: 1.5G Free: (11%)
[1364267105] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364267115] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1364267126] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364267187] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364267198] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=1968.00 KBps
[1364267295] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=3.80% system=3.50% iowait=77.62% idle=15.08%
[1364267296] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;HARD;3;P1.PL OK - mem usage=2.88 %
[1364267296] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi Memory usage;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1364267355] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1364267357] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1364267376] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364267466] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364267605] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364267627] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364267645] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1364267647] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364267707] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1364267735] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.20% system=1.90% iowait=48.00% idle=49.90%
[1364267775] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;HARD;3;P1.PL OK - cpu usage=0.96 %
[1364267788] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364267789] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=46 ms
[1364267811] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364267896] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364268003] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1364268080] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364268107] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364268175] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364268185] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364268261] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364268275] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364268279] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 77%, RTA = 3.07 ms
[1364268321] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364268355] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364268380] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1364268621] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364268625] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364268658] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.88 %
[1364268720] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364268745] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364268767] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364268779] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364268785] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.19 ms
[1364268788] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1364268797] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364268815] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1364268817] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1364269051] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364269056] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364269111] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364269151] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364269175] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1364269195] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364269197] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364269216] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 2.76 ms
[1364269229] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364269248] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=82 ms
[1364269436] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1364269539] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364269605] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=2472.00 KBps
[1364269605] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1364269628] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364269637] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 0.16 ms
[1364269647] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364269655] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;HARD;3;P1.PL OK - cpu usage=1.37 %
[1364269655] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_CPU;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1364269695] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364269699] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364269735] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1364269967] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270027] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364270060] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270077] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270078] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364270080] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=2 ms
[1364270095] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1364270121] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270146] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270275] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1364270345] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1364270455] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364270495] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;HARD;3;P1.PL OK - mem usage=2.88 %
[1364270509] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270516] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2728.00 KBps
[1364270525] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=76 ms
[1364270552] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270567] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364270610] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 79%, RTA = 0.17 ms
[1364270625] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270665] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270705] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1364270892] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364270965] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364270985] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1364271028] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1364271050] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364271096] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364271097] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364271131] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 9.55, 6.10, 5.17
[1364271137] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364271138] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364271227] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1364271455] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1364271479] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1364271501] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364271506] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364271519] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=9.42 %
[1364271519] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364271526] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364271527] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364271562] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 5.60, 5.08, 4.99
[1364271626] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 0.19 ms
[1364271628] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1364271655] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1364271879] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=0 ms
[1364271898] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1364271945] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364271955] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL CRITICAL - Server version unavailable at 'https://10.32.241.183:443/sdk/vimService.wsdl' at /usr/share/perl5/VMware/VICommon.pm line 544.
[1364271961] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1364271966] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272005] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;HARD;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364272005] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1364272055] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272056] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.17 ms
[1364272057] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1364272067] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272068] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272088] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1364272216] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=99 ms
[1364272345] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364272379] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1364272397] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1364272398] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=2241.00 KBps
[1364272435] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272452] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1364272452] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364272457] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364272465] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364272496] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 86%, RTA = 0.25 ms
[1364272496] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=77 ms
[1364272518] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272538] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1364272655] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364272678] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364272695] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (59.5966958211856)
[1364272695] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1364272815] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1364272878] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272880] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364272895] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364272958] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364272977] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1364272986] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1364272988] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=91 ms
[1364273066] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.36 %
[1364273285] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=11 ms
[1364273305] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1364273327] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1364273377] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364273389] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364273417] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.5150631681244)
[1364273448] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364273478] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1364273491] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364273567] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1364273606] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364273685] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1364273732] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1364273810] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364273879] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364273941] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364273985] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364273995] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=108 ms
[1364274038] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364274040] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364274118] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1364274235] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;UNKNOWN;SOFT;2;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 Connect failed: connect: Connection timed out: Connection timed out
[1364274285] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=2417.00 KBps
[1364274317] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364274409] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364274466] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364274505] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1364274507] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364274546] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=105 ms
[1364274557] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.10 %
[1364274716] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1364274825] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364274835] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364274885] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=12 ms
[1364274911] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364274965] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;3;P1.PL OK - overall status=green
[1364275009] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1364275037] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364275077] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1364275085] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364275228] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1364275237] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364275245] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi IO read;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1364275247] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1364275300] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1364275335] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2076.00 KBps
[1364275451] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=167 ms
[1364275456] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;HARD;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364275476] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364275491] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364275525] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1364275576] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364275608] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364275695] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (53.2312925170068)
[1364275712] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364275728] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1364275778] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1364275917] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364275928] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364275936] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1364275965] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364275999] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364276005] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1364276086] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364276087] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1364276165] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=101 ms
[1364276167] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364276172] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364276215] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1364276329] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1364276357] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2326.00 KBps
[1364276357] HOST ALERT: hma-lp.sofia.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 44.36 ms
[1364276415] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364276415] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=0 ms
[1364276455] SERVICE ALERT: hma-lp.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 45.84 ms
[1364276475] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1364276516] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364276596] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364276625] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1364276649] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364276728] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 9.27, 5.69, 4.94
[1364276825] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364276846] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364276908] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1364277028] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1364277049] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364277079] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364277085] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364277108] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364277126] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364277159] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 6.39, 5.18, 4.89
[1364277186] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;UNKNOWN;SOFT;1;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364277248] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=83 ms
[1364277326] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364277332] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1364277332] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1364277385] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=19.34% system=3.79% iowait=71.29% idle=5.58%
[1364277436] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364277475] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364277495] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364277515] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364277539] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1364277568] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364277619] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364277678] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1364277686] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364277698] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364277816] SERVICE ALERT: pintaildb-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.80% system=0.50% iowait=50.65% idle=48.05%
[1364277820] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1364277825] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364277918] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=41 ms
[1364277945] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1364277946] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364277997] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364278089] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1364278089] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364278124] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.88 %
[1364278125] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=86 ms
[1364278186] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364278308] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1364278399] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1364278434] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364278487] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364278487] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1364278537] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364278538] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364278546] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364278555] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;HARD;3;P1.PL OK - cpu usage=2.08 %
[1364278565] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=2 ms
[1364278624] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1364278817] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2063.00 KBps
[1364278848] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364278864] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364278899] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364278916] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1364278977] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279019] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1364279046] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1364279088] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279146] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279156] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1364279197] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1364279234] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364279277] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364279295] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1364279344] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1364279357] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279374] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=2 ms
[1364279404] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279446] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364279518] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279577] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279581] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1364279618] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.45 %
[1364279624] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1364279648] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;UNKNOWN;SOFT;2;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364279707] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364279747] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2657.00 KBps
[1364279774] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364279804] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364279834] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1364279894] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279955] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364279955] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.88 %
[1364280014] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=104 ms
[1364280028] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;3;P1.PL OK - overall status=green
[1364280054] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=135 ms
[1364280159] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364280186] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364280198] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364280257] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1364280274] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364280329] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364280385] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364280426] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364280449] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;SOFT;1;Connection refused
[1364280496] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1364280505] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1364280507] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1364280546] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=62 ms
[1364280555] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364280588] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364280611] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364280655] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=4189.00 KBps
[1364280669] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1364280690] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1364280715] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1364280735] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 81%, RTA = 55.51 ms
[1364280798] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364280844] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1364280856] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1364280864] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364280877] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;SOFT;2;Connection refused
[1364280886] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1364280928] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1364280928] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1364280968] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.30 %
[1364280987] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364280995] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1364281009] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;SOFT;1;CRITICAL - 'Manager:_Lot_Move_Transactions' Pending Normal	      26/03/2013 07:45
[1364281034] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=155 ms
[1364281060] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364281100] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1364281105] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364281120] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;SOFT;1;CRITICAL - 'WIP_Move_Transaction_Manager' has not been scheduled for 5 mins
[1364281197] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1364281205] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 63.85 ms
[1364281214] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 66.53 ms
[1364281274] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;SOFT;1;CRITICAL - 'Process_transaction_interface' Pending Normal	      26/03/2013 07:46
[1364281308] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;HARD;3;Connection refused
[1364281330] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1364281354] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364281370] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1364281371] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;OK;SOFT;3;PROCS OK: 11 processes with UID = 59998 (applmgr), args 'Apache'
[1364281423] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1364281450] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;SOFT;2;CRITICAL - 'Manager:_Lot_Move_Transactions' has not been scheduled for 5 mins
[1364281492] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364281531] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1364281536] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;OK;SOFT;3;PROCS OK: 1 process with UID = 59998 (applmgr), args 'tnslsnr'
[1364281539] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364281562] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;SOFT;2;CRITICAL - 'WIP_Move_Transaction_Manager' has not been scheduled for 5 mins
[1364281627] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364281677] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1364281749] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364281767] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;OK;SOFT;2;OK - 'Process_transaction_interface' Completed Normal      26/03/2013 08:03
[1364281799] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364281817] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;OK;SOFT;3;PROCS OK: 1 process with UID = 59998 (applmgr), args 'FNDCRM'
[1364281840] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1364281859] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 1560 bytes in 0.001 seconds
[1364281901] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1364281907] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;OK;HARD;3;PROCS OK: 50 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1364281921] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364281927] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;OK;SOFT;3;OK - 'Manager:_Lot_Move_Transactions' Completed Normal      26/03/2013 08:03
[1364282038] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;OK;SOFT;3;OK - 'WIP_Move_Transaction_Manager' Completed Normal      26/03/2013 08:04
[1364282116] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364282191] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364282228] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364282268] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1364282390] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364282396] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364282409] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364282526] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364282552] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1364282570] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=2616.00 KBps
[1364282647] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 70%, RTA = 1.28 ms
[1364282664] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364282758] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364282812] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1364282819] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1364282831] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364282839] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364282948] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1364283036] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364283038] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364283040] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1364283059] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364283076] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1364283108] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364283146] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 86%, RTA = 0.17 ms
[1364283180] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364283190] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364283198] HOST ALERT: gambit-test.colo.elex.be;UP;SOFT;2;Server alive
[1364283218] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364283249] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 55.82 ms
[1364283268] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364283309] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364283378] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi Memory usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1364283470] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364283472] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364283477] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364283559] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;1;P1.PL OK - net usage=1749.00 KBps
[1364283606] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;1;Server down!
[1364283616] HOST ALERT: sda-test.sofia.elex.be;UP;SOFT;2;Server alive
[1364283618] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1364283626] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 70%, RTA = 1.00 ms
[1364283652] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364283666] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 73%, RTA = 54.83 ms
[1364283676] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364283736] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1364283767] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364283787] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364283841] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.05 %
[1364283908] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364283927] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364283958] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=126 ms
[1364284103] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1364284127] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.21 ms
[1364284201] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364284219] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364284220] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1364284251] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364284386] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1364284543] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.22 ms
[1364284596] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1364284648] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1364284649] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 81%, RTA = 61.13 ms
[1364284656] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364284677] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;UNKNOWN;SOFT;2;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 Connect failed: connect: Connection timed out: Connection timed out
[1364284709] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1364284873] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364284876] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364284930] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364285037] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 81%, RTA = 0.18 ms
[1364285057] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.88 %
[1364285069] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364285073] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 75%, RTA = 54.30 ms
[1364285073] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=102 ms
[1364285089] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1364285139] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1364285167] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364285192] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364285214] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1364285402] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1364285464] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1364285503] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364285584] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.21 ms
[1364285600] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364285618] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1364285632] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1364285642] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364285678] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=1 ms
[1364285689] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1364285721] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364285721] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364285749] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364285762] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364285871] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364285883] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=100 ms
[1364285979] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364286028] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.52 %
[1364286058] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 0.17 ms
[1364286058] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;PING;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1364286109] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1364286128] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364286140] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.88 %
[1364286140] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364286152] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1364286222] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364286230] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 56.84 ms
[1364286231] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364286300] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364286501] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;WARNING;SOFT;1;WARNING - InFlightCount=19
[1364286578] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286579] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286579] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286581] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286582] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286582] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286583] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286583] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286621] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286624] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364286625] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286625] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286626] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286629] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286629] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286638] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286639] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286641] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286642] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1364286658] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1364286660] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286662] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286664] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286665] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286690] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 54%, RTA = 54.78 ms
[1364286711] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1364286731] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286748] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=2326.00 KBps
[1364286774] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286780] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286788] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286789] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286789] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286791] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286791] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286792] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286799] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286808] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286821] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364286832] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364286906] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364286918] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364286961] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - null
[1364286971] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1364286978] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364286978] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286979] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286980] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286982] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286999] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286999] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364286999] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287000] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287001] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287002] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287002] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287003] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287004] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287004] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287005] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287009] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287010] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287018] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287021] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287022] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287022] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287022] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287023] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287023] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287024] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287025] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287025] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287028] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287028] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287029] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287029] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287042] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287042] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287043] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287045] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287046] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287047] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287048] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - null
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - null
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - null
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - null
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - null
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - null
[1364287049] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - null
[1364287050] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - null
[1364287050] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - null
[1364287050] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - null
[1364287050] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - null
[1364287050] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - null
[1364287050] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - null
[1364287050] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - null
[1364287068] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364287118] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364287119] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287119] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287120] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287121] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287121] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287122] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364287128] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;3;TCP OK - 0.001 second response time on port 61616
[1364287129] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287129] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287132] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1364287138] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287139] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287139] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287141] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1364287158] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364287158] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287159] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287159] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287160] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287168] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 0.002 second response time on port 61616
[1364287181] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287222] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287260] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;SOFT;1;CRITICAL -  amkor.philippines.confirm - nothing found in the logfile. amkor.taiwan.confirm - nothing found in the logfile. atlantic.confirm - nothing found in the logfile. optopac.confirm - nothing found in the logfile. mpd.confirm - nothing found in the logfile. senodia.confirm - nothing found in the logfile. unisem.confirm - nothing found in the logfile. pactech.confirm - nothing found in the logfile. pactech.processed.from - nothing found in the logfile. ispl.confirm - nothing found in the logfile. xfab.from - nothing found in the logfile. heimann.confirm - nothing found in the logfile.
[1364287261] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1364287270] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287272] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287281] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287281] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287282] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287282] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287282] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287288] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1364287289] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287301] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287303] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287308] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364287309] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287325] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364287329] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364287338] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287404] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287410] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1364287440] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1364287479] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287479] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287479] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287480] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287480] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287488] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;3;TCP OK - 2.998 second response time on port 61616
[1364287503] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287503] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287503] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287503] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287503] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287503] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287504] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287504] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287504] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287505] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287505] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287508] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287508] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287524] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287524] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287525] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287528] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287528] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1364287528] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287529] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287530] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287541] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364287541] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287541] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287542] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287542] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287543] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287548] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287548] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287549] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287549] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287549] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287550] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287551] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287559] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364287561] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287562] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287569] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;3;TCP OK - 0.001 second response time on port 61616
[1364287572] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287578] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287579] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287579] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287580] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287581] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287582] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287582] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287583] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287589] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364287599] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364287603] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287608] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364287608] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364287609] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364287619] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364287620] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364287624] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287625] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287626] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287627] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364287631] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287631] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 0.000 second response time on port 61616
[1364287643] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1364287663] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287664] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287665] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287666] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364287682] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287719] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364287721] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia.confirm
[1364287722] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1364287729] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1364287730] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1364287739] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1364287740] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1364287741] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1364287749] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287751] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287760] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287772] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1364287779] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;SOFT;2;CRITICAL -  amkor.philippines.confirm - nothing found in the logfile. amkor.taiwan.confirm - nothing found in the logfile. amkor.korea.confirm - nothing found in the logfile. atlantic.confirm - nothing found in the logfile. optopac.confirm - nothing found in the logfile. mpd.confirm - nothing found in the logfile. senodia.confirm - nothing found in the logfile. pactech.confirm - nothing found in the logfile. pactech.processed.from - nothing found in the logfile. ispl.confirm - nothing found in the logfile. umc-import.from - nothing found in the logfile. leister.from - nothing found in the logfile. vishay.confirm - nothing found in the logfile. xfab.from - nothing found in the logfile.
[1364287799] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1364287834] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364287854] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364287856] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287872] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287872] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364287874] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364287920] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287920] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287921] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287921] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287921] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287941] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287941] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287941] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287941] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287941] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287941] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287949] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287949] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287949] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287950] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287951] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287952] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287953] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287969] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287969] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287970] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287971] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364287971] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287972] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287972] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287972] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287972] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287972] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287973] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287973] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287973] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287973] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364287973] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364287981] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364287982] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364287983] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364287983] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364287984] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364287992] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1364287999] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1364287999] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1364288000] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1364288001] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1364288002] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1364288005] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364288007] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1364288021] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364288043] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364288062] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364288065] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection refused
[1364288112] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=2.16 %
[1364288172] SERVICE ALERT: ewaf-test.colo.elex.be;SERVICEMIX;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 999 (smx-fuse), args '/bin/java'
[1364288231] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;HARD;3;CRITICAL -  amkor.philippines.confirm - nothing found in the logfile. amkor.taiwan.confirm - nothing found in the logfile. amkor.korea.confirm - nothing found in the logfile. atlantic.confirm - nothing found in the logfile. optopac.confirm - nothing found in the logfile. carsem1.confirm - nothing found in the logfile. carsem2.confirm - nothing found in the logfile. mpd.confirm - nothing found in the logfile. senodia.confirm - nothing found in the logfile. unisem.confirm - nothing found in the logfile. pactech.confirm - nothing found in the logfile. pactech.processed.from - nothing found in the logfile. ispl.confirm - nothing found in the logfile. umc-import.from - nothing found in the logfile. leister.from - nothing found in the logfile. vishay.confirm - nothing found in the logfile. xfab.from - nothing found in the logfile. heimann.confirm - nothing found in the logfile.
[1364288233] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1364288260] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1364288261] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364288262] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1364288270] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1364288310] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364288310] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;Connection refused
[1364288311] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364288324] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364288324] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364288333] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364288400] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364288409] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364288413] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;Server down!
[1364288420] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1364288420] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1364288420] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1364288420] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1364288420] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1364288440] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1364288440] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1364288440] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1364288441] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1364288441] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=heimann.confirm
[1364288441] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1364288441] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1364288441] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1364288441] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1364288443] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1364288449] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1364288451] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1364288451] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1364288461] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1364288461] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1364288469] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1364288470] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia.confirm
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1364288471] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1364288472] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1364288479] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364288500] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1364288512] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1364288609] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;HARD;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364288610] SERVICE ALERT: ewaf-test.colo.elex.be;SERVICEMIX;OK;SOFT;2;PROCS OK: 1 process with UID = 999 (smx-fuse), args '/bin/java'
[1364288639] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364288660] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1364288689] HOST ALERT: gambit-test.colo.elex.be;UP;SOFT;2;Server alive
[1364288692] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1364288695] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1364288696] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1364288699] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;Connection refused
[1364288701] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364288710] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1364288720] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364288720] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.87 %
[1364288749] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;HARD;3;Connection refused
[1364288771] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364288812] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364288812] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=3254.00 KBps
[1364288851] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 61.06 ms
[1364289022] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289132] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1364289139] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289165] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;OK;SOFT;2;TCP OK - 0.009 second response time on port 6445
[1364289169] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289179] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364289189] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289210] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;No route to host
[1364289223] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;Server down!
[1364289229] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289229] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;HARD;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289239] SERVICE ALERT: gambit-test.colo.elex.be;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.32.241.183)
[1364289249] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;OK;HARD;3;OK - all polling (s)ftp connections are fresh.
[1364289251] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1364289264] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;HARD;3;TCP OK - 0.000 second response time on port 61616
[1364289279] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289300] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 72%, RTA = 58.22 ms
[1364289315] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364289359] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289369] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364289370] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - null
[1364289373] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289373] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289373] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289374] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289374] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289374] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289374] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289374] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289374] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289376] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289376] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289376] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289376] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289376] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1364289379] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;Connection refused
[1364289401] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1364289409] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;Connection refused
[1364289414] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1364289562] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=83 ms
[1364289569] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;2;Server down!
[1364289590] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1364289593] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1364289599] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1364289649] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;1;No route to host
[1364289750] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;SOFT;1;CRITICAL -  amkor.philippines.confirm - nothing found in the logfile. amkor.taiwan.confirm - nothing found in the logfile. amkor.korea.confirm - nothing found in the logfile. atlantic.confirm - nothing found in the logfile. optopac.confirm - nothing found in the logfile. carsem1.confirm - nothing found in the logfile. carsem2.confirm - nothing found in the logfile. mpd.confirm - nothing found in the logfile. senodia.confirm - nothing found in the logfile. unisem.confirm - nothing found in the logfile. pactech.confirm - nothing found in the logfile. pactech.processed.from - nothing found in the logfile. ispl.confirm - nothing found in the logfile. umc-import.from - nothing found in the logfile. leister.from - nothing found in the logfile. vishay.confirm - nothing found in the logfile. xfab.from - nothing found in the logfile. heimann.confirm - nothing found in the logfile.
[1364289751] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364289751] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364289754] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364289754] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364289755] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364289756] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364289756] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364289761] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;3;TCP OK - 0.000 second response time on port 61616
[1364289776] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289776] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289779] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289779] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289779] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289789] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289791] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289791] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289791] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289791] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289791] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289792] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289792] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289793] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289793] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289794] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289794] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289794] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289795] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1364289812] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289812] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289812] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289812] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289812] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289813] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289813] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289813] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289813] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289816] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289816] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289819] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289819] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289819] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289849] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364289853] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364289992] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364289992] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364289992] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364289992] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364289992] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364289995] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290000] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290000] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290000] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290000] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290001] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290010] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290019] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;3;Server down!
[1364290049] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290049] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290061] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290069] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290070] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290070] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290070] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290071] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290071] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290071] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290071] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290072] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290081] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290083] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290090] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290105] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290110] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290110] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290110] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290112] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290112] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290125] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290125] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290125] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290144] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290145] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290145] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364290145] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290161] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364290180] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;OK;SOFT;2;OK - all polling (s)ftp connections are fresh.
[1364290191] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364290240] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364290284] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1364290460] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;4;Server down!
[1364290542] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1364290561] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 12.36, 11.02, 7.33
[1364290584] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1364290670] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1364290715] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1364290915] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;5;Server down!
[1364291000] SERVICE ALERT: atlan.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 4.44, 6.55, 6.71
[1364291023] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1364291032] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291032] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291033] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291033] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291033] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291040] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1364291041] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291042] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291042] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291043] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291043] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291050] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291051] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291072] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1364291093] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291093] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291109] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1364291110] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291113] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291114] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291114] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291115] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291115] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291115] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291119] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291119] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291119] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291130] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291132] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291133] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291153] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291153] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291153] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291154] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291155] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291156] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291173] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291173] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291173] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291190] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291190] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291190] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291190] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291212] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1364291231] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291281] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291281] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291291] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291292] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291293] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291293] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291294] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291313] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291313] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291313] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291314] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291314] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291320] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291320] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291320] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291320] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291320] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291321] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291322] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291322] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291323] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291323] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291324] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291330] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291330] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291340] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291341] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291341] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291341] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291341] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291341] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291342] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291342] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291342] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291350] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291350] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291351] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291351] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291351] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291360] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;6;Server down!
[1364291372] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 9.34, 7.23, 5.98
[1364291380] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291391] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;(Service Check Timed Out)
[1364291471] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291471] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291471] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291471] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291471] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291473] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291473] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1364291474] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291480] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291481] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291481] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291492] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291492] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291509] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=60 ms
[1364291519] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1364291532] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291533] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291549] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364291551] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291551] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291552] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291552] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291554] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291554] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291554] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291562] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291562] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291562] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291569] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 9.008 second response time on port 61616
[1364291570] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291570] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291571] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291589] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291589] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291590] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291590] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291590] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291590] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291609] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291610] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291610] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291629] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291629] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291630] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291630] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291646] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364291649] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1364291671] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291692] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364291719] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291719] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291720] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364291729] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291730] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291730] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291741] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291741] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291749] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364291749] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364291749] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291749] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291750] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364291750] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291750] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291759] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection reset
[1364291759] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 0.001 second response time on port 61616
[1364291763] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291763] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291764] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291764] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291764] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291764] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291765] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291765] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291765] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291766] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291770] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291770] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291782] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364291783] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291783] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291783] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291783] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291784] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1364291784] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291784] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291784] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291791] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291791] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291791] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291791] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291791] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291800] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364291809] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;7;Server down!
[1364291811] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 3.61, 5.32, 5.63
[1364291823] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291829] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;3;TCP OK - 0.001 second response time on port 61616
[1364291840] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;2;(Service Check Timed Out)
[1364291859] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364291911] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=59 ms
[1364291911] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291912] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291912] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291913] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291914] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291914] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291915] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291919] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291920] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291921] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291933] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291934] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291952] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1364291959] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1364291959] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1364291974] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291975] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364291999] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292000] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292001] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292002] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292002] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292003] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292004] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364292004] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292009] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364292010] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292010] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292013] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292013] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292040] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - Connection reset
[1364292043] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292044] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292045] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292046] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292048] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292052] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292054] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292055] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292091] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292092] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292094] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292098] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292119] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1364292120] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=71 ms
[1364292128] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364292136] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292162] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364292184] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;SOFT;1;CRITICAL -  amkor.philippines.confirm - last conection too old(>15min). amkor.taiwan.confirm - last conection too old(>15min). amkor.korea.confirm - last conection too old(>15min). atlantic.confirm - last conection too old(>15min). optopac.confirm - last conection too old(>15min). carsem1.confirm - last conection too old(>15min). carsem2.confirm - last conection too old(>15min). mpd.confirm - last conection too old(>15min). senodia.confirm - last conection too old(>15min). unisem.confirm - last conection too old(>15min). pactech.confirm - last conection too old(>15min). pactech.processed.from - last conection too old(>15min). ispl.confirm - last conection too old(>15min). umc-import.from - last conection too old(>15min). leister.from - last conection too old(>15min). vishay.confirm - last conection too old(>15min). xfab.from - last conection too old(>15min). heimann.confirm - last conection too old(>15min).
[1364292201] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292202] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292206] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292207] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292207] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292217] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292218] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292219] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292224] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364292225] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292226] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292227] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292227] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292237] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292238] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292239] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292240] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292241] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292241] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292242] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292243] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292243] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292245] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292246] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292246] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292247] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292257] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292257] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292268] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292269] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292270] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292270] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292271] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292272] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292272] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292273] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1364292274] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292275] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292275] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292276] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292277] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292314] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;8;Server down!
[1364292315] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292344] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;HARD;3;(Service Check Timed Out)
[1364292345] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364292484] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=4.80% system=0.20% iowait=82.80% idle=12.20%
[1364292543] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1364292639] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1364292640] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364292691] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;SOFT;2;CRITICAL -  amkor.philippines.confirm - last conection too old(>15min). amkor.taiwan.confirm - last conection too old(>15min). amkor.korea.confirm - last conection too old(>15min). atlantic.confirm - last conection too old(>15min). optopac.confirm - last conection too old(>15min). carsem1.confirm - nothing found in the logfile. carsem2.confirm - nothing found in the logfile. mpd.confirm - last conection too old(>15min). senodia.confirm - last conection too old(>15min). unisem.confirm - last conection too old(>15min). pactech.confirm - last conection too old(>15min). pactech.processed.from - last conection too old(>15min). ispl.confirm - last conection too old(>15min). umc-import.from - last conection too old(>15min). leister.from - last conection too old(>15min). vishay.confirm - last conection too old(>15min). xfab.from - last conection too old(>15min). heimann.confirm - last conection too old(>15min).
[1364292728] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;3;TCP OK - 9.000 second response time on port 61616
[1364292770] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;OK;SOFT;3;OK: Connection to cfengine master OK
[1364292797] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364292818] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;9;Server down!
[1364292833] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364292878] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 66.21 ms
[1364292939] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.20% iowait=0.00% idle=99.80%
[1364292981] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1364293071] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (59.5966958211856)
[1364293128] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;HARD;3;CRITICAL -  amkor.philippines.confirm - last conection too old(>15min). amkor.taiwan.confirm - last conection too old(>15min). amkor.korea.confirm - last conection too old(>15min). atlantic.confirm - last conection too old(>15min). optopac.confirm - last conection too old(>15min). carsem1.confirm - nothing found in the logfile. carsem2.confirm - nothing found in the logfile. mpd.confirm - last conection too old(>15min). senodia.confirm - last conection too old(>15min). unisem.confirm - nothing found in the logfile. pactech.confirm - last conection too old(>15min). pactech.processed.from - last conection too old(>15min). ispl.confirm - last conection too old(>15min). umc-import.from - last conection too old(>15min). leister.from - last conection too old(>15min). vishay.confirm - last conection too old(>15min). xfab.from - last conection too old(>15min). heimann.confirm - last conection too old(>15min).
[1364293173] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1364293259] HOST ALERT: gambit-test.colo.elex.be;DOWN;HARD;10;Server down!
[1364293410] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1364293973] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1364294101] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_CPU;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1364294110] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364294152] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1364294198] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364294461] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1364294501] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1364294532] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1364294578] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1364294611] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 0.000 second response time on port 61616
[1364294836] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi SWAP;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1364294872] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1364294891] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364294932] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364294940] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364294963] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1364295008] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1364295361] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1364295368] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1364295381] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1364295599] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364295638] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;SSH;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1364295791] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364295794] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1364295811] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1364295841] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364295841] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Ambient Temperature;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364295911] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1364296050] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 2.997 second response time on port 61616
[1364296133] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1364296269] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;PING;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1364296332] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Ambient Temperature;OK;SOFT;2;OK System-Temp-Ambient (1450) 14.50 Centigrade
[1364296332] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1364296334] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1364296341] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1364296355] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364296396] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1364296531] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364296582] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364296812] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1364296831] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1364296850] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364296873] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1364297019] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1364297074] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364297193] SERVICE ALERT: nki-laptop.sofia.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364297210] HOST ALERT: nki-laptop.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364297300] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1364297320] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1364297329] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364297329] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1364297519] HOST ALERT: nki-laptop.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 42.77 ms
[1364297551] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364297639] SERVICE ALERT: nki-laptop.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 42.95 ms
[1364297825] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1364297930] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1364298069] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;HARD;3;TCP OK - 8.997 second response time on port 61616
[1364298259] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi Runtime status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1364298294] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1364298319] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364298624] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1364298673] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1364298759] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1364298781] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364298794] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364298855] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364298891] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1364298962] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi Network usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1364298989] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 58.01 ms
[1364299101] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364299112] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1364299219] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1364299379] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364299462] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1364299489] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 81%, RTA = 58.60 ms
[1364299516] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;SOFT;2;TCP OK - 0.000 second response time on port 61616
[1364299660] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1364299859] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364299941] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=93 ms
[1364299957] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 73%, RTA = 54.64 ms
[1364300117] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1364300130] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364300168] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300169] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300169] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300169] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300169] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300169] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300169] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300169] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300171] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300171] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300173] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300187] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300188] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300188] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300191] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300192] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300192] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300209] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300210] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300210] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300210] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300211] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300212] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300230] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300231] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300232] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300249] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300249] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300251] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364300251] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300256] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300268] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;OK;HARD;3;OK - all polling (s)ftp connections are fresh.
[1364300276] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1364300291] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300306] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364300337] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300346] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300347] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300347] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300347] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300348] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1364300357] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300357] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300366] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300367] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300367] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300367] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300368] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300372] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 12.38, 10.23, 6.68
[1364300373] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=89 ms
[1364300374] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300377] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300377] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300386] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300386] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300386] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300387] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300387] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300388] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300388] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300389] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300390] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300391] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300398] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300407] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300407] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300407] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300407] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300407] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300407] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300408] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300408] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300408] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300408] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300408] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300409] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300412] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300447] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300466] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300559] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364300578] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300578] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300578] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300578] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300580] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300580] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300580] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300582] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1364300582] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300587] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1364300607] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364300679] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1364300728] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1364300766] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364300771] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1364300799] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 5.02, 6.28, 6.12
[1364300863] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=94 ms
[1364300863] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1364300986] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364301038] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364301157] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;HARD;1;P1.PL OK - io write latency=0 ms
[1364301170] HOST ALERT: gambit-test.colo.elex.be;UP;HARD;1;Server alive
[1364301199] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1364301246] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;HARD;1;P1.PL OK - cpu usage=1.28 %
[1364301266] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364301279] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;1;P1.PL OK - overall status=green
[1364301326] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;1;SSH OK - dropbear_0.52 (protocol 2.0)
[1364301356] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1364301378] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;HARD;1;P1.PL OK - io read latency=0 ms
[1364301416] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364301437] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.24 ms
[1364301468] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364301469] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364301496] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;1;P1.PL OK - swap usage=0.00 MB
[1364301517] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364301532] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1364301612] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364301632] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=85 ms
[1364301678] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364301678] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364301678] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364301679] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364301679] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364301687] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1364301698] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1364301699] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1364301699] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1364301701] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1364301719] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1364301719] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1364301720] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1364301728] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1364301738] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;SOFT;1;CRITICAL -  amkor.philippines.confirm - nothing found in the logfile. amkor.taiwan.confirm - nothing found in the logfile. amkor.korea.confirm - nothing found in the logfile. atlantic.confirm - nothing found in the logfile. optopac.confirm - nothing found in the logfile. carsem1.confirm - nothing found in the logfile. carsem2.confirm - nothing found in the logfile. mpd.confirm - nothing found in the logfile. senodia.confirm - nothing found in the logfile. unisem.confirm - nothing found in the logfile. pactech.confirm - nothing found in the logfile. pactech.processed.from - nothing found in the logfile. ispl.confirm - nothing found in the logfile. umc-import.from - nothing found in the logfile. leister.from - nothing found in the logfile. vishay.confirm - nothing found in the logfile. xfab.from - nothing found in the logfile. heimann.confirm - nothing found in the logfile.
[1364301761] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1364301810] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia.confirm
[1364301816] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1364301817] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1364301818] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1364301818] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1364301819] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364301829] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1364301829] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1364301829] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1364301838] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1364301838] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1364301838] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1364301838] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1364301848] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364301848] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364301849] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1364301856] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1364301858] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1364301858] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=heimann.confirm
[1364301858] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1364301859] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1364301859] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1364301860] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1364301860] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1364301860] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1364301860] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1364301866] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1364301866] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1364301868] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 1.23 ms
[1364301878] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1364301878] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1364301879] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1364301879] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1364301879] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1364301879] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1364301879] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1364301880] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1364301880] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1364301880] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia.confirm
[1364301881] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1364301881] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1364301881] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1364301888] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1364301920] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1364301926] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364301939] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1364301958] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364301970] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1364302020] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;Connection refused
[1364302049] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1364302050] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1364302050] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1364302051] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1364302051] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1364302058] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1364302058] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1364302058] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1364302059] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1364302059] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1364302067] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1364302077] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1364302078] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1364302110] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1364302112] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1364302112] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1364302116] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;Connection refused
[1364302117] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1364302117] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1364302126] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1364302126] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1364302127] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1364302129] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1364302129] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=heimann
[1364302130] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=heimann
[1364302130] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1364302133] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1364302133] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=heimann.confirm
[1364302136] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;HARD;3;Connection refused
[1364302138] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1364302138] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1364302139] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1364302159] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364302178] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;SOFT;2;CRITICAL -  amkor.philippines.confirm - nothing found in the logfile. amkor.taiwan.confirm - nothing found in the logfile. amkor.korea.confirm - nothing found in the logfile. atlantic.confirm - nothing found in the logfile. optopac.confirm - nothing found in the logfile. carsem1.confirm - nothing found in the logfile. carsem2.confirm - nothing found in the logfile. mpd.confirm - nothing found in the logfile. senodia.confirm - nothing found in the logfile. unisem.confirm - nothing found in the logfile. pactech.confirm - nothing found in the logfile. pactech.processed.from - nothing found in the logfile. ispl.confirm - nothing found in the logfile. umc-import.from - nothing found in the logfile. leister.from - nothing found in the logfile. vishay.confirm - nothing found in the logfile. xfab.from - nothing found in the logfile. heimann.confirm - nothing found in the logfile.
[1364302249] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364302301] SERVICE ALERT: hudson2-test.colo.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1364302307] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 80%, RTA = 0.19 ms
[1364302346] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=4090.00 KBps
[1364302348] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1364302380] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364302428] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364302481] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1364302536] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1364302558] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=10 ms
[1364302577] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 Server closed connection without sending any data back
[1364302606] SERVICE ALERT: ewaf-test.colo.elex.be;EWAF_POLLING_FTP;CRITICAL;HARD;3;CRITICAL -  amkor.philippines.confirm - nothing found in the logfile. amkor.taiwan.confirm - nothing found in the logfile. amkor.korea.confirm - nothing found in the logfile. atlantic.confirm - nothing found in the logfile. optopac.confirm - nothing found in the logfile. carsem1.confirm - nothing found in the logfile. carsem2.confirm - nothing found in the logfile. mpd.confirm - nothing found in the logfile. senodia.confirm - nothing found in the logfile. unisem.confirm - nothing found in the logfile. pactech.confirm - nothing found in the logfile. pactech.processed.from - nothing found in the logfile. ispl.confirm - nothing found in the logfile. umc-import.from - nothing found in the logfile. leister.from - nothing found in the logfile. vishay.confirm - nothing found in the logfile. xfab.from - nothing found in the logfile. heimann.confirm - nothing found in the logfile.
[1364302618] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1364302688] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1364302756] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=0.98 %
[1364302763] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364302782] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 76%, RTA = 2.30 ms
[1364302817] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364302846] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364302887] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364302888] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364302889] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364302909] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 59.19 ms
[1364302941] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1364302956] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364302958] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1364303079] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303098] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1364303101] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1364303197] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1364303201] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364303262] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364303317] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.26 ms
[1364303330] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303357] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303357] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303358] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303378] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303389] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 13.48, 3.43, 1.29
[1364303409] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1364303458] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;HARD;3;TCP OK - 0.000 second response time on port 61616
[1364303458] SERVICE FLAPPING ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1364303470] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303477] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303477] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303483] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364303487] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303497] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1364303509] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303527] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303561] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303562] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1364303562] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303568] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303568] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303568] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303581] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303582] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303582] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303588] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303588] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303588] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303589] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303602] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303602] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303608] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303609] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303618] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303618] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303618] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303619] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303619] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303619] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303620] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303621] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303621] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303629] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303629] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303629] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303629] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303629] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303641] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303641] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303642] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303642] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303642] SERVICE ALERT: ewaf-test.colo.elex.be;senodia.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303642] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303642] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303642] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303647] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303647] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1364303683] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303717] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303728] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364303747] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364303747] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=1960.00 KBps
[1364303747] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.86 %
[1364303758] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 0.21 ms
[1364303772] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303777] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.08 %
[1364303799] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303800] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303800] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303807] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303813] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303813] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303813] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303818] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303818] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303819] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364303819] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.01, 1.15, 1.09
[1364303820] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303830] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303830] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303848] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1364303862] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364303869] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364303869] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364303869] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364303869] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364303869] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303869] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303870] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364303870] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303871] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303871] SERVICE ALERT: ewaf-test.colo.elex.be;heimann_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303878] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303882] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364303882] SERVICE ALERT: ewaf-test.colo.elex.be;heimann.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364303890] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;OK;SOFT;3;OK - InFlightCount=0
[1364303890] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364303890] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;OK;SOFT;3;OK - QueueSize=0
[1364303910] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364303920] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1364304001] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1364304042] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364304082] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1364304149] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364304158] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364304198] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 0.17 ms
[1364304227] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=2 ms
[1364304280] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1364304351] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1364304478] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364304491] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1364304618] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364304667] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 64%, RTA = 0.18 ms
[1364304667] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1364304769] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1364304787] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364304787] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364304788] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364304837] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364304857] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=8 ms
[1364304879] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1364304920] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1364304947] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364305029] SERVICE ALERT: pintail2-test.erfurt.elex.be;SYSLOG;CRITICAL;SOFT;1;Connection refused or timed out
[1364305048] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364305127] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364305170] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.17 ms
[1364305201] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1364305228] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364305228] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364305228] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364305267] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=96 ms
[1364305311] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1364305383] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1364305383] SERVICE FLAPPING ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1364305383] SERVICE FLAPPING ALERT: ldap-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (26.0% change >= 20.0% threshold)
[1364305383] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1364305383] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (31.1% change >= 20.0% threshold)
[1364305383] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (31.7% change >= 20.0% threshold)
[1364305575] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364305635] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1364305645] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364305655] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=44 ms
[1364305655] SERVICE ALERT: imc-test.colo.elex.be;CFAGENT;CRITICAL;SOFT;1;/var/lib/cfengine2/outputs/previous is missing!
[1364305719] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 70.06 ms
[1364305724] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364305818] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1364305825] SERVICE ALERT: imc-test.colo.elex.be;CFAGENT;CRITICAL;SOFT;2;/var/lib/cfengine2/outputs/previous is missing!
[1364305905] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1364305905] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1364305924] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1364306044] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.86 %
[1364306056] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1364306067] SERVICE ALERT: imc-test.colo.elex.be;CFAGENT;CRITICAL;HARD;3;/var/lib/cfengine2/outputs/previous is missing!
[1364306139] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364306164] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=107 ms
[1364306345] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364306378] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1364306407] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364306464] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1364306498] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364306574] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.02 %
[1364306609] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=77 ms
[1364306664] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364306726] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=1923.00 KBps
[1364306735] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364306754] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1364306811] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.17 ms
[1364306834] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364306844] SERVICE ALERT: pintail2-test.erfurt.elex.be;SYSLOG;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'syslogd'
[1364307021] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364307219] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1364307260] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1364307280] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364307291] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364307295] SERVICE ALERT: bifur-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364307309] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 76%, RTA = 0.23 ms
[1364307309] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1364307432] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.86 %
[1364307474] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364307511] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1364307543] HOST ALERT: gambit-test.colo.elex.be;UP;SOFT;2;Server alive
[1364307570] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364307650] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364307702] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364307704] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1364307759] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364307760] SERVICE ALERT: bifur-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=4 ms
[1364307829] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364307861] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364307879] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1364307939] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364307979] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364307994] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364308016] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364308040] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1364308122] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (47.8377065111759)
[1364308123] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308123] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308129] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308129] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308129] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308130] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308130] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308133] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308140] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308140] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - Connection refused
[1364308160] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1364308161] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1364308162] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1364308162] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1364308162] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1364308243] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1364308259] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1364308300] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364308309] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364308319] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364308330] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364308349] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1364308390] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 83%, RTA = 0.21 ms
[1364308403] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364308413] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364308452] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364308463] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364308491] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364308510] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364308546] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1364308546] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1364308610] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308611] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364308611] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308611] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364308611] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364308611] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308612] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308623] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308623] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308623] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364308650] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364308650] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308651] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1364308651] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364308651] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1364308741] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1364308794] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364308799] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1364308815] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1364308880] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 4.28 ms
[1364308895] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2576.00 KBps
[1364308895] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364308949] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364308979] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364309011] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364309016] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364309035] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1364309055] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1364309144] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364309232] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1364309239] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364309325] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364309349] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 72%, RTA = 0.17 ms
[1364309389] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=11 ms
[1364309451] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364309469] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1364309503] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1364309511] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;HARD;3;P1.PL OK - mem usage=2.86 %
[1364309563] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1364309641] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364309821] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364309865] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2441.00 KBps
[1364309891] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 1.18 ms
[1364309929] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1364309964] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1364309984] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364310045] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1364310050] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364310080] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364310103] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364310125] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=108 ms
[1364310215] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 54.52 ms
[1364310306] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.22 %
[1364310397] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.16 ms
[1364310539] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364310546] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364310555] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364310588] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=111 ms
[1364310595] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364310618] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364310626] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1364310677] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=105 ms
[1364310686] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364310745] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 55.04 ms
[1364310807] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364310920] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364310937] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 86%, RTA = 1.65 ms
[1364310988] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1364311046] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1364311056] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;3;P1.PL CRITICAL - Server version unavailable at 'https://10.32.241.183:443/sdk/vimService.wsdl' at /usr/share/perl5/VMware/VICommon.pm line 544.
[1364311059] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364311085] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1364311106] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364311140] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1364311148] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364311177] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1364311300] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1364311307] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 54.83 ms
[1364311362] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1364311485] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1364311536] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1364311549] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;HARD;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364311569] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=52 ms
[1364311588] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=105 ms
[1364311610] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=10 ms
[1364311639] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1364311658] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364311705] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364311705] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi Network usage;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1364311790] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 73%, RTA = 54.64 ms
[1364311850] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364311970] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;CRITICAL;HARD;3;CRITICAL: Unable to establish connection with cfengine master server
[1364312068] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1364312131] HOST ALERT: hma-lp.sofia.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1364312135] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364312136] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364312148] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364312155] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=2954.00 KBps
[1364312175] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364312256] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1364312327] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1364312585] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364312625] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364312629] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1364312667] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364312739] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1364312800] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364312821] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364313077] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 12.35, 6.99, 5.91
[1364313106] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (47.8377065111759)
[1364313108] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364313157] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364313159] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364313167] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364313186] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364313207] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1364313215] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364313227] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;HARD;3;CRITICAL: Unable to establish connection with cfengine master server
[1364313287] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364313307] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1364313307] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter  Blades communication;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364313317] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1364313555] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 7.35, 7.30, 6.46
[1364313588] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1364313637] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364313648] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364313675] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364313700] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364313745] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1364313787] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1364313788] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter  Blades communication;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1364313790] SERVICE FLAPPING ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1364313826] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364313927] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364313947] SERVICE ALERT: murray-test.sensors.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 13.52, 17.09, 10.63
[1364313956] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;UNKNOWN;SOFT;1;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 Connect failed: connect: Connection timed out: Connection timed out
[1364314171] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364314185] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364314186] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1364314287] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364314314] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter  Blades communication;OK;SOFT;3;OK Blades-Comm (0) OK
[1364314405] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364314411] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364314423] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1364314441] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=0.93 %
[1364314473] SERVICE ALERT: murray-test.sensors.elex.be;LOAD;OK;SOFT;2;OK - load average: 4.94, 7.74, 8.81
[1364314492] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364314731] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364314732] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364314815] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364314843] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364314877] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364314896] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1364314917] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=2273.00 KBps
[1364315002] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364315034] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364315153] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364315195] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364315285] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1364315316] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1364315373] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1364315481] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.63 %
[1364315508] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364315601] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364315611] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364315662] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364315751] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=69 ms
[1364315891] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1364315912] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2998.00 KBps
[1364316012] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.86 %
[1364316031] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1364316032] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364316071] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364316256] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364316274] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1364316344] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364316352] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364316483] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;3;P1.PL OK - overall status=green
[1364316484] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1364316495] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364316712] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364316791] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1364316817] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1364316821] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364316865] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364316965] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=2.05 %
[1364317021] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1364317068] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364317122] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364317202] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364317202] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi SWAP;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1364317252] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=2 ms
[1364317273] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=115 ms
[1364317301] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1364317325] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 16.54, 9.88, 7.36
[1364317353] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1364317521] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364317574] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364317574] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364317574] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364317642] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364317823] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=64 ms
[1364317851] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1364317874] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 6.22, 7.22, 7.25
[1364317876] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;UNKNOWN;HARD;3;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 Connect failed: connect: Connection timed out: Connection timed out
[1364317902] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=71 ms
[1364318032] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1364318035] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364318113] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=133 ms
[1364318114] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364318114] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364318126] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=0.89 %
[1364318301] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364318317] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=0 ms
[1364318375] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1364318401] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1364318401] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1364318423] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 12.59, 10.22, 8.55
[1364318452] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=107 ms
[1364318497] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364318577] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=2428.00 KBps
[1364318594] SERVICE ALERT: gambit-test.colo.elex.be;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1364318605] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1364318768] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364318902] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 6.80, 7.11, 7.61
[1364318921] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1364318945] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1364319002] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1364319103] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364319142] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 96%, RTA = 0.21 ms
[1364319235] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.87 %
[1364319403] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1364319431] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1364319586] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364319614] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364319641] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 10.04, 5.93, 5.29
[1364319674] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364319701] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364319703] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1364319921] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=96 ms
[1364319952] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1364320077] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364320083] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364320123] SERVICE ALERT: esb-a-test.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 2.43, 4.06, 4.68
[1364320161] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1364320161] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2057.00 KBps
[1364320186] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1364320206] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364320243] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364320346] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364320443] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1364320471] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1364320576] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364320675] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1364320695] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1364320703] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364320742] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364320791] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1364320834] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364320893] SERVICE ALERT: imc-test.colo.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1364321011] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1364321082] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 17.48, 14.03, 9.17
[1364321097] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364321121] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364321173] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=1493.00 KBps
[1364321229] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1364321274] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364321502] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364321560] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364321590] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1364321611] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1364321630] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 9.53, 11.38, 10.11
[1364321652] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364321662] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1364321692] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364321733] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364321794] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1364321795] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364321893] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (54.7862001943635)
[1364321938] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364321998] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364322078] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1364322078] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;UNKNOWN;SOFT;1;**ePN /usr/lib/nagios/plugins/check_esx: "SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364322119] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 11.09, 11.24, 10.55
[1364322142] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364322159] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;2;P1.PL OK - swap usage=0.00 MB
[1364322200] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364322220] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364322221] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364322269] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;HARD;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364322269] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1364322281] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1364322282] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1364322391] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364322533] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364322533] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364322560] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364322560] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1364322574] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.22 %
[1364322638] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=7 ms
[1364322654] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1364322669] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 7.21, 8.42, 9.58
[1364322710] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1364322712] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364322772] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=12 ms
[1364322833] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364322870] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364322881] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1364323018] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364323021] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364323051] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364323099] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364323132] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1364323250] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364323301] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1364323310] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364323318] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1364323478] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;3;P1.PL OK - mem usage=2.86 %
[1364323599] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364323650] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1364323790] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1364323798] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=68 ms
[1364323815] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364323902] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364324038] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;3;P1.PL OK - overall status=green
[1364324059] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364324198] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1364324262] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=56 ms
[1364324308] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364324339] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364324409] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364324440] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=2.80% system=2.00% iowait=90.51% idle=4.70%
[1364324553] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364324561] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364324678] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1364324748] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=5 ms
[1364324792] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364324930] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;3;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364324941] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1364324946] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364324962] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364324986] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.30% system=0.10% iowait=0.00% idle=99.60%
[1364325002] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364325131] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2464.00 KBps
[1364325222] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1364325240] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=0 ms
[1364325314] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364325422] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1364325430] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364325450] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;3;P1.PL OK - cpu usage=1.51 %
[1364325451] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364325470] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364325482] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364325773] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1364325831] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364325866] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1364325914] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1364325922] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364325950] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364325970] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364325976] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1364326255] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364326262] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1364326324] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364326354] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1364326355] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 78%, RTA = 0.17 ms
[1364326360] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364326460] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1364326463] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364326463] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;SSH;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1364326482] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364326525] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364326560] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364326582] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364326590] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1364326704] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.87 %
[1364326805] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364326840] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1364326962] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1364327001] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.49 %
[1364327014] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364327072] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364327084] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364327140] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364327292] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364327330] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1364327344] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364327416] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364327434] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1364327491] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1364327520] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364327564] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364327620] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1364327773] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364327820] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364327860] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=2253.00 KBps
[1364327892] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364327900] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1364327970] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 1.56 ms
[1364328005] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1364328011] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1364328052] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1364328063] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364328063] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364328065] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364328065] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=113 ms
[1364328080] HOST ALERT: gambit-test.colo.elex.be;UP;SOFT;2;Server alive
[1364328092] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=62 ms
[1364328102] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364328113] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 17.23, 8.86, 6.06
[1364328131] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.32.241.183:443/sdk/webService': Perhaps host is not a vCenter or ESX server
[1364328213] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364328271] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;3;P1.PL OK - overall status=green
[1364328311] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1364328350] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364328403] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 96%, RTA = 0.16 ms
[1364328545] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1364328555] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=0.84 %
[1364328558] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;1;SSH OK - dropbear_0.52 (protocol 2.0)
[1364328558] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;HARD;1;P1.PL OK - io write latency=8 ms
[1364328558] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1364328586] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364328606] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 4.50, 5.62, 5.66
[1364328616] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1364328617] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364328630] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364328769] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1364328804] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1364328834] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1364328888] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 0.17 ms
[1364328926] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364328944] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364329054] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL CRITICAL - Server version unavailable at 'https://10.32.241.183:443/sdk/vimService.wsdl' at /usr/share/perl5/VMware/VICommon.pm line 544.
[1364329056] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364329075] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329076] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1364329104] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL CRITICAL - Server version unavailable at 'https://10.32.241.183:443/sdk/vimService.wsdl' at /usr/share/perl5/VMware/VICommon.pm line 544.
[1364329117] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329194] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364329207] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1364329226] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;HARD;3;P1.PL OK - mem usage=2.87 %
[1364329236] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=64 ms
[1364329254] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=2.99% system=1.20% iowait=79.84% idle=15.97%
[1364329326] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2470.00 KBps
[1364329366] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364329374] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1364329434] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;HARD;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364329477] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364329505] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1364329507] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329526] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=96 ms
[1364329547] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329614] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (59.5966958211856)
[1364329684] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;2;CPU IOWAIT CRITICAL : user=1.70% system=1.00% iowait=84.82% idle=12.49%
[1364329699] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1364329734] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1364329795] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1364329814] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1364329864] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 68%, RTA = 0.21 ms
[1364329865] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329865] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329905] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364329918] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364329947] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329966] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=86 ms
[1364329986] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364329988] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364330077] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=192 ms
[1364330078] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364330094] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1364330114] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;3;CPU STATISTICS OK : user=3.50% system=0.70% iowait=46.45% idle=49.35%
[1364330127] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1364330295] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364330295] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;2;P1.PL OK - net usage=2790.00 KBps
[1364330326] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1364330399] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=69 ms
[1364330424] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364330457] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.32 %
[1364330514] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1364330528] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1364330625] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364330716] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364330797] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=181 ms
[1364330807] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364330864] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.15 ms
[1364330926] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=116 ms
[1364330935] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=7 ms
[1364330957] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1364330986] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=1 ms
[1364330986] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi IO read;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1364331005] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364331066] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1364331077] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=55 ms
[1364331149] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1364331234] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364331277] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364331306] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 0.24 ms
[1364331366] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364331368] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1364331388] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1364331454] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364331464] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=128 ms
[1364331517] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364331544] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364331559] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1364331587] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=112 ms
[1364331614] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=3.80% system=0.90% iowait=80.40% idle=14.90%
[1364331726] SERVICE ALERT: lechuck-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1364331737] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1364331755] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;HARD;3;PING WARNING - Packet loss = 86%, RTA = 0.17 ms
[1364331796] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364331815] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364331816] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364331892] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1364331895] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;HARD;3;P1.PL OK - swap usage=0.00 MB
[1364331907] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1364331907] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364332014] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364332051] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1364332108] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;2;CPU IOWAIT CRITICAL : user=0.90% system=0.60% iowait=84.88% idle=13.61%
[1364332128] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 16.53, 18.42, 12.30
[1364332147] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=44 ms
[1364332230] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364332237] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364332310] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1364332312] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364332350] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364332381] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1364332407] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;HARD;3;P1.PL OK - cpu usage=1.50 %
[1364332457] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=1 ms
[1364332529] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1364332538] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364332567] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;HARD;3;CPU IOWAIT CRITICAL : user=2.69% system=2.50% iowait=94.81% idle=0.00%
[1364332572] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 17.32, 16.63, 13.68
[1364332729] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364332739] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1364332797] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364332821] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1364332958] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364332980] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364332987] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364333002] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 13.91, 13.56, 13.28
[1364333169] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1364333170] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1364333189] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=3877.00 KBps
[1364333293] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364333307] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364333318] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1364333373] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=0 ms
[1364333420] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364333429] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364333457] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364333498] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364333498] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 11.66, 10.32, 11.57
[1364333498] SERVICE FLAPPING ALERT: dvorak-test.colo.elex.be;LOAD;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1364333548] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;HARD;3;CPU IOWAIT WARNING : user=3.70% system=2.10% iowait=79.68% idle=14.51%
[1364333628] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1364333730] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1364333748] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1364333759] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364333821] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364333859] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364333888] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364333899] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364333940] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364333941] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1364333991] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.44, 6.45, 9.14
[1364334047] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;HARD;3;CPU IOWAIT CRITICAL : user=2.20% system=3.10% iowait=94.71% idle=0.00%
[1364334121] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1364334163] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1364334222] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334297] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334319] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1364334373] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=93 ms
[1364334378] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334410] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL CRITICAL - SOAP request error - possibly a protocol issue: 500 SSL read timeout:
[1364334428] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334499] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1364334538] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;HARD;3;CPU IOWAIT WARNING : user=5.89% system=2.30% iowait=76.92% idle=14.89%
[1364334538] SERVICE FLAPPING ALERT: orpmelei-test.apps.elex.be;CPU_STATS;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1364334553] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1364334617] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364334651] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364334733] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1364334810] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1364334828] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334838] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364334851] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334852] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334892] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364334917] HOST ALERT: gambit-test.colo.elex.be;DOWN;SOFT;1;Server down!
[1364334932] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364334971] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 11.67, 10.65, 9.23
[1364334988] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1364335028] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;HARD;3;CPU IOWAIT CRITICAL : user=2.20% system=1.80% iowait=95.91% idle=0.10%
[1364335081] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1364335109] HOST ALERT: gambit-test.colo.elex.be;UP;SOFT;2;Server alive
[1364335163] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1364335222] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;HARD;3;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364335282] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335289] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335310] SERVICE ALERT: janus-test.colo.elex.be;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364335330] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335368] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335369] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335410] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 28.32, 17.93, 12.22
[1364335428] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1364335589] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1364335636] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1364335730] SERVICE ALERT: gambit-test.colo.elex.be;ESXi SWAP;OK;SOFT;3;P1.PL OK - swap usage=0.00 MB
[1364335788] SERVICE ALERT: janus-test.colo.elex.be;HTTP;OK;SOFT;2;HTTP OK HTTP/1.1 200 OK - 8265 bytes in 0.684 seconds
[1364335805] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;HARD;3;P1.PL OK - cpu usage=0.73 %
[1364335805] SERVICE FLAPPING ALERT: gambit-test.colo.elex.be;ESXi_CPU;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1364335824] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335828] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335832] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1364335834] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;HARD;3;P1.PL OK - net usage=1183.00 KBps
[1364335848] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1364335864] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364335896] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 11.17, 13.86, 12.89
[1364335912] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=169 ms
[1364336075] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;HARD;3;CPU STATISTICS OK : user=3.99% system=1.40% iowait=65.07% idle=29.54%
[1364336127] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1364336215] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364336235] SERVICE ALERT: semana-test.sofia.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 13.15, 11.14, 8.16
[1364336266] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=0 ms
[1364336324] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364336345] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1364336386] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 5.20, 9.56, 11.43
[1364336518] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1364336558] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1364336564] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=2.40% system=1.50% iowait=91.81% idle=4.30%
[1364336666] SERVICE ALERT: semana-test.sofia.elex.be;LOAD;OK;SOFT;2;OK - load average: 4.81, 7.81, 7.95
[1364336766] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364336767] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364336775] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;HARD;3;P1.PL OK - overall status=green
[1364336836] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1364336845] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364336846] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364336879] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 5.23, 5.27, 8.42
[1364336995] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;2;CPU IOWAIT CRITICAL : user=6.00% system=2.50% iowait=86.00% idle=5.50%
[1364337015] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1364337048] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1364337145] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;OK;SOFT;2;P1.PL OK - Storages : 'datastore1'(free)=470115.00 MB (99.86%)
[1364337145] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Network usage;OK;SOFT;3;P1.PL OK - net usage=2980.00 KBps
[1364337154] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1364337234] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1364337235] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;OK;SOFT;2;P1.PL OK - cpu usage=1.20 %
[1364337297] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364337305] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364337328] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=153 ms
[1364337332] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1364337434] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;HARD;3;CPU IOWAIT WARNING : user=3.09% system=1.60% iowait=73.78% idle=21.54%
[1364337449] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1364337484] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1364337576] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (47.8377065111759)
[1364337676] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364337685] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364337746] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364337757] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1364337777] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364337777] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364337784] SERVICE ALERT: gambit-test.colo.elex.be;PING;OK;HARD;3;PING OK - Packet loss = 70%, RTA = 0.22 ms
[1364337825] SERVICE ALERT: gambit-test.colo.elex.be;SSH;OK;HARD;3;SSH OK - dropbear_0.52 (protocol 2.0)
[1364337881] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364337924] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;HARD;3;CPU STATISTICS OK : user=1.40% system=1.20% iowait=47.95% idle=49.45%
[1364337976] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1364338066] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Memory usage;OK;SOFT;2;P1.PL OK - mem usage=2.86 %
[1364338116] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1364338137] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=0 ms
[1364338138] SERVICE ALERT: bifur-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364338218] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364338218] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;SOFT;2;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364338234] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364338254] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=101 ms
[1364338286] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1364338305] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1364338319] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1364338325] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1364338409] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1364338431] SERVICE ALERT: inkless-test.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:inkless-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/inkless-test.elex.be in copy
[1364338431] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1364338516] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1364338534] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1364338547] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=113 ms
[1364338569] SERVICE ALERT: bifur-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=6 ms
[1364338624] SERVICE ALERT: gambit-test.colo.elex.be;ESXi Runtime status;OK;SOFT;2;P1.PL OK - overall status=green
[1364338656] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364338656] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_CPU;CRITICAL;HARD;3;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364338677] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364338684] SERVICE ALERT: gambit-test.colo.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1364338686] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=151 ms
[1364338709] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1364338756] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1364338756] SERVICE ALERT: gambit-test.colo.elex.be;SSH;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1364338775] SERVICE ALERT: gambit-test.colo.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.23 ms