[1358204402] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358204402] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (53.2312925170068)
[1358204444] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358204455] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=43 ms
[1358204545] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358204697] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358204742] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358204748] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=19 ms
[1358205033] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358205052] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358205078] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358205342] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358205375] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358205432] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1358205786] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=141 ms
[1358205924] SERVICE ALERT: postgres-stage.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358205936] SERVICE ALERT: lechuck-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358206077] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358206205] SERVICE ALERT: postgres-stage.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1358206217] SERVICE ALERT: lechuck-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1358206304] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=105 ms
[1358206368] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=108 ms
[1358206392] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358206427] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358206442] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358206593] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=99 ms
[1358206683] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358206722] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358206722] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=87 ms
[1358206735] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358206858] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358206884] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=87 ms
[1358207022] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358207035] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358207074] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=14 ms
[1358207147] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358207234] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358207293] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358207332] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358207497] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1358207538] SERVICE ALERT: pintail3-test.erfurt.elex.be;CFENGINE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358207582] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358207851] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=95 ms
[1358207863] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358207903] SERVICE ALERT: pintail3-test.erfurt.elex.be;CFENGINE;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1358208011] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358208101] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358208149] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=99 ms
[1358208201] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358208220] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358208301] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358208399] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358208449] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=97 ms
[1358208494] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358208512] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358208659] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1358208751] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358208772] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358208799] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=64 ms
[1358208864] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358208951] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=94 ms
[1358209049] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358209072] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358209155] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358209249] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1358209351] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1358209399] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=217 ms
[1358209503] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=13 ms
[1358209551] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1358209561] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358209600] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1358209691] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358209844] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358209859] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=142 ms
[1358209859] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358209864] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358210104] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358210149] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=196 ms
[1358210152] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1358210209] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358210302] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358210402] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358210443] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358210443] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1358210502] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358210549] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358210601] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1358210753] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358210799] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358210843] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358210891] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358211033] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358211049] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358211092] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358211195] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358211383] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358211402] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1358211451] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1358211481] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358211551] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358211653] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=108 ms
[1358211759] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358211763] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358211899] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358211914] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1358211950] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358212074] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1358212111] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358212193] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358212282] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=66 ms
[1358212304] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358212370] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (53.2312925170068)
[1358212409] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358212510] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358212544] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358212574] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358212599] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1358212653] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 33.50, 22.36, 14.06
[1358212699] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1358212801] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358212841] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358212892] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=99 ms
[1358212950] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 17.52, 23.24, 17.09
[1358212979] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358213132] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1358213152] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358213189] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358213241] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 16.40, 18.76, 16.76
[1358213242] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1358213272] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358213451] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=74 ms
[1358213484] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358213569] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1358213593] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1358213741] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1358213754] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1358213780] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1358213882] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358213941] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.56, 7.97, 12.28
[1358214044] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358214185] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358214263] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1358214333] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358214441] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358214449] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358214479] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358214582] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358214613] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358214731] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358214739] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358214769] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1358214873] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1358214954] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=83 ms
[1358215082] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358215092] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1358215136] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358215216] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1358215282] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1358215433] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358215442] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358215446] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=99 ms
[1358215528] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=91 ms
[1358215605] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358215657] SERVICE ALERT: dino-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=260 ms
[1358215732] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358215742] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358215793] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358215826] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1358215903] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358215933] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358215956] SERVICE ALERT: dino-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1358216083] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1358216084] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358216119] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1358216153] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=41 ms
[1358216222] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=109 ms
[1358216346] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358216382] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358216383] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=60 ms
[1358216446] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=13 ms
[1358216452] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1358216476] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1358216486] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358216514] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1358216642] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=96 ms
[1358216733] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=13 ms
[1358216742] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358216775] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358216934] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=95 ms
[1358217082] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358217096] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358217287] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=62 ms
[1358217372] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358217385] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1358217435] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358217572] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358217616] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358217636] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1358217677] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=60 ms
[1358217723] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358217857] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358217905] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1358217986] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=104 ms
[1358218022] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358218029] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1358218073] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358218147] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1358218227] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358218279] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358218363] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1358218372] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358218554] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358218616] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358218713] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358218714] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358218726] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358218934] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358219025] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358219072] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358219237] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358219335] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1358219359] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358219502] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358219532] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1358219685] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358219734] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358219782] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358219782] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.00% system=0.20% iowait=99.80% idle=0.00%
[1358219802] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358219882] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358219915] SERVICE ALERT: sda-uat.erfurt.elex.be;CACHE_MGR;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358219982] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1358220043] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358220045] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1358220085] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358220095] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358220102] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.20% iowait=0.00% idle=99.80%
[1358220114] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1358220207] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358220253] SERVICE ALERT: sda-uat.erfurt.elex.be;CACHE_MGR;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args '/usr/share/wafermap/bin/cache_mgr'
[1358220276] SERVICE ALERT: onlinedlog-test.erfurt.elex.be;NRPE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358220314] SERVICE ALERT: atlan.erfurt.elex.be;NRPE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358220314] SERVICE ALERT: dlog1-test.erfurt.elex.be;AT;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358220314] SERVICE ALERT: atlan.erfurt.elex.be;CFENGINE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358220317] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1358220373] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1358220383] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (53.2312925170068)
[1358220416] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358220426] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1358220532] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1358220577] SERVICE ALERT: dlog3-test.erfurt.elex.be;CPU_STATS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358220585] SERVICE ALERT: onlinedlog-test.erfurt.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1358220615] SERVICE ALERT: atlan.erfurt.elex.be;NRPE;OK;SOFT;2;PROCS OK: 4 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1358220615] SERVICE ALERT: atlan.erfurt.elex.be;CFENGINE;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1358220623] SERVICE ALERT: dlog1-test.erfurt.elex.be;AT;OK;SOFT;2;PROCS OK: 1 process with UID = 1 (daemon), args '/usr/sbin/atd'
[1358220667] SERVICE ALERT: te-vnc-test.erfurt.elex.be;LOAD;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358220683] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358220686] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1358220729] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=64 ms
[1358220805] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;WARNING;SOFT;1;WARNING - InFlightCount=10
[1358220872] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358220877] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;WARNING;SOFT;1;WARNING - QueueSize=31
[1358220902] SERVICE ALERT: dlog3-test.erfurt.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.00% idle=100.00%
[1358220904] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358220992] SERVICE ALERT: te-vnc-test.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.00, 0.00, 0.00
[1358220999] SERVICE ALERT: pintail-fb-test.erfurt.elex.be;NRPE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358221012] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358221055] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1358221106] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1358221106] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1358221124] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358221195] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358221197] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358221225] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358221316] SERVICE ALERT: pintail-fb-test.erfurt.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1358221375] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358221475] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358221536] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358221555] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358221582] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358221663] SERVICE ALERT: sda-test.erfurt.elex.be;CRON;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358221735] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358221768] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1358221825] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358221846] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358221872] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358221953] SERVICE ALERT: sda-test.erfurt.elex.be;CRON;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cron'
[1358222032] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358222062] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1358222074] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358222186] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1358222282] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358222420] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358222438] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358222464] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=92 ms
[1358222476] SERVICE ALERT: atlan.erfurt.elex.be;CFENGINE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358222544] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358222574] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358222605] SERVICE ALERT: dori-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358222712] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1358222727] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1358222755] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=76 ms
[1358222772] SERVICE ALERT: atlan.erfurt.elex.be;CFENGINE;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1358222897] SERVICE ALERT: dori-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=7 ms
[1358222932] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358223022] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=58 ms
[1358223053] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=71 ms
[1358223063] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358223222] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358223247] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358223372] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=121 ms
[1358223413] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358223512] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1358223535] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1358223638] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358223712] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358223722] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1358223825] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1358223862] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358223862] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1358223877] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358223932] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=122 ms
[1358223985] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358224062] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1358224072] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358224102] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=92 ms
[1358224107] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358224167] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358224212] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358224222] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1358224275] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358224352] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=73 ms
[1358224362] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=84 ms
[1358224402] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358224453] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358224457] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1358224503] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358224638] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358224654] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1358224802] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1358224803] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358224882] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1358225093] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358225152] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358225172] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358225199] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=44 ms
[1358225338] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358225355] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1358225443] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1358225444] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358225465] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=90 ms
[1358225485] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1358225494] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=23 ms
[1358225633] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358225647] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1358225792] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=136 ms
[1358225794] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=90 ms
[1358225875] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1358225996] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358226042] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358226092] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358226112] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=99 ms
[1358226146] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1358226205] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358226362] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358226392] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=75 ms
[1358226422] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358226504] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=65 ms
[1358226543] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358226544] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358226722] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358226753] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=91 ms
[1358226868] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358227013] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358227074] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358227152] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358227342] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358227475] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1358227652] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358227665] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1358227816] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1358227858] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1358227984] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358228064] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358228073] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1358228131] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358228192] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358228207] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358228325] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (54.7862001943635)
[1358228394] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358228412] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (53.2312925170068)
[1358228464] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358228542] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358228577] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1358228737] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358228755] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358228772] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1358228866] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=104 ms
[1358228892] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358229034] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358229052] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=94 ms
[1358229092] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358229122] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1358229126] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1358229166] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=94 ms
[1358229187] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358229267] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358229385] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1358229399] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1358229441] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358229538] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358229579] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358229621] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358229635] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=42 ms
[1358229837] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358229925] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358229958] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1358230115] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358230249] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358230400] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=82 ms
[1358230565] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358230568] SERVICE ALERT: sda-test.sensors.elex.be;XRAS;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358230579] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358230581] SERVICE ALERT: postprocess-test.sensors.elex.be;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358230585] SERVICE ALERT: murray-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 9.18 ms
[1358230596] HOST ALERT: dlog1-test.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358230615] HOST ALERT: dlog1-test.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.64 ms
[1358230640] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358230696] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1358230855] SERVICE ALERT: murray-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.59 ms
[1358230855] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_4.3p2 Debian-9etch3 (protocol 2.0)
[1358230857] SERVICE ALERT: sda-test.sensors.elex.be;XRAS;OK;SOFT;2;TCP OK - 0.009 second response time on port 6933
[1358230857] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;OK;SOFT;2;TCP OK - 0.010 second response time on port 6445
[1358230858] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358230859] SERVICE ALERT: postprocess-test.sensors.elex.be;HTTP;OK;SOFT;2;HTTP OK - HTTP/1.1 302 Found - 0.036 second response time
[1358230931] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358230946] SERVICE ALERT: dlog1-test.sensors.elex.be;DLD;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 2193 (dlog), args '/usr/bin/dld'
[1358231045] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1358231145] SERVICE ALERT: hma-lp.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 49.98 ms
[1358231148] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1358231149] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358231155] HOST ALERT: hma-lp.sofia.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 45.87 ms
[1358231220] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=80 ms
[1358231236] SERVICE ALERT: dlog1-test.sensors.elex.be;DLD;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 2193 (dlog), args '/usr/bin/dld'
[1358231241] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358231337] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358231445] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=60 ms
[1358231498] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358231535] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358231650] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358231735] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358231791] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358231825] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1358231926] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1358231945] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=74 ms
[1358232039] SERVICE ALERT: ftp-test.sofia.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 14.5% (148856 kB) free
[1358232178] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1358232235] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=24 ms
[1358232276] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358232327] SERVICE ALERT: ftp-test.sofia.elex.be;MEMORY;OK;SOFT;2;Memory OK - 69.9% (719688 kB) free
[1358232388] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358232437] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1358232560] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1358232580] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358232677] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1358232785] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358232849] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1358232869] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358232891] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358233005] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358233025] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1358233133] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358233227] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358233253] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358233277] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358233314] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1358233377] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358233426] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358233456] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358233484] SERVICE ALERT: sda-test.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 10.40 ms
[1358233527] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1358233549] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358233573] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1358233616] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1358233676] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358233719] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1358233719] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (28.3% change >= 20.0% threshold)
[1358233719] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (35.3% change >= 20.0% threshold)
[1358233719] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (26.4% change >= 20.0% threshold)
[1358233719] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (30.8% change >= 20.0% threshold)
[1358233719] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1358233719] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (41.6% change >= 20.0% threshold)
[1358233734] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1358233795] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358233872] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=76 ms
[1358233963] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1358234024] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358234053] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358234215] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358234361] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=98 ms
[1358234383] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358234496] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358234511] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;OK;SOFT;2;TCP OK - 0.009 second response time on port 6933
[1358234515] SERVICE ALERT: sda-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.30 ms
[1358234631] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1358234653] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358234684] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358234726] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358234743] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=133 ms
[1358234923] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=130 ms
[1358235025] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358235077] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358235094] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1358235094] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1358235131] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1358235131] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (28.3% change >= 20.0% threshold)
[1358235131] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1358235131] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1358235131] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1358235131] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1358235248] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358235349] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1358235383] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 58.46 ms
[1358235398] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1358235398] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (41.7% change >= 20.0% threshold)
[1358235433] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358235509] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358235553] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358235588] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1358235649] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358235749] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1358235787] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358235793] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=11 ms
[1358235867] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358235885] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1358235885] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (39.6% change >= 20.0% threshold)
[1358235946] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358236005] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1358236065] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358236083] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358236116] SERVICE ALERT: postprocess-test.sensors.elex.be;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358236117] HOST ALERT: mona.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358236243] HOST ALERT: mona.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.59 ms
[1358236263] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1358236334] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1358236385] SERVICE ALERT: postprocess-test.sensors.elex.be;HTTP;OK;SOFT;2;HTTP OK - HTTP/1.1 302 Found - 0.037 second response time
[1358236415] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1358236514] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1358236559] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1358236560] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358236560] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1358236560] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358236560] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358236561] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358236561] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358236561] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358236561] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358236564] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1358236584] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1358236585] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1358236585] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1358236586] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1358236599] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 25.10, 8.33, 2.95
[1358236600] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1358236606] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1358236606] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358236606] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1358236615] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1358236615] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1358236619] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1358236619] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358236620] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358236680] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=92 ms
[1358236695] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1358236777] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1358236804] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358236824] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 61.38 ms
[1358236858] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - No LoginModules configured for karaf
[1358236875] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358236875] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358236875] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358236875] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358236896] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358236897] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358236897] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1358236899] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236906] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236906] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236906] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236917] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236917] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 4.69, 5.90, 3.61
[1358236917] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236917] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236917] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236920] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236920] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - Connection refused
[1358236935] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1358236935] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358236936] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358236977] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1358236978] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358236984] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358236985] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1358236995] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358236996] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358236996] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358236996] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358236996] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1358237018] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358237018] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358237018] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358237018] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358237024] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1358237055] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358237067] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1358237084] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358237085] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1358237085] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358237085] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358237087] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358237087] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1358237087] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1358237087] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1358237088] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1358237095] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358237095] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358237098] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358237099] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358237108] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1358237115] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 73%, RTA = 68.53 ms
[1358237134] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;1;Connection refused
[1358237135] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358237136] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358237136] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1358237136] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1358237136] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358237147] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1358237148] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358237149] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1358237155] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1358237159] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358237166] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358237166] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358237166] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1358237167] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358237168] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358237168] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358237168] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358237168] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358237184] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358237184] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358237184] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358237186] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358237187] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1358237188] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1358237194] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358237195] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358237195] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358237195] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358237196] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358237196] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358237196] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358237197] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1358237197] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358237197] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1358237198] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358237198] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1358237199] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1358237207] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1358237207] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1358237207] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1358237217] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;3;OK - load average: 0.04, 2.20, 2.63
[1358237217] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic.confirm
[1358237219] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1358237219] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358237220] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1358237225] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1358237225] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem
[1358237236] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1358237237] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358237237] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358237277] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358237279] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358237280] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358237281] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1358237296] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358237296] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358237296] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358237298] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358237298] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1358237317] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358237320] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358237320] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358237320] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358237320] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358237320] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1358237357] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358237386] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358237387] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1358237387] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358237387] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358237395] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358237395] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1358237395] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1358237395] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1358237395] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1358237395] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358237395] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358237406] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358237407] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1358237416] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1358237438] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;SOFT;2;Connection refused
[1358237438] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358237439] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358237439] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1358237445] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1358237445] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358237455] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1358237456] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358237456] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1358237464] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358237474] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358237474] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358237475] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=senodia
[1358237476] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358237477] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358237477] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358237477] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358237477] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358237485] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358237495] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358237495] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358237497] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358237500] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358237500] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1358237500] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358237500] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;SOFT;2;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1358237577] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358237578] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358237579] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1358237586] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358237586] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358237594] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358237595] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358237596] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1358237615] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358237619] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358237619] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358237620] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358237620] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358237620] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1358237658] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358237679] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358237680] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1358237680] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358237680] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358237696] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358237697] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan
[1358237697] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd.confirm
[1358237697] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=vishay
[1358237698] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1358237698] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358237699] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358237702] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358237717] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1358237738] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;CRITICAL;HARD;3;Connection refused
[1358237739] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358237739] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358237740] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl.confirm
[1358237746] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1358237747] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358237755] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech.confirm
[1358237756] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1358237756] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358237758] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358237765] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358237774] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358237779] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358237779] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358237780] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358237780] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358237796] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358237797] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358237798] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358237798] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358237799] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358237804] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358237804] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1358237804] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358237805] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;CRITICAL;HARD;3;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=umc
[1358237834] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358237917] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1358238045] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238046] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238046] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238047] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238056] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238056] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238056] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238056] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238056] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238057] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238057] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238065] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238066] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358238068] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358238077] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238098] SERVICE ALERT: ewaf-test.colo.elex.be;ACTIVEMQ;OK;HARD;3;TCP OK - 0.000 second response time on port 61616
[1358238098] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238099] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238104] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238107] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238107] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238115] SERVICE ALERT: ewaf-test.colo.elex.be;pactech.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238115] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238115] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238118] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238128] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358238136] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238136] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238137] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238137] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238147] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238154] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238155] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238155] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238158] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238158] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238158] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238158] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238188] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238188] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238188] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238194] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238230] SERVICE ALERT: ewaf-test.colo.elex.be;umc_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238288] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 2.45, 5.30, 3.88
[1358238358] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1358238395] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem1.confirm
[1358238396] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.unhandled
[1358238396] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac.confirm
[1358238396] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358238418] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358238477] HOST ALERT: stella.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358238477] HOST ALERT: tanja.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358238494] HOST ALERT: stella.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.85 ms
[1358238524] HOST ALERT: tanja.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.60 ms
[1358238586] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 1.82, 8.16, 6.12
[1358238617] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238617] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238619] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238619] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238619] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238624] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238624] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238625] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238625] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238625] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238626] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238636] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238636] SERVICE ALERT: ewaf-test.colo.elex.be;senodia_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238636] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238646] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238647] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238647] SERVICE ALERT: ewaf-test.colo.elex.be;mpd.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238647] SERVICE ALERT: ewaf-test.colo.elex.be;vishay_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238647] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238648] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238649] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238649] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238649] SERVICE ALERT: ewaf-test.colo.elex.be;carsem_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238655] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238655] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238657] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238658] SERVICE ALERT: ewaf-test.colo.elex.be;ispl.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238664] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238664] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238664] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238664] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238688] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238688] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238688] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238688] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238688] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;HARD;3;OK - QueueSize=0
[1358238689] SERVICE ALERT: ewaf-test.colo.elex.be;carsem1.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358238689] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.unhandled_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358238689] SERVICE ALERT: ewaf-test.colo.elex.be;optopac.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358238690] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358238706] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1358238710] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358238712] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358238725] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_InFlightCount;OK;HARD;3;OK - InFlightCount=0
[1358238859] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=133 ms
[1358238877] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 0.01, 3.03, 4.44
[1358238977] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358238997] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1358239004] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1358239067] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1358239154] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=153 ms
[1358239186] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358239226] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 0.00, 0.94, 3.04
[1358239264] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358239288] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=82 ms
[1358239357] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1358239446] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1358239477] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358239555] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=62 ms
[1358239654] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1358239666] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=87 ms
[1358239690] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358239755] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358239769] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358239794] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358239907] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1358239944] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 85%, RTA = 73.05 ms
[1358239957] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1358239986] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358240006] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358240060] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358240121] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=108 ms
[1358240166] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1358240325] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358240355] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358240407] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358240416] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358240455] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358240466] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=81 ms
[1358240497] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1358240617] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358240655] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1358240746] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358240764] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1358240768] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358240795] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;OK;SOFT;3;OK: Connection to cfengine master OK
[1358240824] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=89 ms
[1358240894] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358240915] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1358240995] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358241059] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1358241079] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=99 ms
[1358241084] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358241084] SERVICE ALERT: esb-a-test.kuching.elex.be;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.60.112.35)
[1358241097] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;1;Server down!
[1358241124] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 66.78 ms
[1358241125] SERVICE ALERT: esb-a-test.kuching.elex.be;NRPE;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241144] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=7 ms
[1358241164] SERVICE ALERT: esb-a-test.kuching.elex.be;MEMORY;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241326] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1358241334] SERVICE ALERT: ftp-test.kuching.elex.be;CFENGINE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241334] SERVICE ALERT: ftp-test.kuching.elex.be;DISK;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241336] SERVICE ALERT: esb-a-test.kuching.elex.be;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1358241344] SERVICE ALERT: ftp-test.kuching.elex.be;SWAP;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241344] SERVICE ALERT: ftp-test.kuching.elex.be;AT;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241347] SERVICE ALERT: ftp-test.kuching.elex.be;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.60.112.53)
[1358241347] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;1;Server down!
[1358241354] SERVICE ALERT: ftp-test.kuching.elex.be;NRPE;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241374] SERVICE ALERT: ftp-test.kuching.elex.be;IOSTAT;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241387] SERVICE ALERT: ftp-test.kuching.elex.be;FTP;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1358241395] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;2;Server down!
[1358241404] SERVICE ALERT: ftp-test.kuching.elex.be;CRON;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358241405] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=87 ms
[1358241406] SERVICE ALERT: esb-a-test.kuching.elex.be;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.60.112.35)
[1358241408] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358241408] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=44 ms
[1358241414] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;2;Server down!
[1358241426] SERVICE ALERT: ftp-test.kuching.elex.be;FTP_FETCH;UNKNOWN;HARD;1;**ePN /usr/lib/nagios/plugins/check_ftp.pl: "ERROR: Cannot conect to 10.60.112.53".
[1358241434] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 63.08 ms
[1358241469] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358241574] SERVICE ALERT: ftp-test.kuching.elex.be;SSH;CRITICAL;HARD;1;No route to host
[1358241586] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1358241666] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358241674] SERVICE ALERT: ftp-test.kuching.elex.be;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.60.112.53)
[1358241737] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=7 ms
[1358241739] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;3;Server down!
[1358241754] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;3;Server down!
[1358241794] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358241966] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1358241984] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358242075] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;4;Server down!
[1358242085] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;4;Server down!
[1358242156] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=131 ms
[1358242270] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=99 ms
[1358242348] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358242397] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;5;Server down!
[1358242404] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;5;Server down!
[1358242539] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358242554] SERVICE ALERT: christina-test.colo.elex.be;DISK;OK;HARD;3;OK - all disk
[1358242594] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=74 ms
[1358242678] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1358242744] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;6;Server down!
[1358242744] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;6;Server down!
[1358242859] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1358242916] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1358242919] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358242969] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=92 ms
[1358242994] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1358243075] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;7;Server down!
[1358243076] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;7;Server down!
[1358243180] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358243197] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358243238] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358243321] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi_DATASTORE1;CRITICAL;SOFT;1;P1.PL CRITICAL - Error connecting to server at 'https://10.49.32.52:443/sdk/webService': Connection refused
[1358243358] HOST ALERT: ironhide-test.erfurt.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1358243361] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=82 ms
[1358243384] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358243405] SERVICE ALERT: ironhide-test.erfurt.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1358243420] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;8;Server down!
[1358243424] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;8;Server down!
[1358243436] SERVICE ALERT: ironhide-test.erfurt.elex.be;ESXi Runtime status;CRITICAL;HARD;1;P1.PL UNKNOWN - plugin timed out (timeout 30s)
[1358243456] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1358243456] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (47.9% change >= 20.0% threshold)
[1358243456] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1358243456] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (32.0% change >= 20.0% threshold)
[1358243456] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (51.3% change >= 20.0% threshold)
[1358243488] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358243498] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;2;Server down!
[1358243508] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;2;Server down!
[1358243520] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358243588] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1358243590] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;3;Server down!
[1358243599] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358243600] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;3;Server down!
[1358243662] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358243662] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1358243684] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1358243684] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;4;Server down!
[1358243728] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;4;Server down!
[1358243824] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358243845] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358243848] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;5;Server down!
[1358243859] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1358243859] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1358243859] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (47.9% change >= 20.0% threshold)
[1358243859] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1358243859] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (32.0% change >= 20.0% threshold)
[1358243893] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1358243893] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (46.5% change >= 20.0% threshold)
[1358243901] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358243911] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358243931] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;2;Server down!
[1358243951] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 58.83 ms
[1358243961] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;2;Server down!
[1358244003] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358244041] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;3;Server down!
[1358244041] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358244061] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1358244071] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;3;Server down!
[1358244091] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358244113] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 58.23 ms
[1358244154] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=60 ms
[1358244172] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;4;Server down!
[1358244181] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=67 ms
[1358244191] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358244221] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;4;Server down!
[1358244222] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358244311] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1358244343] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358244352] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;5;Server down!
[1358244355] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=60 ms
[1358244391] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358244431] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;5;Server down!
[1358244583] SERVICE ALERT: betty.printers.erfurt.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358244594] SERVICE ALERT: ishy.printers.erfurt.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358244597] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=69 ms
[1358244601] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358244612] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358244613] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;6;Server down!
[1358244614] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=111 ms
[1358244642] SERVICE ALERT: esb-b-test.kuching.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 240.67 ms
[1358244643] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358244664] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;1;Server down!
[1358244681] SERVICE ALERT: esb-b-test.kuching.elex.be;MEMORY;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358244716] SERVICE ALERT: esb-b-test.kuching.elex.be;CFENGINE;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358244721] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;6;Server down!
[1358244742] SERVICE ALERT: esb-b-test.kuching.elex.be;DISK;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358244762] SERVICE ALERT: esb-b-test.kuching.elex.be;CPU_STATS;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358244763] SERVICE ALERT: esb-b-test.kuching.elex.be;INODES;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358244771] SERVICE ALERT: esb-b-test.kuching.elex.be;NRPE;CRITICAL;HARD;1;Connection refused or timed out
[1358244772] SERVICE ALERT: esb-b-test.kuching.elex.be;SYSLOG;CRITICAL;HARD;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358244873] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358244874] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1358244914] SERVICE ALERT: betty.printers.erfurt.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1358244932] SERVICE ALERT: ishy.printers.erfurt.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1358244941] SERVICE ALERT: esb-b-test.kuching.elex.be;SSH;CRITICAL;HARD;1;No route to host
[1358244953] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358244955] SERVICE ALERT: esb-b-test.kuching.elex.be;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.60.112.34)
[1358244963] SERVICE ALERT: irmina.printers.erfurt.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358244971] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;2;Server down!
[1358244972] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;7;Server down!
[1358244972] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1358244981] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358245001] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1358245002] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358245061] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358245061] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1358245094] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;7;Server down!
[1358245191] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1358245234] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358245243] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358245261] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1358245321] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1358245321] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1358245323] SERVICE ALERT: irmina.printers.erfurt.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1358245353] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;3;Server down!
[1358245356] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;8;Server down!
[1358245356] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1358245362] SERVICE ALERT: bumblebee-test.erfurt.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358245381] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;1;Server down!
[1358245384] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1358245402] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;2;Server down!
[1358245424] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358245424] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358245481] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;8;Server down!
[1358245491] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1358245511] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;3;Server down!
[1358245561] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;4;Server down!
[1358245603] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1358245631] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1358245641] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;5;Server down!
[1358245666] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1358245723] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;6;Server down!
[1358245752] SERVICE ALERT: bumblebee-test.erfurt.elex.be;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1358245763] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;4;Server down!
[1358245765] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;7;Server down!
[1358245766] HOST ALERT: ftp-test.kuching.elex.be;DOWN;SOFT;9;Server down!
[1358245767] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1358245771] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1358245781] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358245792] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1358245826] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358245865] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;8;Server down!
[1358245885] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=99 ms
[1358245901] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;SOFT;9;Server down!
[1358245924] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;SOFT;9;Server down!
[1358245941] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1358245961] HOST ALERT: bumblebee-test.erfurt.elex.be;DOWN;HARD;10;Server down!
[1358246045] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1358246061] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1358246131] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=112 ms
[1358246131] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;5;Server down!
[1358246133] HOST ALERT: ftp-test.kuching.elex.be;DOWN;HARD;10;Server down!
[1358246134] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1358246164] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1358246182] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1358246231] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358246233] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358246271] HOST ALERT: esb-a-test.kuching.elex.be;DOWN;HARD;10;Server down!
[1358246341] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1358246393] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1358246401] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 70.07 ms
[1358246441] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1358246454] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=113 ms
[1358246466] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;6;Server down!
[1358246471] HOST ALERT: ishy.printers.erfurt.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1358246496] HOST ALERT: irmina.printers.erfurt.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1358246541] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=104 ms
[1358246561] HOST ALERT: betty.printers.erfurt.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1358246626] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358246685] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358246701] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 60.39 ms
[1358246773] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;7;Server down!
[1358246802] HOST ALERT: bumblebee-test.erfurt.elex.be;UP;HARD;1;Server alive
[1358246806] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1358246831] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1358246883] SERVICE ALERT: bumblebee-test.erfurt.elex.be;SSH;OK;HARD;1;SSH OK - dropbear_0.52 (protocol 2.0)
[1358246891] HOST ALERT: betty.printers.erfurt.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 36.83 ms
[1358246911] SERVICE ALERT: betty.printers.erfurt.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 36.46 ms
[1358246914] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358246943] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1358246974] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 73%, RTA = 59.22 ms
[1358246975] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358247054] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1358247072] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;8;Server down!
[1358247151] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1358247226] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358247264] SERVICE ALERT: ishy.printers.erfurt.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 36.71 ms
[1358247271] HOST ALERT: ishy.printers.erfurt.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 39.69 ms
[1358247394] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1358247442] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;SOFT;9;Server down!
[1358247611] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1358247663] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1358247771] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358247791] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358247792] HOST ALERT: esb-b-test.kuching.elex.be;DOWN;HARD;10;Server down!
[1358247981] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358247992] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=73 ms
[1358248001] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358248091] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358248114] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358248296] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358248322] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358248414] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1358248601] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (55.9037900874636)
[1358248651] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1358248685] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358248727] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1358248786] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358248866] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=124 ms
[1358249021] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358249078] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358249156] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358249261] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 58.15 ms
[1358249351] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358249366] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358249385] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358249447] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1358249544] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 76%, RTA = 60.08 ms
[1358249634] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358249661] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=97 ms
[1358249675] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358249701] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358249855] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358249923] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358249946] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1358249965] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=56 ms
[1358249991] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358250127] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358250147] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358250162] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358250206] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1358250276] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1358250414] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358250443] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358250643] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1358250661] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1358250841] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358251051] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=122 ms
[1358251066] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358251193] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1358251203] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358251246] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358251345] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358251361] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358251484] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1358251485] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=8 ms
[1358251535] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358251636] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=60 ms
[1358251800] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1358251800] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (30.3% change >= 20.0% threshold)
[1358251800] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1358251800] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (36.7% change >= 20.0% threshold)
[1358251800] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (53.5% change >= 20.0% threshold)
[1358251859] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1358251938] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358251985] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1358252081] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1358252088] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358252088] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_SEARCH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358252097] SERVICE ALERT: tanja.printers.sensors.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 15.00 ms
[1358252145] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358252151] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358252195] SERVICE ALERT: tanja.printers.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.91 ms
[1358252198] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;OK;SOFT;2;FTP OK - 0.021 second response time on port 21 [220 Welcome to the Melexis Data Exchange FTP service.]
[1358252198] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_SEARCH;OK;SOFT;2;TCP OK - 0.009 second response time on port 6911
[1358252205] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=59 ms
[1358252288] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358252293] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1358252293] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (31.4% change >= 20.0% threshold)
[1358252293] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1358252293] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (36.7% change >= 20.0% threshold)
[1358252293] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (53.5% change >= 20.0% threshold)
[1358252305] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358252306] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358252326] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1358252358] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1358252446] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1358252446] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1358252528] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358252628] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358252675] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1358252687] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358252808] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358252845] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=91 ms
[1358252858] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1358252876] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358252956] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358252966] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=106 ms
[1358253099] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358253118] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358253145] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1358253206] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358253215] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=94 ms
[1358253216] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358253362] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1358253381] SERVICE ALERT: irmina.printers.erfurt.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 43.17 ms
[1358253386] HOST ALERT: irmina.printers.erfurt.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 36.62 ms
[1358253449] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358253475] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358253487] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=50 ms
[1358253490] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1358253649] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=95 ms
[1358253736] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1358253756] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358254005] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358254009] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358254027] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1358254107] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358254180] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1358254239] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358254299] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358254355] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358254377] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=92 ms
[1358254399] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=88 ms
[1358254586] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1358254587] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358254698] HOST ALERT: mondra.printers.erfurt.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358254724] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358254748] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=82 ms
[1358254764] HOST ALERT: mondra.printers.erfurt.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 36.12 ms
[1358254905] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358254905] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358255016] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358255240] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358255245] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358255255] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1358255255] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=99 ms
[1358255454] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358255464] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358255534] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358255536] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358255549] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1358255606] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358255754] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1358255836] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=106 ms
[1358255880] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358255897] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1358255957] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1358255957] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1358255957] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (39.1% change >= 20.0% threshold)
[1358255957] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1358255957] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (39.3% change >= 20.0% threshold)
[1358255957] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (49.3% change >= 20.0% threshold)
[1358256189] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1358256222] SERVICE ALERT: riverbed.colo.elex.be;RIVERBED_STATUS;WARNING;SOFT;1;Steelhead 2050 (2050L): Healthy, optimisation service: runninghealth_status=Healthy, service_status=running,
[1358256311] SERVICE ALERT: riverbed.colo.elex.be;RIVERBED_STATUS;OK;SOFT;2;Steelhead 2050 (2050L): Healthy, optimisation service: runninghealth_status=Healthy, service_status=running,
[1358256390] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=72 ms
[1358256509] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358256564] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1358256609] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1358256633] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1358256751] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358256760] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358256822] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358256871] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1358256979] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1358256985] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1358257030] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358257049] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358257086] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358257121] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1358257121] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358257269] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1358257319] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358257320] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1358257340] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1358257371] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1358257401] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=63 ms
[1358257402] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358257609] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358257643] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1358257661] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=12 ms
[1358257749] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1358257751] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358257865] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1358257865] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1358257865] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1358257865] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1358257865] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (48.2% change >= 20.0% threshold)
[1358257865] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (44.4% change >= 20.0% threshold)
[1358257908] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=53 ms
[1358257916] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1358257988] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358257997] SERVICE ALERT: dss-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358258016] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_SEARCH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358258017] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358258017] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1358258019] SERVICE ALERT: pintail-fb-test.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358258028] SERVICE ALERT: sda32-test.sensors.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358258125] SERVICE ALERT: dss-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1358258129] SERVICE ALERT: pintail-fb-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.34 ms
[1358258136] SERVICE ALERT: sda32-test.sensors.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 11.34 ms
[1358258149] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_SEARCH;OK;SOFT;2;TCP OK - 0.009 second response time on port 6911
[1358258150] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1358258178] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358258190] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1358258346] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=96 ms
[1358258368] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1358258375] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358258418] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=61 ms
[1358258449] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358258537] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1358258568] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358258578] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1358258610] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1358258655] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358258818] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1358258828] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358258875] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=90 ms
[1358259067] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358259128] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1358259167] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358259175] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358259357] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1358259406] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358259458] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358259469] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1358259507] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358259580] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358259635] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358259709] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358259751] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358259807] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358259829] HOST ALERT: hma-lp.sofia.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1358259856] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358259925] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358259975] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358260045] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=76 ms
[1358260045] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1358260076] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358260207] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358260250] SERVICE ALERT: sda-test.sofia.elex.be;IOSTAT;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358260250] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1358260258] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1358260285] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1358260315] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1358260325] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;2;(Host Check Timed Out)
[1358260355] SERVICE ALERT: sda-test.sofia.elex.be;CFENGINE;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260358] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=88 ms
[1358260358] SERVICE ALERT: sda-test.sofia.elex.be;POSTFIX;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260377] SERVICE ALERT: sda-test.sofia.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1358260386] SERVICE ALERT: sda-test.sofia.elex.be;CACHE_MGR;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260408] SERVICE ALERT: sda-test.sofia.elex.be;CPU_STATS;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260408] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1358260435] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1358260456] SERVICE ALERT: sda-test.sofia.elex.be;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1358260465] SERVICE ALERT: sda-test.sofia.elex.be;AT;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260487] SERVICE ALERT: sda-test.sofia.elex.be;LOAD;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260496] SERVICE ALERT: sda-test.sofia.elex.be;CRON;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260527] SERVICE ALERT: sda-test.sofia.elex.be;NRPE;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260527] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358260568] SERVICE ALERT: sda-test.sofia.elex.be;MEMORY;CRITICAL;HARD;1;(Service Check Timed Out)
[1358260638] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1358260646] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;3;(Host Check Timed Out)
[1358260665] SERVICE FLAPPING ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1358260665] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (29.7% change >= 20.0% threshold)
[1358260665] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (26.8% change >= 20.0% threshold)
[1358260665] SERVICE FLAPPING ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1358260665] SERVICE FLAPPING ALERT: zonda-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (41.2% change >= 20.0% threshold)
[1358260681] SERVICE ALERT: sda-test.sofia.elex.be;WMAP_FILE;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1358260749] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;2;(Host Check Timed Out)
[1358260838] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358260869] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;3;(Host Check Timed Out)
[1358260922] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358261009] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;4;(Host Check Timed Out)
[1358261038] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1358261092] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1358261174] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1358261179] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;5;(Host Check Timed Out)
[1358261221] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358261251] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358261330] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358261340] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358261389] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;6;(Host Check Timed Out)
[1358261401] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1358261510] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358261533] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358261548] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358261618] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1358261639] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;7;(Host Check Timed Out)
[1358261699] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358261768] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358261811] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=87 ms
[1358261853] SERVICE ALERT: sda-test.sofia.elex.be;WMAP_SEARCH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1358261914] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358261931] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358262009] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;8;(Host Check Timed Out)
[1358262042] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=72 ms
[1358262042] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (41.8% change >= 20.0% threshold)
[1358262088] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1358262101] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358262119] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1358262211] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1358262229] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358262348] HOST ALERT: sda-test.sofia.elex.be;DOWN;SOFT;9;(Host Check Timed Out)
[1358262378] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358262401] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1358262403] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1358262451] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358262474] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358262511] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358262681] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358262689] HOST ALERT: sda-test.sofia.elex.be;DOWN;HARD;10;(Host Check Timed Out)
[1358262758] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358262772] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358262774] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358262789] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358262890] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358262998] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=36 ms
[1358263061] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1358263078] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358263081] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358263190] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1358263352] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1358263374] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=76 ms
[1358263415] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358263489] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1358263509] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358263600] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358263668] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1358263732] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1358263753] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1358263833] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1358263869] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1358263872] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358263910] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358263918] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 80%, RTA = 66.09 ms
[1358264058] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 58.54 ms
[1358264098] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358264120] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=96 ms
[1358264140] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1358264232] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358264259] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 60.73 ms
[1358264260] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor
[1358264269] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.korea.confirm
[1358264269] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358264270] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.pactechconti.continental.regensburg
[1358264270] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358264271] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=leister
[1358264271] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358264271] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358264280] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem
[1358264281] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=pactech
[1358264289] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.philippines.confirm
[1358264290] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.inkless
[1358264292] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=ispl
[1358264292] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=partner
[1358264293] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.customerdeliveries.splitted
[1358264321] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.trigger
[1358264321] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=xfab.in
[1358264330] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=atlantic
[1358264330] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=exceptions.aggregate
[1358264330] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 19.12, 5.46, 1.87
[1358264330] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=mpd
[1358264331] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=shipped
[1358264359] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=carsem2.confirm
[1358264359] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=inkless.trigger
[1358264360] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=optopac
[1358264360] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=unisem.confirm
[1358264371] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=amkor.taiwan.confirm
[1358264372] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=com.melexis.viiper.fillcollectionplan.x02sawing
[1358264372] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=lothold
[1358264373] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=postprocessing.mapmerge.erfurt.in
[1358264373] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;CRITICAL;SOFT;1;CRITICAL - org.apache.activemq:BrokerName=default,Type=Queue,Destination=viiper.materialtransactions.splitted
[1358264412] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1358264440] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1358264551] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358264574] SERVICE ALERT: ewaf-test.colo.elex.be;amkor_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264588] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.korea.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264588] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264589] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.pactechconti.continental.regensburg_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264590] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264590] SERVICE ALERT: ewaf-test.colo.elex.be;leister_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264591] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264591] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264599] SERVICE ALERT: ewaf-test.colo.elex.be;unisem_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264600] SERVICE ALERT: ewaf-test.colo.elex.be;pactech_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264608] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.philippines.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264609] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.inkless_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264611] SERVICE ALERT: ewaf-test.colo.elex.be;ispl_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264611] SERVICE ALERT: ewaf-test.colo.elex.be;partner_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264611] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.customerdeliveries.splitted_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264640] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.trigger_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264640] SERVICE ALERT: ewaf-test.colo.elex.be;xfab.in_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264641] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 0.54, 4.80, 3.03
[1358264641] SERVICE ALERT: ewaf-test.colo.elex.be;atlantic_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264642] SERVICE ALERT: ewaf-test.colo.elex.be;exceptions.aggregate_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264642] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264649] SERVICE ALERT: ewaf-test.colo.elex.be;shipped_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264658] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 67.87 ms
[1358264670] SERVICE ALERT: ewaf-test.colo.elex.be;carsem2.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264671] SERVICE ALERT: ewaf-test.colo.elex.be;inkless.trigger_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264671] SERVICE ALERT: ewaf-test.colo.elex.be;optopac_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264672] SERVICE ALERT: ewaf-test.colo.elex.be;unisem.confirm_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264681] SERVICE ALERT: ewaf-test.colo.elex.be;amkor.taiwan.confirm_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264688] SERVICE ALERT: ewaf-test.colo.elex.be;com.melexis.viiper.fillcollectionplan.x02sawing_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264688] SERVICE ALERT: ewaf-test.colo.elex.be;lothold_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264688] SERVICE ALERT: ewaf-test.colo.elex.be;postprocessing.mapmerge.erfurt.in_Queue_InFlightCount;OK;SOFT;2;OK - InFlightCount=0
[1358264689] SERVICE ALERT: ewaf-test.colo.elex.be;viiper.materialtransactions.splitted_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358264749] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=82 ms
[1358264781] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358264782] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358264809] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358264920] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358264949] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;3;OK - load average: 0.01, 1.72, 2.17
[1358264969] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 58.36 ms
[1358265091] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358265092] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1358265112] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358265172] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358265231] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358265258] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358265263] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;3;PING OK - Packet loss = 58%, RTA = 61.24 ms
[1358265471] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1358265497] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1358265540] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358265569] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358265658] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=102 ms
[1358265767] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358265882] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1358265887] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358265979] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=98 ms
[1358266007] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358266010] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=45 ms
[1358266127] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1358266380] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358266409] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358266557] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358266732] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358266831] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358266919] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358266962] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358267090] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358267182] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358267183] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358267217] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=44 ms
[1358267269] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1358267310] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 85%, RTA = 70.40 ms
[1358267312] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358267334] SERVICE ALERT: ewaf-test.colo.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 19.1% (777896 kB) free
[1358267380] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;1;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358267540] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358267567] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358267624] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1358267698] SERVICE ALERT: ewaf-test.colo.elex.be;MEMORY;WARNING;SOFT;2;Memory WARNING - 18.6% (754188 kB) free
[1358267738] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;2;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358267757] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358267914] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1358267949] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;WARNING;SOFT;1;WARNING - QueueSize=15
[1358267972] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1358268038] SERVICE ALERT: ewaf-test.colo.elex.be;MEMORY;WARNING;HARD;3;Memory WARNING - 18.3% (741736 kB) free
[1358268071] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358268080] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1358268110] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358268227] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358268262] SERVICE ALERT: ewaf-test.colo.elex.be;mpd_Queue_QueueSize;OK;SOFT;2;OK - QueueSize=0
[1358268350] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358268352] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1358268440] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358268701] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358268707] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358268997] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=48 ms
[1358269040] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1358269363] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=4 ms
[1358269410] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1358269410] SERVICE FLAPPING ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1358269412] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358269539] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358269565] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358269837] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358269899] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358269929] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358270097] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358270157] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=116 ms
[1358270167] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=106 ms
[1358270411] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358270471] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=55 ms
[1358270482] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358270502] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358270651] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358270717] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358270731] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1358270821] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1358270864] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358270969] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358270969] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1358271062] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=89 ms
[1358271141] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1358271187] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1358271351] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358271389] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358271504] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1358271668] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358271903] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1358271979] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (55.2721088435374)
[1358272027] SERVICE ALERT: openlink-test.sofia.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=14.80% system=4.40% iowait=74.20% idle=%
[1358272240] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358272348] SERVICE ALERT: openlink-test.sofia.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.60% idle=%
[1358272428] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1358272523] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1358272562] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=82 ms
[1358272805] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358272807] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1358272899] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=83 ms
[1358272939] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1358273187] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358273195] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1358273225] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358273505] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=83 ms
[1358273516] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=78 ms
[1358273603] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358273826] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1358273836] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=49 ms
[1358273842] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358273858] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358273872] SERVICE ALERT: pintail1-test.erfurt.elex.be;POSTFIX;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358273917] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358273933] SERVICE ALERT: pintail1-test.erfurt.elex.be;CRON;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358274158] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1358274176] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358274186] SERVICE ALERT: pintail1-test.erfurt.elex.be;POSTFIX;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args '/postfix/master'
[1358274254] SERVICE ALERT: pintail1-test.erfurt.elex.be;CRON;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cron'
[1358274304] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358274333] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358274417] SERVICE ALERT: ftp-test.erfurt.elex.be;DISK;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358274484] SERVICE ALERT: dlog1-test.erfurt.elex.be;CRON;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358274517] SERVICE ALERT: ftp-test.erfurt.elex.be;SWAP;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358274539] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358274544] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358274585] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358274596] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=97 ms
[1358274622] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=85 ms
[1358274655] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358274673] SERVICE ALERT: dlog1-test.erfurt.elex.be;MEMORY;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358274734] SERVICE ALERT: ftp-test.erfurt.elex.be;DISK;OK;SOFT;2;OK - all disk
[1358274762] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 70.26 ms
[1358274804] SERVICE ALERT: dlog1-test.erfurt.elex.be;CRON;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cron'
[1358274817] SERVICE ALERT: dlog1-test.erfurt.elex.be;SWAP;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358274852] SERVICE ALERT: ftp-test.erfurt.elex.be;SWAP;OK;SOFT;2;SWAP OK - 100% free (1004 MB out of 1004 MB)
[1358274862] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358274913] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=74 ms
[1358274944] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=71 ms
[1358274945] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=42 ms
[1358274978] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1358274978] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=9 ms
[1358274995] SERVICE ALERT: dlog1-test.erfurt.elex.be;MEMORY;OK;SOFT;2;Memory OK - 91.1% (3704648 kB) free
[1358275134] SERVICE ALERT: dlog1-test.erfurt.elex.be;SWAP;OK;SOFT;2;SWAP OK - 100% free (1004 MB out of 1004 MB)
[1358275183] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1358275234] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=44 ms
[1358275264] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1358275539] SERVICE ALERT: dlog1-test.erfurt.elex.be;NRPE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358275614] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1358275742] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358275845] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1358275857] SERVICE ALERT: dlog1-test.erfurt.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1358275895] SERVICE ALERT: pintail2-test.erfurt.elex.be;CFENGINE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358275994] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358276055] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358276074] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1358276094] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358276106] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1358276136] SERVICE ALERT: pintail1-test.erfurt.elex.be;CRON;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358276157] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1358276196] SERVICE ALERT: pintail2-test.erfurt.elex.be;CFENGINE;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1358276319] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1358276345] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1358276445] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358276484] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358276487] SERVICE ALERT: pintail1-test.erfurt.elex.be;CRON;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'cron'
[1358276506] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1358276519] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=129 ms
[1358276610] SERVICE ALERT: pintail3-test.erfurt.elex.be;LOAD;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1358276695] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358276707] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1358276777] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358276817] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1358276906] SERVICE ALERT: pintail3-test.erfurt.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.00, 0.00, 0.00
[1358276987] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358277004] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358277105] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358277168] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358277190] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1358277299] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1358277405] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358277454] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1358277464] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1358277487] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358277708] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1358277750] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1358277754] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1358277844] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358277866] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358278008] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1358278008] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1358278044] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=78 ms
[1358278049] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1358278114] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1358278134] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358278160] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1358278187] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=46 ms
[1358278237] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358278354] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=88 ms
[1358278399] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1358278414] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=82 ms
[1358278494] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=9 ms
[1358278536] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1358278555] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358278646] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1358278704] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1358278846] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358279226] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358279357] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1358279426] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358279464] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1358279524] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358279556] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358279655] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358279660] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358279719] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358279758] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358279774] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1358279855] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358279998] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;OK;SOFT;2;TCP OK - 0.009 second response time on port 6445
[1358280069] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=109 ms
[1358280074] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1358280111] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1358280187] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1358280268] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1358280368] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1358280482] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=108 ms
[1358280559] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=83 ms
[1358280659] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1358280769] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1358280820] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1358280820] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1358280851] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=88 ms
[1358281012] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1358281012] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1358281127] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358281172] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358281208] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1358281368] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1358281417] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1358281470] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358281660] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358281703] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358281710] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1358281848] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1358282001] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1358282019] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1358282147] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1358282269] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1358282309] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=14 ms
[1358282417] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1358282439] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1358282527] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358282562] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358282701] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1358282702] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1358282710] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1358282787] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1358282821] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1358282998] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1358283007] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=73 ms
[1358283013] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1358283087] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358283121] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1358283191] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358283270] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1358283312] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=92 ms
[1358283498] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1358283502] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1358283601] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358283611] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358283661] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=72 ms
[1358283817] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=2.90% system=7.49% iowait=72.63% idle=16.98%
[1358283835] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1358283904] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1358283965] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 24.70, 19.66, 11.65
[1358284111] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358284111] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.30% system=0.30% iowait=49.95% idle=49.45%
[1358284196] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358284262] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 24.64, 22.83, 15.20
[1358284364] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358284403] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1358284433] SERVICE ALERT: diehard-test.erfurt.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358284481] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1358284487] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1358284552] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 26.57, 25.11, 18.28
[1358284611] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358284674] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1358284694] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1358284717] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1358284733] SERVICE ALERT: diehard-test.erfurt.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=6 ms
[1358284778] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1358284831] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=72 ms
[1358284911] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1358284964] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358285014] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358285046] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358285123] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1358285261] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358285322] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=112 ms
[1358285342] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1358285367] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358285416] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=146 ms
[1358285484] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1358285551] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358285608] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358285612] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 7.04, 15.21, 19.16
[1358285613] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1358285627] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358285633] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=37 ms
[1358285665] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1358285837] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358285901] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1358285913] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1358285919] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1358285956] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1358286125] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1358286133] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358286306] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1358286475] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358286605] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1358286611] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1358286611] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1358286614] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1358286771] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=80 ms
[1358286904] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1358286911] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1358286972] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=100 ms
[1358287078] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1358287254] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358287323] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1358287323] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1358287326] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=73 ms
[1358287393] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1358287477] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1358287553] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1358287614] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1358287614] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1358287617] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1358287687] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1358287746] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1358287833] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=134 ms
[1358287845] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1358287973] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (54.7862001943635)
[1358288044] HOST ALERT: stella.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358288044] HOST ALERT: tanja.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1358288044] SERVICE ALERT: lechuck-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358288057] SERVICE ALERT: stella.printers.sensors.elex.be;PING;WARNING;HARD;1;PING WARNING - Packet loss = 90%, RTA = 9.85 ms
[1358288063] HOST ALERT: elaine-test.sensors.elex.be;DOWN;SOFT;1;(Host Check Timed Out)
[1358288064] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1358288078] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1358288103] HOST ALERT: elaine-test.sensors.elex.be;UP;SOFT;2;Server alive
[1358288127] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=106 ms
[1358288143] HOST ALERT: tanja.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 10.62 ms
[1358288145] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 10.18, 8.67, 11.45
[1358288173] HOST ALERT: stella.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.79 ms
[1358288203] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1358288329] SERVICE ALERT: lechuck-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - dropbear_0.52 (protocol 2.0)
[1358288354] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;OK;SOFT;2;TCP OK - 0.009 second response time on port 6913
[1358288375] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1358288387] SERVICE ALERT: stella.printers.sensors.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 9.62 ms
[1358288425] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=106 ms
[1358288495] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1358288630] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=69 ms
[1358288667] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1358288676] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358288678] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1358288776] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1358288923] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1358288973] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1358288974] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=103 ms
[1358289033] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1358289133] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1358289265] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1358289323] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1358289677] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1358289937] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1358289973] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=111 ms
[1358289974] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1358290033] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1358290079] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358290193] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1358290234] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1358290265] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1358290266] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1358290270] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1358290270] SERVICE FLAPPING ALERT: jaguar-test.kiev.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1358290325] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1358290375] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1358290488] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1358290554] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1358290554] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms