[1360191649] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1360191681] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360191788] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1360191829] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360191959] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360191967] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360191999] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1360192141] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360192159] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1360192249] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360192271] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1360192277] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360192277] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 18.64, 12.71, 8.30
[1360192533] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360192569] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360192589] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360192590] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 17.91, 14.38, 10.27
[1360192617] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=8.48% system=2.00% iowait=84.73% idle=4.79%
[1360192651] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360192851] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=84 ms
[1360192970] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;3;OK - load average: 8.30, 11.23, 10.39
[1360192995] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=43.01% system=9.18% iowait=34.23% idle=13.47%
[1360193036] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1360193227] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1360193387] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360193410] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1360193585] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=204 ms
[1360193601] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360193696] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1360193720] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=125 ms
[1360193778] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1360193899] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360194008] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1360194039] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1360194087] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360194347] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1360194408] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1360194666] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360194686] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360194785] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360194970] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360194998] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360195037] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360195098] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360195248] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360195287] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360195355] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360195415] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=88 ms
[1360195505] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360195567] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360195588] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=75 ms
[1360195668] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360195897] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360196139] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360196165] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=104 ms
[1360196215] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1360196248] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360196447] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360196582] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=169 ms
[1360196597] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360196624] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360196746] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360196784] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360196846] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=90 ms
[1360196922] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=1 ms
[1360196947] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1360197063] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360197065] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360197104] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=10 ms
[1360197164] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360197347] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360197374] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360197396] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1360197542] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360197665] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=124 ms
[1360197673] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360197693] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360197853] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1360197977] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1360197977] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360198148] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360198162] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;OK;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360198169] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=77 ms
[1360198354] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1360198433] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360198446] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360198464] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360198546] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1360198749] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360198764] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1360198813] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360199076] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=44 ms
[1360199123] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360199167] SERVICE ALERT: lenor-test.sofia.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 18.26, 14.69, 9.42
[1360199212] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360199295] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1360199475] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360199479] SERVICE ALERT: lenor-test.sofia.elex.be;LOAD;OK;SOFT;2;OK - load average: 11.58, 13.29, 10.23
[1360199500] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360199533] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360199577] SERVICE ALERT: fili-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360199604] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360199765] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 18.36, 15.93, 11.50
[1360199786] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360199795] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=55 ms
[1360199816] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360199826] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1360199858] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1360199896] SERVICE ALERT: fili-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=7 ms
[1360199904] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360200108] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 17.58, 19.07, 14.58
[1360200173] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360200195] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360200235] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1360200283] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360200373] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360200422] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1360200485] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 15.42, 17.34, 15.38
[1360200552] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=94 ms
[1360200572] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360200596] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1360200676] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360200734] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=73 ms
[1360200878] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360200892] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360200974] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=17 ms
[1360201052] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360201195] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360201234] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 6.45, 8.53, 12.05
[1360201263] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1360201302] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=60 ms
[1360201387] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360201572] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360201686] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360201706] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360201724] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360201794] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1360202003] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360202035] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=74 ms
[1360202054] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=109 ms
[1360202114] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=60 ms
[1360202322] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360202353] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1360202423] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1360202424] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=85 ms
[1360202454] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360202642] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=95 ms
[1360202724] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360202772] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1360202805] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=32 ms
[1360202957] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=58 ms
[1360202977] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360203044] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=70 ms
[1360203072] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360203089] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360203293] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360203332] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1360203355] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1360203382] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360203553] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1360203556] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=99 ms
[1360203927] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360203930] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360204076] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360204147] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1360204169] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360204251] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=101 ms
[1360204272] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360204392] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360204466] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360204487] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360204588] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360204626] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1360204657] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (55.9037900874636)
[1360204776] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360204859] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360204999] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1360205150] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1360205166] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360205168] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360205422] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=128 ms
[1360205466] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=76 ms
[1360205480] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360205688] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360205738] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=79 ms
[1360205746] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360205786] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1360206007] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360206058] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360206066] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360206090] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360206228] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360206382] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1360206408] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360206538] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360206540] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=8 ms
[1360206759] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1360206850] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360206919] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360207041] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=173 ms
[1360207136] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360207146] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360207187] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1360207237] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360207239] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360207452] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=13 ms
[1360207544] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1360207563] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360207605] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360207636] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360207652] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360207852] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360207854] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360207885] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=46 ms
[1360207905] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360207974] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360208042] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360208172] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1360208176] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360208362] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360208375] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360208445] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360208485] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1360208637] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360208693] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360208732] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360208736] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360208762] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360209045] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360209052] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360209066] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=91 ms
[1360209243] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1360209358] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360209358] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1360209365] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1360209384] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360209384] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360209562] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360209663] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 19.18, 9.03, 7.37
[1360209738] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360209765] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360209772] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360209942] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360209982] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 7.87, 12.93, 10.19
[1360210062] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360210072] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360210086] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360210092] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360210127] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360210255] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1360210376] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=79 ms
[1360210394] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360210452] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360210463] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360210472] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360210574] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=71 ms
[1360210757] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1360210777] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1360210782] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=51 ms
[1360211160] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360211168] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=81 ms
[1360211390] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=14 ms
[1360211572] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360211769] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360211888] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1360211918] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1360212079] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=73 ms
[1360212208] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1360212252] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360212288] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=76 ms
[1360212291] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=90 ms
[1360212401] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1360212570] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360212580] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1360212580] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1360212600] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=80 ms
[1360212611] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1360212813] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;1;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360212921] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360212928] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=44 ms
[1360213130] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;SOFT;2;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360213298] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=92 ms
[1360213338] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360213451] SERVICE ALERT: bamboo-xpeqt-test.colo.elex.be;CFAGENT;WARNING;HARD;3;WARNING: cfengine:bamboo-xpeqt-test: Can't stat /var/lib/cfengine2/masterfiles/files/apache/bamboo-xpeqt-test.colo.elex.be in copy
[1360213543] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360213551] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1360213642] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=176 ms
[1360213684] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360213759] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360213898] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360213963] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360213999] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1360214099] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360214250] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=47 ms
[1360214301] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360214309] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360214555] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360214600] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=94 ms
[1360214717] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1360214719] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=69 ms
[1360214725] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=14 ms
[1360214941] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360215018] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=43 ms
[1360215079] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=33 ms
[1360215139] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360215140] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360215362] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1360215495] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360215627] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (54.7862001943635)
[1360215688] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1360215875] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360215918] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360216018] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360216188] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360216237] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1360216239] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360216340] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360216429] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360216445] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360216550] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360216555] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360216558] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=37 ms
[1360216567] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360216745] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360216765] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=11 ms
[1360216815] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360216885] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360216916] SERVICE ALERT: hma-lp.sofia.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 50.09 ms
[1360216925] HOST ALERT: hma-lp.sofia.elex.be;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 49.37 ms
[1360217138] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1360217146] HOST ALERT: easyship.printers.tess.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360217181] SERVICE ALERT: easyship.printers.tess.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1360217300] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360217315] HOST ALERT: easyship.printers.tess.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1360217315] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360217449] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360217476] HOST ALERT: easyship.printers.tess.elex.be;UP;SOFT;3;PING OK - Packet loss = 88%, RTA = 9.48 ms
[1360217528] SERVICE ALERT: easyship.printers.tess.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.55 ms
[1360217565] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360217617] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1360217628] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1360217636] SERVICE ALERT: mirona.printers.erfurt.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360217761] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1360217875] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=85 ms
[1360217927] SERVICE ALERT: mirona.printers.erfurt.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 37.28 ms
[1360217931] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360217945] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1360218006] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360218199] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1360218259] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1360218367] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360218386] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360218637] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360218682] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360218702] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=40 ms
[1360218950] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360218951] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=96 ms
[1360218995] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=83 ms
[1360219265] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360219267] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=80 ms
[1360219330] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360219639] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=115 ms
[1360219639] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1360219649] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;2;Connection refused
[1360219859] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360219946] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360219962] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;HARD;3;Connection refused
[1360220007] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360220017] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=74 ms
[1360220119] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1360220125] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 68.07 ms
[1360220145] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360220167] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.5150631681244)
[1360220255] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=115 ms
[1360220325] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=44 ms
[1360220388] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=106 ms
[1360220465] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1360220497] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360220578] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360220635] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=85 ms
[1360220645] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1360220776] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360220788] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1360220820] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360220905] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360221009] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=11 ms
[1360221025] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360221135] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1360221148] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=92 ms
[1360221165] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1360221337] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360221387] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360221461] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360221562] SERVICE ALERT: cmdb-test.elex.be;SYSLOG;CRITICAL;SOFT;1;Connection refused or timed out
[1360221565] HOST ALERT: cmdb-test.elex.be;DOWN;SOFT;1;Server down!
[1360221605] SERVICE ALERT: cmdb-test.elex.be;CFENGINE;CRITICAL;HARD;1;Connection refused or timed out
[1360221635] SERVICE ALERT: cmdb-test.elex.be;NRPE;CRITICAL;HARD;1;Connection refused or timed out
[1360221650] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1360221650] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=62 ms
[1360221757] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1360221764] SERVICE ALERT: cmdb-test.elex.be;SSH;CRITICAL;HARD;1;No route to host
[1360221828] SERVICE ALERT: cmdb-test.elex.be;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.32.241.146)
[1360221879] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360221904] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360221964] HOST ALERT: cmdb-test.elex.be;DOWN;SOFT;2;Server down!
[1360221978] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360222036] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360222094] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1360222106] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1360222138] SERVICE ALERT: cmdb-test.elex.be;SSH;OK;HARD;1;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze2 (protocol 2.0)
[1360222144] HOST ALERT: cmdb-test.elex.be;UP;SOFT;3;Server alive
[1360222198] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360222204] SERVICE ALERT: cmdb-test.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1360222218] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360222280] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360222296] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360222446] SERVICE ALERT: cmdb-test.elex.be;NRPE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1360222465] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360222478] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=71 ms
[1360222525] SERVICE ALERT: cmdb-test.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360222529] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360222544] HOST ALERT: cmdb-test.elex.be;DOWN;SOFT;1;Server down!
[1360222604] HOST ALERT: cmdb-test.elex.be;DOWN;SOFT;2;Server down!
[1360222604] HOST FLAPPING ALERT: cmdb-test.elex.be;STARTED; Host appears to have started flapping (20.0% change > 20.0% threshold)
[1360222778] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360222794] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=89 ms
[1360222835] SERVICE ALERT: cmdb-test.elex.be;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1360222849] HOST ALERT: cmdb-test.elex.be;DOWN;SOFT;3;Server down!
[1360222975] HOST ALERT: cmdb-test.elex.be;DOWN;SOFT;4;Server down!
[1360223086] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360223107] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1360223124] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=50 ms
[1360223164] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=85 ms
[1360223179] HOST ALERT: cmdb-test.elex.be;UP;SOFT;5;Server alive
[1360223216] SERVICE ALERT: cmdb-test.elex.be;SSH;OK;HARD;1;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze2 (protocol 2.0)
[1360223216] SERVICE FLAPPING ALERT: cmdb-test.elex.be;SSH;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1360223357] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360223400] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360223475] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360223525] SERVICE ALERT: cmdb-test.elex.be;NRPE;OK;HARD;1;PROCS OK: 4 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360223546] SERVICE ALERT: cmdb-test.elex.be;CFENGINE;OK;HARD;1;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1360223667] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;2;Connection refused
[1360223667] HOST FLAPPING ALERT: cmdb-test.elex.be;STOPPED; Host appears to have stopped flapping (4.9% change < 5.0% threshold)
[1360223837] SERVICE ALERT: cmdb-test.elex.be;SYSLOG;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'syslogd'
[1360223874] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=213 ms
[1360223979] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 61.22 ms
[1360224055] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360224247] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360224268] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1360224295] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 61%, RTA = 61.08 ms
[1360224384] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=57 ms
[1360224386] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360224396] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=95 ms
[1360224475] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 60.81 ms
[1360224475] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360224550] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360224568] SERVICE ALERT: azhira.printers.erfurt.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 36.75 ms
[1360224574] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360224585] HOST ALERT: azhira.printers.erfurt.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360224598] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360224624] HOST ALERT: azhira.printers.erfurt.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1360224651] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360224704] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1360224716] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1360224786] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360224865] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360224884] SERVICE ALERT: azhira.printers.erfurt.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1360224886] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;CRITICAL;SOFT;1;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 42% CRITICAL,Waste Toner Box is at CRITICAL Level
[1360224916] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1360224920] HOST ALERT: azhira.printers.erfurt.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1360224954] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360224959] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360224968] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=128 ms
[1360225004] HOST ALERT: azhira.printers.erfurt.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1360225008] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=31 ms
[1360225098] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360225204] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1360225234] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;CRITICAL;SOFT;2;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 42% CRITICAL,Waste Toner Box is at CRITICAL Level
[1360225317] HOST ALERT: azhira.printers.erfurt.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1360225351] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360225370] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360225383] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1360225492] HOST ALERT: azhira.printers.erfurt.elex.be;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 36.83 ms
[1360225525] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=65 ms
[1360225556] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1360225715] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;CRITICAL;HARD;3;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 42% CRITICAL,Waste Toner Box is at CRITICAL Level
[1360225821] SERVICE ALERT: azhira.printers.erfurt.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 36.90 ms
[1360225862] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360225883] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=111 ms
[1360225893] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360225951] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360225975] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=60 ms
[1360226071] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 67.28 ms
[1360226128] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360226192] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1360226233] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360226264] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360226311] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1360226381] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1360226411] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 70%, RTA = 74.67 ms
[1360226474] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360226632] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=44 ms
[1360226754] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360226822] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (55.2721088435374)
[1360226974] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1360227043] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360227106] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1360227222] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 63.36 ms
[1360227404] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360227412] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360227571] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 62.48 ms
[1360227613] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360227721] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=108 ms
[1360227772] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=61 ms
[1360227793] SERVICE ALERT: tipa.printers.erfurt.elex.be;PRINTER_STATUS;OK;HARD;3;OK,C-KIT DCC2725 M is at 88% OK,C-KIT DCC2725 Y is at 57% OK,C-KIT DCC2740 K is at 42% OK,Waste Toner Box is FULL
[1360227833] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360227903] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360227936] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360227938] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360227963] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360228062] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=40 ms
[1360228121] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=14 ms
[1360228172] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360228236] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360228271] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360228274] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360228494] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1360228719] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360228738] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360228961] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360229043] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360229191] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360229194] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1360229215] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360229403] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360229457] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360229537] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360229543] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1360229544] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=75 ms
[1360229610] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360229732] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1360229850] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1360229927] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1360229937] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360230121] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360230181] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1360230183] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360230198] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360230198] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;1;Connection refused
[1360230258] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;SOFT;1;Connection refused
[1360230260] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=87 ms
[1360230261] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360230319] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360230359] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360230373] SERVICE ALERT: dss-test.colo.elex.be;DSS_MPD;OK;HARD;3;PROCS OK: 1 process with UID = 2737 (syncmgr), args '/mnt/ed/ct/lnx/rh/53/dm/current/syncinc/bin/mpd'
[1360230387] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;CRITICAL;SOFT;1;CRITICAL: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentB - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentC - Agent Execution Failure, Aborting Mapping/IMPOR.
[1360230407] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1360230441] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360230477] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360230500] SERVICE ALERT: dss-test.colo.elex.be;HTTP_DSS;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 490 bytes in 0.134 seconds
[1360230509] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360230519] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360230519] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360230522] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;CRITICAL: SQL query failed!
[1360230528] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360230537] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;2;Connection refused
[1360230592] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;SOFT;2;Connection refused
[1360230598] SERVICE ALERT: dss-test.colo.elex.be;DSS_POSTGRES;OK;HARD;3;PROCS OK: 10 processes with UID = 2737 (syncmgr), args 'postgres'
[1360230607] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360230658] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=21 ms
[1360230659] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360230703] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360230707] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360230717] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;CRITICAL;SOFT;2;CRITICAL: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentB - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentC - Agent Execution Failure, Aborting Mapping/IMPOR.
[1360230730] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1360230768] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=29 ms
[1360230801] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360230833] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360230838] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360230839] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360230840] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;2;CRITICAL: SQL query failed!
[1360230852] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'Apache'
[1360230853] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;HARD;3;Connection refused
[1360230877] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360230909] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;CRITICAL;HARD;3;Connection refused
[1360230929] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDCRM'
[1360230967] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1360230987] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360231012] SERVICE FLAPPING ALERT: cmdb-test.elex.be;SSH;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360231018] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360231028] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=64 ms
[1360231049] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;CRITICAL;HARD;3;CRITICAL: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentB - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentC - Agent Execution Failure, Aborting Mapping/IMPOR.
[1360231051] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=90 ms
[1360231069] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 59998 (applmgr), args 'tnslsnr'
[1360231119] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360231148] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360231161] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 87%, RTA = 81.69 ms
[1360231179] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360231193] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360231198] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360231317] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360231337] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=101 ms
[1360231365] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=62 ms
[1360231399] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360231448] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360231650] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360231862] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360231927] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=70 ms
[1360232127] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1360232138] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=4.05% system=2.15% iowait=76.17% idle=17.63%
[1360232138] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 35.24, 21.59, 12.51
[1360232179] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360232242] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360232247] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 13.83, 10.01, 5.06
[1360232287] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360232448] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360232457] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;WARNING;SOFT;2;CPU IOWAIT WARNING : user=22.30% system=2.70% iowait=71.80% idle=3.20%
[1360232462] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 36.25, 30.94, 19.18
[1360232538] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360232587] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1360232626] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 7.70, 10.02, 6.89
[1360232665] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360232745] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360232815] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1360232825] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;OK;SOFT;3;CPU STATISTICS OK : user=27.61% system=2.85% iowait=52.42% idle=17.12%
[1360232835] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 29.32, 31.00, 23.31
[1360232847] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=98 ms
[1360232903] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360232912] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360232941] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360232963] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 11.85, 11.10, 8.28
[1360233005] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1360233072] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360233165] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=6 ms
[1360233241] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1360233261] SERVICE ALERT: cmdb-test.elex.be;SYSLOG;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1360233322] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1360233323] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1360233323] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1360233394] HOST ALERT: cassius.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360233396] SERVICE ALERT: murray-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360233397] SERVICE ALERT: pintail1-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360233421] HOST ALERT: cassius.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.25 ms
[1360233434] SERVICE ALERT: ewaf-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=4.76% system=0.79% iowait=94.45% idle=0.00%
[1360233453] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 7.08, 4.20, 1.80
[1360233566] SERVICE ALERT: cmdb-test.elex.be;SYSLOG;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'syslogd'
[1360233644] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=54 ms
[1360233665] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=109 ms
[1360233702] SERVICE ALERT: murray-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1360233702] SERVICE ALERT: pintail1-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1360233732] SERVICE ALERT: ewaf-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=1.58% system=0.00% iowait=6.62% idle=91.80%
[1360233749] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.86, 2.84, 1.97
[1360233769] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 1760 bytes in 0.040 seconds
[1360233954] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1360233961] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360233983] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360234274] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1360234297] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=51 ms
[1360234333] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=103 ms
[1360234371] SERVICE ALERT: gtools-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.20% system=0.60% iowait=96.60% idle=2.60%
[1360234472] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 7.58, 8.70, 9.42
[1360234575] SERVICE ALERT: bifur-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360234577] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=88 ms
[1360234591] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1360234653] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=76 ms
[1360234676] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1360234691] SERVICE ALERT: gtools-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.20% system=0.20% iowait=0.00% idle=99.60%
[1360234741] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 69.57 ms
[1360234893] SERVICE ALERT: bifur-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=5 ms
[1360234894] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360234977] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1360234984] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360235048] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 62.13 ms
[1360235064] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1360235225] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360235237] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;OK;HARD;3;OK - load average: 1.80, 3.76, 6.36
[1360235289] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360235314] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1360235485] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;OK;HARD;3;TCP OK - 0.000 second response time on port 1521
[1360235494] SERVICE ALERT: mainrouter-test.kiev.elex.be;NRPE;OK;HARD;1;PROCS OK: 4 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360235496] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 7.22, 8.91, 16.04
[1360235504] HOST ALERT: mainrouter-test.kiev.elex.be;UP;HARD;1;Server alive
[1360235520] SERVICE ALERT: mainrouter-test.kiev.elex.be;CPU_STATS;OK;HARD;1;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.00% idle=100.00%
[1360235521] SERVICE FLAPPING ALERT: mainrouter-test.kiev.elex.be;CPU_STATS;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1360235534] SERVICE ALERT: mainrouter-test.kiev.elex.be;DISK;OK;HARD;1;OK - all disk
[1360235536] SERVICE ALERT: mainrouter-test.kiev.elex.be;AT;OK;HARD;1;PROCS OK: 1 process with UID = 1 (daemon), args '/usr/sbin/atd'
[1360235536] SERVICE ALERT: mainrouter-test.kiev.elex.be;SWAP;OK;HARD;1;SWAP OK - 100% free (1004 MB out of 1004 MB)
[1360235536] SERVICE FLAPPING ALERT: mainrouter-test.kiev.elex.be;SWAP;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1360235537] SERVICE ALERT: mainrouter-test.kiev.elex.be;SSH;OK;HARD;1;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze2 (protocol 2.0)
[1360235538] SERVICE ALERT: mainrouter-test.kiev.elex.be;IOSTAT;OK;HARD;1;SDA:rsec/s=0,wsec/s=0,%util=0:SDB:rsec/s=0,wsec/s=0,%util=0
[1360235544] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1360235557] SERVICE ALERT: mainrouter-test.kiev.elex.be;CRON;OK;HARD;1;PROCS OK: 1 process with UID = 0 (root), args 'cron'
[1360235557] SERVICE FLAPPING ALERT: mainrouter-test.kiev.elex.be;CRON;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1360235606] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;2;Connection refused
[1360235615] SERVICE ALERT: mainrouter-test.kiev.elex.be;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 55.04 ms
[1360235634] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1360235645] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360235696] SERVICE ALERT: mainrouter-test.kiev.elex.be;SQUID;CRITICAL;SOFT;1;Connection refused
[1360235737] SERVICE ALERT: mainrouter-test.kiev.elex.be;CFENGINE;OK;HARD;1;PROCS OK: 1 process with UID = 0 (root), args 'cfservd'
[1360235749] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;OK;HARD;3;PROCS OK: 1 process with UID = 700 (oracle), args 'tnslsnr'
[1360235795] SERVICE ALERT: mainrouter-test.kiev.elex.be;PPTP;OK;SOFT;2;PROCS OK: 1 process with UID = 0 (root), args 'pptpd'
[1360235796] SERVICE ALERT: mainrouter-test.kiev.elex.be;MEMORY;OK;SOFT;2;Memory OK - 89.1% (914668 kB) free
[1360235796] SERVICE ALERT: mainrouter-test.kiev.elex.be;INODES;OK;SOFT;2;OK - all disk
[1360235796] SERVICE ALERT: mainrouter-test.kiev.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.00, 0.00, 0.00
[1360235806] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;OK;HARD;3;OK: no outstanding alert(s) in EM
[1360235806] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;OK;HARD;3;OK - ARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.3.0 - 64bit Production)
[1360235824] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360235825] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 66.28 ms
[1360235864] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=41 ms
[1360235876] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.15, 5.52, 12.06
[1360235921] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;HARD;3;Connection refused
[1360235956] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.5150631681244)
[1360236020] SERVICE ALERT: mainrouter-test.kiev.elex.be;SQUID;CRITICAL;SOFT;2;Connection refused
[1360236046] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1360236158] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=25 ms
[1360236167] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 67.89 ms
[1360236178] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360236292] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1360236420] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360236420] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1360236448] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1360236520] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360236549] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=114 ms
[1360236653] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360236703] SERVICE ALERT: dss-test.colo.elex.be;HTTP_DSS;CRITICAL;SOFT;1;Connection refused
[1360236768] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1360236790] SERVICE ALERT: dss-test.colo.elex.be;DSS_POSTGRES;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 2737 (syncmgr), args 'postgres'
[1360236863] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=102 ms
[1360236960] SERVICE ALERT: dss-test.colo.elex.be;DSS_MPD;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 2737 (syncmgr), args '/mnt/ed/ct/lnx/rh/53/dm/current/syncinc/bin/mpd'
[1360236971] SERVICE ALERT: galaxy-test.tess.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360237021] SERVICE ALERT: dss-test.colo.elex.be;HTTP_DSS;CRITICAL;SOFT;2;Connection refused
[1360237082] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=46 ms
[1360237110] SERVICE ALERT: dss-test.colo.elex.be;DSS_POSTGRES;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 2737 (syncmgr), args 'postgres'
[1360237181] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=64 ms
[1360237210] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 11.94, 9.17, 3.85
[1360237249] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360237268] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360237271] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360237272] SERVICE ALERT: dss-test.colo.elex.be;DSS_MPD;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 2737 (syncmgr), args '/mnt/ed/ct/lnx/rh/53/dm/current/syncinc/bin/mpd'
[1360237334] SERVICE ALERT: dss-test.colo.elex.be;HTTP_DSS;CRITICAL;HARD;3;Connection refused
[1360237424] SERVICE ALERT: dss-test.colo.elex.be;DSS_POSTGRES;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 2737 (syncmgr), args 'postgres'
[1360237461] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1360237529] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.05, 3.13, 2.71
[1360237550] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=83 ms
[1360237559] SERVICE ALERT: postprocess-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1360237559] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1360237588] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360237588] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360237591] SERVICE ALERT: dss-test.colo.elex.be;DSS_MPD;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 2737 (syncmgr), args '/mnt/ed/ct/lnx/rh/53/dm/current/syncinc/bin/mpd'
[1360237802] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360237838] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=81 ms
[1360237868] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1360237962] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360238122] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360238158] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360238182] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=80 ms
[1360238281] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=224 ms
[1360238311] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1360238410] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360238526] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=54 ms
[1360238570] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1360238608] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1360238641] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360238739] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360238858] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=13 ms
[1360238891] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360238928] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360238988] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1360239108] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360239170] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360239211] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360239245] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=36 ms
[1360239328] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360239368] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1360239428] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360239441] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 1803 bytes in 59.111 seconds
[1360239490] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360239644] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360239722] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360239871] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360240018] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1360240027] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360240127] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360240206] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360240222] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=68 ms
[1360240258] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360240326] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360240357] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360240405] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360240457] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=30 ms
[1360240537] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;OK;SOFT;2;HTTP OK HTTP/1.1 200 OK - 1803 bytes in 0.028 seconds
[1360240555] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=75 ms
[1360240558] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360240585] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360240669] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360240741] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360240836] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 78.23 ms
[1360240868] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1360240899] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1360241023] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360241149] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360241167] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360241218] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=56 ms
[1360241262] SERVICE ALERT: temp1.sofia.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360241282] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;UNKNOWN;SOFT;1;UNKOWN: error reading value
[1360241349] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360241498] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360241518] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360241546] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;OK;HARD;3;OK: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Synchronization Successful. HRAgentB - Synchronization Successful. HRAgentC - Synchronization Successful.
[1360241575] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=34 ms
[1360241595] SERVICE ALERT: temp1.sofia.elex.be;TEMPERATURE;OK;SOFT;2;OK: 16.17
[1360241595] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 28%, RTA = 63.40 ms
[1360241707] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;OK;SOFT;2;HTTP OK HTTP/1.1 200 OK - 1803 bytes in 0.038 seconds
[1360241736] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=100 ms
[1360241965] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;CRITICAL;SOFT;1;CRITICAL: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Synchronization Successful. HRAgentB - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentC - Agent Execution Failure, Aborting Mapping/IMPOR.
[1360242047] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360242048] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360242108] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;CRITICAL: SQL query failed!
[1360242108] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;1;Connection refused
[1360242116] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360242116] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360242116] HOST FLAPPING ALERT: temp1.sofia.elex.be;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1360242168] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;1;Connection refused
[1360242288] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;CRITICAL;SOFT;2;CRITICAL: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentB - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentC - Agent Execution Failure, Aborting Mapping/IMPOR.
[1360242327] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360242369] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360242370] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1360242398] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 85%, RTA = 69.56 ms
[1360242429] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;2;CRITICAL: SQL query failed!
[1360242435] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;SOFT;2;Connection refused
[1360242445] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360242490] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;2;Connection refused
[1360242518] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=16 ms
[1360242567] SERVICE ALERT: pintail-fb-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360242589] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360242589] SERVICE ALERT: sda-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360242590] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360242605] SERVICE FLAPPING ALERT: mainrouter-test.kiev.elex.be;SWAP;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1360242605] SERVICE FLAPPING ALERT: mainrouter-test.kiev.elex.be;CPU_STATS;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1360242607] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360242616] SERVICE ALERT: viiperbi-test.apps.elex.be;OIDEXPSYNC;CRITICAL;HARD;3;CRITICAL: SYNCTOLDAPMASTER41 - Synchronization Successful. HRAgentA - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentB - Agent Execution Failure, Aborting Mapping/IMPOR. HRAgentC - Agent Execution Failure, Aborting Mapping/IMPOR.
[1360242649] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360242696] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360242697] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=48 ms
[1360242706] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360242764] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;CRITICAL;HARD;3;Connection refused
[1360242764] SERVICE FLAPPING ALERT: cmdb-test.elex.be;HTTP_8080;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1360242764] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360242805] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360242902] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;HARD;3;Connection refused
[1360242949] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=76 ms
[1360242955] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360242999] SERVICE ALERT: pintail-fb-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1360243021] SERVICE ALERT: dlog1-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.1p1 Debian-5 (protocol 2.0)
[1360243021] SERVICE ALERT: sda-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze1 (protocol 2.0)
[1360243021] SERVICE ALERT: ftp-test.sensors.elex.be;FTP;OK;SOFT;2;FTP OK - 0.021 second response time on port 21 [220 Welcome to the Melexis Data Exchange FTP service.]
[1360243049] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360243051] SERVICE ALERT: dss-test.colo.elex.be;DSS_POSTGRES;OK;HARD;3;PROCS OK: 5 processes with UID = 2737 (syncmgr), args 'postgres'
[1360243090] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1360243139] SERVICE FLAPPING ALERT: mainrouter-test.kiev.elex.be;CRON;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360243167] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=46 ms
[1360243340] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 65.82 ms
[1360243361] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL CRITICAL - io read latency=284 ms
[1360243365] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360243452] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1360243522] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 42.09, 30.47, 18.55
[1360243568] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360243572] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1360243642] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;WARNING;SOFT;1;WARNING: load average: 1.70, 5.74, 3.18
[1360243676] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 14.85, 13.18, 8.18
[1360243742] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 65.55 ms
[1360243758] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1360243805] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360243883] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 35.47, 38.84, 26.55
[1360243922] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360243964] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.15, 1.91, 2.23
[1360244026] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 1.52, 7.00, 7.24
[1360244050] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (55.9037900874636)
[1360244092] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=19 ms
[1360244092] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1360244158] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=90 ms
[1360244160] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360244226] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 36.18, 32.39, 26.69
[1360244319] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360244384] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=22.95% system=2.60% iowait=74.45% idle=0.00%
[1360244388] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 4.12, 6.53, 4.22
[1360244389] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 12.08, 10.47, 8.71
[1360244524] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360244524] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=18 ms
[1360244527] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;CRITICAL;SOFT;1;(Service Check Timed Out)
[1360244700] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360244710] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=1.00% system=0.40% iowait=98.60% idle=0.00%
[1360244760] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=66.90% system=4.85% iowait=11.70% idle=16.55%
[1360244764] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 0.30, 2.02, 2.88
[1360244832] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1360244832] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1360244872] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;OK;SOFT;2;OK: Opening BINARY mode data connection for ftp_check (17520000 bytes). File send OK.
[1360244890] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360244950] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360245072] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=119 ms
[1360245080] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=19.84% idle=80.16%
[1360245164] SERVICE ALERT: mondra.printers.erfurt.elex.be;PRINTER_STATUS;OK;HARD;3;OK,C-KIT DCC2740 C is at 40% OK,C-KIT DCC2740 M is at 67% OK,C-KIT DCC2740 Y is at 55% OK,Waste Toner Box is FULL
[1360245201] SERVICE ALERT: gtools-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.40% system=0.00% iowait=99.60% idle=0.00%
[1360245253] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360245272] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360245310] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360245310] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1360245358] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360245416] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360245421] HOST FLAPPING ALERT: temp1.sofia.elex.be;STOPPED; Host appears to have stopped flapping (0.0% change < 5.0% threshold)
[1360245527] SERVICE ALERT: dss-test.colo.elex.be;HTTP_DSS;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 490 bytes in 6.176 seconds
[1360245550] SERVICE ALERT: gtools-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=33.20% idle=66.80%
[1360245590] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=23.79% system=2.70% iowait=73.51% idle=0.00%
[1360245601] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360245676] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1360245718] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360245725] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360245822] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360245929] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245929] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_SEARCH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245931] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245931] SERVICE ALERT: postgres-stage.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245939] SERVICE ALERT: dss-test.colo.elex.be;DSS_MPD;OK;HARD;3;PROCS OK: 1 process with UID = 2737 (syncmgr), args '/mnt/ed/ct/lnx/rh/53/dm/current/syncinc/bin/mpd'
[1360245963] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245963] SERVICE ALERT: dlog-pp-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245964] SERVICE ALERT: postgres-stage.sensors.elex.be;POSTGRES_TCP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245968] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360245999] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=77 ms
[1360246098] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1360246099] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=60 ms
[1360246108] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=53.12% system=5.30% iowait=29.09% idle=12.49%
[1360246120] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=27 ms
[1360246189] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=16 ms
[1360246248] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1360246259] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360246278] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360246310] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_FILE;OK;SOFT;2;TCP OK - 0.009 second response time on port 6913
[1360246310] SERVICE ALERT: sda-test.sensors.elex.be;WMAP_SEARCH;OK;SOFT;2;TCP OK - 0.011 second response time on port 6911
[1360246310] SERVICE ALERT: sda32-test.sensors.elex.be;XRAS;OK;SOFT;2;TCP OK - 0.009 second response time on port 6933
[1360246310] SERVICE ALERT: postgres-stage.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.6 (protocol 2.0)
[1360246310] SERVICE ALERT: dlog1-test.sensors.elex.be;GRIDENGINE_EXEC;OK;SOFT;2;TCP OK - 0.009 second response time on port 6445
[1360246311] SERVICE ALERT: dlog-pp-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_5.5p1 Debian-6+squeeze2 (protocol 2.0)
[1360246311] SERVICE ALERT: postgres-stage.sensors.elex.be;POSTGRES_TCP;OK;SOFT;2;TCP OK - 0.009 second response time on port 5432
[1360246311] SERVICE ALERT: contactmonitoring-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_4.3p2 Debian-9etch3 (protocol 2.0)
[1360246360] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;2;P1.PL WARNING - io read latency=46 ms
[1360246411] SERVICE ALERT: temp1.sofia.elex.be;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 67.06 ms
[1360246429] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360246437] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1360246590] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=113 ms
[1360246597] SERVICE ALERT: hma-lp.sofia.elex.be;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1360246620] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 9.54, 6.82, 8.61
[1360246620] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360246627] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360246692] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=0 ms
[1360246731] SERVICE ALERT: temp1.sofia.elex.be;PING;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 68.20 ms
[1360246800] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1360246838] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360246890] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.00% system=0.40% iowait=99.60% idle=0.00%
[1360246970] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360246990] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360247030] SERVICE ALERT: gambit-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1360247031] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms
[1360247067] SERVICE ALERT: jira4-test.colo.elex.be;JIRA_AUTH;UNKNOWN;HARD;3;UNKNOWN Error
[1360247090] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;OK;HARD;3;OK - load average: 2.23, 3.43, 6.28
[1360247200] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;OK;HARD;3;PROCS OK: 1 process with UID = 700 (oracle), args 'tnslsnr'
[1360247231] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1360247249] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=66 ms
[1360247279] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;OK;HARD;3;OK: no outstanding alert(s) in EM
[1360247290] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;OK;HARD;3;OK - ARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.3.0 - 64bit Production)
[1360247298] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=2.79% system=1.40% iowait=59.88% idle=35.93%
[1360247331] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;OK;HARD;3;TCP OK - 0.000 second response time on port 1521
[1360247347] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 4.61, 8.16, 18.73
[1360247351] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1360247367] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1360247440] SERVICE ALERT: jira4-test.colo.elex.be;JIRA_AUTH;CRITICAL;HARD;3;User jis is not authenticated!
[1360247449] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=70 ms
[1360247530] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=12 ms
[1360247597] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=68 ms
[1360247598] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1360247678] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (54.7862001943635)
[1360247701] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360247777] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360247800] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360247877] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360247908] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360247949] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1360247991] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=30 ms
[1360248068] SERVICE ALERT: borneo-test.kuching.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=140 ms
[1360248110] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360248128] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360248171] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.39, 5.25, 10.69
[1360248240] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360248299] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360248369] SERVICE ALERT: bipmelei-test.apps.elex.be;NRPE;UNKNOWN;SOFT;1;User name nrpe was not found
[1360248371] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1360248431] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360248569] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360248600] SERVICE ALERT: borneo-test.kuching.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=0 ms
[1360248658] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360248699] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=66 ms
[1360248888] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.32.241.17'
[1360248929] SERVICE ALERT: bipmelei-test.apps.elex.be;NRPE;OK;SOFT;2;OK - 3 processes running with UID = 18018 (nrpe), args /usr/sbin/nrpe
[1360248944] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360248947] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1360248958] SERVICE ALERT: pmelei-test.apps.elex.be;NRPE;UNKNOWN;SOFT;1;User name nrpe was not found
[1360248989] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=80 ms
[1360249009] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (57.555879494655)
[1360249010] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360249090] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360249167] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360249179] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 90%, RTA = 61.02 ms
[1360249298] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;CRITICAL: SQL query failed!
[1360249302] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360249363] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;1;Connection refused
[1360249435] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360249460] SERVICE ALERT: pmelei-test.apps.elex.be;NRPE;OK;SOFT;2;OK - 3 processes running with UID = 18018 (nrpe), args /usr/sbin/nrpe
[1360249462] HOST ALERT: hma-lp.sofia.elex.be;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1360249492] SERVICE ALERT: openft-test.elex.be;NRPE;UNKNOWN;SOFT;1;check_procs: User name was not found - nrpe
[1360249504] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=15 ms
[1360249543] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360249543] HOST FLAPPING ALERT: temp1.sofia.elex.be;STARTED; Host appears to have started flapping (23.8% change > 20.0% threshold)
[1360249613] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1360249682] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360249751] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 85%, RTA = 82.81 ms
[1360249772] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;2;CRITICAL: SQL query failed!
[1360249773] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360249790] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=33 ms
[1360249793] SERVICE FLAPPING ALERT: cmdb-test.elex.be;HTTP_8080;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1360249816] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;2;Connection refused
[1360249843] SERVICE ALERT: pfus2-test.colo.elex.be;HTTP;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 503 Service Temporarily Unavailable
[1360249917] HOST ALERT: hma-lp.sofia.elex.be;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1360249920] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360249922] SERVICE ALERT: openft-test.elex.be;NRPE;OK;SOFT;2;PROCS OK: 3 processes with UID = 18018 (nrpe), args '/usr/sbin/nrpe'
[1360250072] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360250074] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360250154] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360250155] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360250213] SERVICE ALERT: starcraft-test.colo.elex.be;MEMORY;WARNING;SOFT;1;Memory WARNING - 19.8% (3261096 kB) free
[1360250214] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;HARD;3;Connection refused
[1360250236] SERVICE ALERT: pfus2-test.colo.elex.be;HTTP;OK;SOFT;2;HTTP OK - HTTP/1.1 302 Found - 0.003 second response time
[1360250315] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360250454] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=52 ms
[1360250483] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=100 ms
[1360250605] SERVICE ALERT: starcraft-test.colo.elex.be;MEMORY;OK;SOFT;2;Memory OK - 28.2% (4650216 kB) free
[1360250656] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=58 ms
[1360250675] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360250697] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=47 ms
[1360250797] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360250816] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360250992] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=126 ms
[1360250999] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360251017] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360251039] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360251150] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 43.74, 31.73, 17.84
[1360251216] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360251240] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (55.2721088435374)
[1360251240] SERVICE FLAPPING ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1360251292] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 13.23, 10.18, 5.46
[1360251330] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL CRITICAL - io read latency=121 ms
[1360251374] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360251394] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=92 ms
[1360251470] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 45.44, 39.46, 25.16
[1360251499] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.00% system=0.00% iowait=100.00% idle=0.00%
[1360251532] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360251532] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=53 ms
[1360251611] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 13.22, 12.12, 7.69
[1360251644] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;3;P1.PL OK - io read latency=38 ms
[1360251699] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360251716] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1360251790] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 36.09, 37.76, 28.56
[1360251790] SERVICE FLAPPING ALERT: dvorak-test.colo.elex.be;LOAD;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1360251814] SERVICE ALERT: crowd-test.colo.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.20% system=0.00% iowait=33.13% idle=66.67%
[1360251840] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360251852] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=57 ms
[1360251892] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1360251930] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 10.54, 11.55, 8.77
[1360251930] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;LOAD;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1360251952] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360251959] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360252021] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=56 ms
[1360252032] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1360252059] SERVICE ALERT: bipmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.80% system=0.60% iowait=94.41% idle=%
[1360252209] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360252232] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=109 ms
[1360252274] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360252279] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360252335] HOST ALERT: tanja.printers.sensors.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360252339] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=28 ms
[1360252340] SERVICE ALERT: sda32-test.sensors.elex.be;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1360252379] HOST ALERT: tanja.printers.sensors.elex.be;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.39 ms
[1360252379] SERVICE ALERT: bipmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;2;CPU IOWAIT CRITICAL : user=0.80% system=0.60% iowait=95.80% idle=%
[1360252451] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=93 ms
[1360252499] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;CRITICAL;SOFT;1;(Service Check Timed Out)
[1360252584] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360252604] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1360252609] HOST FLAPPING ALERT: temp1.sofia.elex.be;STOPPED; Host appears to have stopped flapping (0.0% change < 5.0% threshold)
[1360252649] SERVICE ALERT: sda32-test.sensors.elex.be;SSH;OK;SOFT;2;SSH OK - OpenSSH_3.8.1p1 Debian-8.sarge.4 (protocol 2.0)
[1360252699] SERVICE ALERT: bipmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;3;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.40% idle=%
[1360252769] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;OK;SOFT;2;OK: Opening BINARY mode data connection for ftp_check (17520000 bytes). File send OK.
[1360252771] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360252789] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360252952] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=45 ms
[1360253112] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 3.89, 5.89, 8.08
[1360253162] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360253181] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 10.08, 8.89, 4.14
[1360253210] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360253329] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1360253459] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360253529] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=68 ms
[1360253533] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;OK;HARD;3;OK - load average: 3.02, 3.81, 6.29
[1360253541] SERVICE ALERT: ewaf-test.colo.elex.be;LOAD;OK;SOFT;2;OK - load average: 2.27, 3.40, 3.08
[1360253574] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360253755] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360253781] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360253892] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360253913] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=15 ms
[1360253959] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360254070] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;OK;HARD;3;PROCS OK: 1 process with UID = 700 (oracle), args 'tnslsnr'
[1360254082] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360254104] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=125 ms
[1360254135] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;OK;HARD;3;OK: no outstanding alert(s) in EM
[1360254142] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;OK;HARD;3;OK - ARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.3.0 - 64bit Production)
[1360254155] HOST ALERT: temp1.sofia.elex.be;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1360254170] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;OK;HARD;3;TCP OK - 0.000 second response time on port 1521
[1360254172] HOST ALERT: temp1.sofia.elex.be;UP;SOFT;2;PING OK - Packet loss = 85%, RTA = 80.13 ms
[1360254210] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360254279] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=67 ms
[1360254450] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360254484] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=83 ms
[1360254520] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;CRITICAL: SQL query failed!
[1360254522] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360254533] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 6.62, 9.00, 15.95
[1360254551] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;1;Connection refused
[1360254602] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360254659] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360254672] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360254769] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360254840] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;2;CRITICAL: SQL query failed!
[1360254843] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360254849] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360254864] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=27 ms
[1360254870] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;2;Connection refused
[1360254971] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360254975] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360254991] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1360255091] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360255091] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;TNSLSNR;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1360255093] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=120 ms
[1360255100] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360255169] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360255189] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;HARD;3;Connection refused
[1360255189] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;TNS;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1360255193] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;CRITICAL;SOFT;1;(Service Check Timed Out)
[1360255199] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360255199] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1360255199] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360255199] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1360255243] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=78 ms
[1360255284] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=91 ms
[1360255286] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 42.49, 30.19, 21.06
[1360255306] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1360255359] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360255410] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=31 ms
[1360255410] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360255430] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 12.28, 9.20, 5.52
[1360255471] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360255503] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;OK;SOFT;2;OK: Opening BINARY mode data connection for ftp_check (17520000 bytes). File send OK.
[1360255543] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360255560] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1360255603] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=39 ms
[1360255679] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=33 ms
[1360255686] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=35 ms
[1360255750] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 14.01, 11.87, 7.70
[1360255786] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=59 ms
[1360255786] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=9 ms
[1360255829] SERVICE ALERT: bipmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.40% system=0.80% iowait=85.83% idle=%
[1360255842] SERVICE ALERT: rman-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;(Service Check Timed Out)
[1360255862] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=78 ms
[1360255881] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=32 ms
[1360256052] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=46 ms
[1360256064] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=102 ms
[1360256071] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 11.36, 11.22, 8.64
[1360256105] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL CRITICAL - io read latency=116 ms
[1360256127] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1360256149] SERVICE ALERT: rman-test.apps.elex.be;EM_ALERTS;OK;SOFT;2;OK: no outstanding alert(s) in EM
[1360256149] SERVICE ALERT: bipmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=12.00% idle=%
[1360256186] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360256280] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;CRITICAL;SOFT;1;(Service Check Timed Out)
[1360256372] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360256372] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1360256384] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=47 ms
[1360256467] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360256468] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL CRITICAL - io read latency=147 ms
[1360256499] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1360256640] SERVICE ALERT: openft-test.elex.be;FTP_FETCH;OK;SOFT;2;OK: Opening BINARY mode data connection for ftp_check (17520000 bytes). File send OK.
[1360256648] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;HARD;3;OK Power1-Allocation (55.9037900874636)
[1360256765] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=30 ms
[1360256815] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360256836] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360256838] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360256846] SERVICE ALERT: merlin-test.sofia.elex.be;CFAGENT;OK;SOFT;3;OK: Connection to cfengine master OK
[1360256866] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;HARD;3;P1.PL WARNING - io read latency=55 ms
[1360257135] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360257147] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=111 ms
[1360257156] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360257179] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1360257246] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;HARD;3;P1.PL CRITICAL - io read latency=116 ms
[1360257267] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 3.75, 6.80, 8.63
[1360257385] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=63 ms
[1360257458] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360257468] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=23 ms
[1360257518] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360257558] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360257626] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=0 ms
[1360257626] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO read;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1360257705] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1360257778] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360257837] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360257848] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360257885] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=59 ms
[1360258027] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1360258027] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;OK;HARD;3;OK - load average: 1.30, 2.87, 5.55
[1360258098] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1360258158] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=62 ms
[1360258168] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360258176] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;OK;HARD;3;PROCS OK: 1 process with UID = 700 (oracle), args 'tnslsnr'
[1360258198] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=201 ms
[1360258250] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;OK;HARD;3;OK: no outstanding alert(s) in EM
[1360258257] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;OK;HARD;3;OK - ARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.3.0 - 64bit Production)
[1360258276] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;OK;HARD;3;TCP OK - 0.000 second response time on port 1521
[1360258376] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 5.17, 6.46, 16.20
[1360258485] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=35 ms
[1360258579] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=67 ms
[1360258757] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 4.97, 5.65, 12.58
[1360258781] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360258865] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1360258918] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=28 ms
[1360258958] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=23 ms
[1360259100] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=83 ms
[1360259185] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360259267] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360259337] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360259421] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=22 ms
[1360259508] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1360259586] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (55.2721088435374)
[1360259659] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=42 ms
[1360259739] SERVICE ALERT: dvorak-test.colo.elex.be;DISK;WARNING;SOFT;1;[ /vservers ] -> Size: 124G Used: 98G Avail: 27G Free: (20%)
[1360259799] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360259882] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=73 ms
[1360259977] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=45 ms
[1360259977] SERVICE FLAPPING ALERT: geofront-test.colo.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1360260009] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360260118] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=55 ms
[1360260121] SERVICE ALERT: dvorak-test.colo.elex.be;DISK;WARNING;SOFT;2;[ /vservers ] -> Size: 124G Used: 99G Avail: 26G Free: (20%)
[1360260185] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=103 ms
[1360260273] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1360260375] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=74 ms
[1360260388] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360260428] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=34 ms
[1360260451] SERVICE ALERT: oraclexpeqt-stage.sofia.elex.be;CFAGENT;OK;HARD;3;OK: Connection to cfengine master OK
[1360260475] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360260476] SERVICE ALERT: dvorak-test.colo.elex.be;DISK;WARNING;HARD;3;[ /vservers ] -> Size: 124G Used: 99G Avail: 26G Free: (20%)
[1360260515] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=28 ms
[1360260517] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360260597] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=53 ms
[1360260695] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=8 ms
[1360260767] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=98 ms
[1360260836] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.9037900874636)
[1360260852] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360260975] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=40 ms
[1360261015] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360261089] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360261135] SERVICE ALERT: linkiev-test.kiev.elex.be;BULIBSYNC;OK;HARD;3;OK: /home/pol/bulibsync_my/TEST.txt was updated in the last 24 hours.
[1360261175] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=31 ms
[1360261272] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360261335] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360261550] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=91 ms
[1360261560] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360261590] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360261657] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;3;OK Power1-Allocation (57.555879494655)
[1360261867] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360261879] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360261968] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360262115] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360262197] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=93 ms
[1360262289] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=34 ms
[1360262429] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=95 ms
[1360262570] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=44 ms
[1360262625] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360262747] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=54 ms
[1360262786] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;TNSLSNR;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1360262858] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1360262860] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1360262878] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;TNS;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1360262937] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=43 ms
[1360263252] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=50 ms
[1360263415] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360263415] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360263436] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360263495] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=29 ms
[1360263768] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=36 ms
[1360263790] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1360263797] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.35.80.176'
[1360263805] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=17 ms
[1360263832] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;2;ERROR: No response from remote host '10.32.241.17'
[1360263897] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=45 ms
[1360263945] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360264108] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=100 ms
[1360264126] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;HARD;3;ERROR: No response from remote host '10.35.80.176'
[1360264151] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 2 usage;OK;SOFT;3;OK Power2-Allocation (55.2721088435374)
[1360264158] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360264212] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=7 ms
[1360264265] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360264425] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360264476] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=29 ms
[1360264558] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360264798] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360264829] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=39 ms
[1360264829] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1360264849] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=44 ms
[1360264855] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360264875] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=16 ms
[1360264947] SERVICE FLAPPING ALERT: shifu-test.colo.elex.be;ESXi IO read;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360264949] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=69 ms
[1360264956] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;LOAD;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1360265165] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1360265167] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=152 ms
[1360265176] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360265178] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=74 ms
[1360265208] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=163 ms
[1360265255] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=182 ms
[1360265267] SERVICE ALERT: bumblebee-test.erfurt.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=18 ms
[1360265299] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=72 ms
[1360265486] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360265497] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=71 ms
[1360265507] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360265522] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360265565] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=27 ms
[1360265615] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=7 ms
[1360265640] SERVICE ALERT: mainrouter-test.kiev.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.20% system=0.20% iowait=99.60% idle=0.00%
[1360265765] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=347 ms
[1360265765] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=238 ms
[1360265775] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=341 ms
[1360265779] SERVICE ALERT: ftp-test.sofia.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=0.00% system=0.20% iowait=70.60% idle=29.20%
[1360265816] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=23 ms
[1360265829] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360265955] SERVICE ALERT: mainrouter-test.kiev.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=0.00% idle=100.00%
[1360266029] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360266058] SERVICE FLAPPING ALERT: dvorak-test.colo.elex.be;LOAD;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1360266081] SERVICE ALERT: medix-test.sofia.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=3 ms
[1360266085] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=39 ms
[1360266087] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360266095] SERVICE ALERT: ftp-test.sofia.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.20% iowait=0.60% idle=99.20%
[1360266319] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=182 ms
[1360266349] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=62 ms
[1360266370] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=604 ms
[1360266467] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360266637] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=26 ms
[1360266656] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360266668] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=70 ms
[1360266685] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=11 ms
[1360266785] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360266856] SERVICE ALERT: cmdb-test.elex.be;HTTP_8080;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 2385 bytes in 0.019 seconds
[1360266945] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360266977] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360267160] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=48 ms
[1360267327] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360267488] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=12 ms
[1360267548] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=222 ms
[1360267707] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360267834] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360267866] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=160 ms
[1360267980] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;HARD;3;OK Power2-Allocation (54.7862001943635)
[1360268026] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360268044] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=60 ms
[1360268154] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=69 ms
[1360268241] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=42 ms
[1360268364] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360268464] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=85 ms
[1360268776] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=79 ms
[1360268844] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=26 ms
[1360268927] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360268994] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1360269098] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360269239] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=77 ms
[1360269270] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360269368] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360269559] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1360269594] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=40 ms
[1360269687] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=20 ms
[1360269749] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360269849] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=82 ms
[1360269977] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360270067] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=12 ms
[1360270165] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360270292] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=41 ms
[1360270314] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=56 ms
[1360270436] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360270444] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=77 ms
[1360270508] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360270517] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;CRITICAL: SQL query failed!
[1360270518] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;1;CRITICAL - SQL query error!
[1360270528] SERVICE ALERT: kili-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360270535] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;1;Connection refused
[1360270538] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=54 ms
[1360270610] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=21 ms
[1360270629] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360270748] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360270757] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=65 ms
[1360270847] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=96 ms
[1360270871] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360270874] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.49.32.16'
[1360270886] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;2;CRITICAL: SQL query failed!
[1360270889] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;SOFT;2;CRITICAL - SQL query error!
[1360270898] SERVICE ALERT: kili-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=5 ms
[1360270914] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;SOFT;2;Connection refused
[1360270924] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;2;P1.PL CRITICAL - io write latency=104 ms
[1360271035] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.20% system=0.10% iowait=99.70% idle=0.00%
[1360271067] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360271075] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 33.53, 22.77, 13.00
[1360271099] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;SOFT;1;CRITICAL: load average: 12.04, 8.07, 4.02
[1360271104] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;1;P1.PL CRITICAL - io read latency=141 ms
[1360271127] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with UID = 700 (oracle), args 'tnslsnr'
[1360271137] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=61 ms
[1360271224] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360271247] SERVICE ALERT: blade1mm1.erfurt.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.5150631681244)
[1360271264] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;CRITICAL;HARD;3;Connection refused
[1360271265] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;CRITICAL;HARD;3;CRITICAL: SQL query failed!
[1360271265] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;CRITICAL;HARD;3;CRITICAL - SQL query error!
[1360271266] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=19 ms
[1360271354] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;2;CPU IOWAIT CRITICAL : user=0.30% system=0.90% iowait=98.80% idle=0.00%
[1360271387] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=15 ms
[1360271395] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;SOFT;2;CRITICAL: load average: 38.29, 32.92, 20.25
[1360271425] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;SOFT;2;WARNING: load average: 11.55, 10.55, 6.32
[1360271426] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;CRITICAL;SOFT;2;P1.PL CRITICAL - io read latency=115 ms
[1360271438] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=52 ms
[1360271517] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=18 ms
[1360271545] SERVICE ALERT: pmelei-test.apps.elex.be;CPU_STATS;CRITICAL;SOFT;1;CPU IOWAIT CRITICAL : user=0.80% system=0.40% iowait=94.61% idle=%
[1360271598] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1360271638] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=87 ms
[1360271644] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360271664] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;WARNING;SOFT;1;CPU IOWAIT WARNING : user=0.00% system=0.90% iowait=74.20% idle=24.90%
[1360271674] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;CRITICAL;HARD;3;CPU IOWAIT CRITICAL : user=0.40% system=0.30% iowait=99.30% idle=0.00%
[1360271715] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 34.74, 34.49, 24.67
[1360271746] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 12.11, 11.71, 8.07
[1360271746] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;WARNING;HARD;3;P1.PL WARNING - io read latency=59 ms
[1360271754] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360271807] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360271836] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360271864] SERVICE ALERT: pmelei-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=0.00% system=0.00% iowait=17.23% idle=%
[1360271897] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=53 ms
[1360271916] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=44 ms
[1360271918] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=57 ms
[1360271958] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360271959] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=24 ms
[1360271965] SERVICE ALERT: rman-test.apps.elex.be;EM_ALERTS;CRITICAL;SOFT;1;(Service Check Timed Out)
[1360271984] SERVICE ALERT: viiperdb-test.apps.elex.be;CPU_STATS;OK;SOFT;2;CPU STATISTICS OK : user=23.10% system=2.65% iowait=50.20% idle=24.05%
[1360272119] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 9.32, 10.56, 8.81
[1360272119] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO read;OK;HARD;3;P1.PL OK - io read latency=0 ms
[1360272127] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=32 ms
[1360272156] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 2 usage;OK;SOFT;2;OK Power2-Allocation (54.7862001943635)
[1360272208] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=48 ms
[1360272234] SERVICE ALERT: shakira-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360272234] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=80 ms
[1360272254] SERVICE ALERT: rman-test.apps.elex.be;EM_ALERTS;OK;SOFT;2;OK: no outstanding alert(s) in EM
[1360272409] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.32.241.17'
[1360272434] SERVICE ALERT: orpmelei-test.apps.elex.be;CPU_STATS;OK;HARD;3;CPU STATISTICS OK : user=0.30% system=0.10% iowait=13.81% idle=85.79%
[1360272498] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;CRITICAL;HARD;3;CRITICAL: load average: 10.45, 10.80, 9.54
[1360272528] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=17 ms
[1360272614] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=22 ms
[1360272728] SERVICE ALERT: blade2mm1.colo.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (55.9037900874636)
[1360272876] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 7.99, 9.83, 9.58
[1360272876] SERVICE FLAPPING ALERT: viiperdb-test.apps.elex.be;LOAD;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1360272897] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=49 ms
[1360272994] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=64 ms
[1360273089] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=84 ms
[1360273214] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360273214] SERVICE FLAPPING ALERT: elaine-test.sensors.elex.be;ESXi IO write;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1360273264] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=66 ms
[1360273311] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=63 ms
[1360273408] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=21 ms
[1360273584] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=38 ms
[1360273628] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1360273628] SERVICE ALERT: viiperdb-test.apps.elex.be;LOAD;OK;HARD;3;OK - load average: 3.14, 4.22, 6.52
[1360274007] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=59 ms
[1360274044] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=80 ms
[1360274148] SERVICE ALERT: viiperdb-test.apps.elex.be;TNSLSNR;OK;HARD;3;PROCS OK: 1 process with UID = 700 (oracle), args 'tnslsnr'
[1360274166] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=97 ms
[1360274268] SERVICE ALERT: viiperdb-test.apps.elex.be;EM_ALERTS;OK;HARD;3;OK: no outstanding alert(s) in EM
[1360274269] SERVICE ALERT: viiperdb-test.apps.elex.be;ARCHIVELOG;OK;HARD;3;OK - ARCHIVELOG (Oracle Database 11g Enterprise Edition Release 11.2.0.3.0 - 64bit Production)
[1360274276] SERVICE ALERT: viiperdb-test.apps.elex.be;TNS;OK;HARD;3;TCP OK - 0.000 second response time on port 1521
[1360274327] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=68 ms
[1360274358] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=25 ms
[1360274547] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=81 ms
[1360274575] SERVICE ALERT: viiper-test.apps.elex.be;TNSLSNR_PROCESS;OK;HARD;3;PROCS OK: 1 process with UID = 59998 (applmgr), args 'tnslsnr'
[1360274717] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=31 ms
[1360274806] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=42 ms
[1360274820] SERVICE ALERT: viiper-test.apps.elex.be;CRM_Process;OK;HARD;3;PROCS OK: 1 process with UID = 59998 (applmgr), args 'FNDCRM'
[1360274878] SERVICE ALERT: viiper-test.apps.elex.be;FNDLIBR_Process;OK;HARD;3;PROCS OK: 5 processes with UID = 59998 (applmgr), args 'FNDLIBR'
[1360274886] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;CRITICAL;HARD;3;P1.PL CRITICAL - io write latency=160 ms
[1360275076] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Process_transaction_interface;OK;HARD;3;OK - 'Process_transaction_interface' Completed Normal      07/02/2013 23:03
[1360275120] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=22 ms
[1360275137] SERVICE ALERT: viiper-test.apps.elex.be;APPS_APACHE;OK;HARD;3;PROCS OK: 12 processes with UID = 59998 (applmgr), args 'Apache'
[1360275156] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_WIP_Move_Transaction_Manager;OK;HARD;3;OK - 'WIP_Move_Transaction_Manager' Completed Normal      07/02/2013 23:04
[1360275176] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;WARNING;HARD;3;WARNING: load average: 12.18, 13.94, 17.33
[1360275187] SERVICE ALERT: viiper-test.apps.elex.be;APPS_HTTP;OK;HARD;3;HTTP OK HTTP/1.1 200 OK - 1560 bytes in 0.001 seconds
[1360275258] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=20 ms
[1360275346] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360275406] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;CRITICAL;SOFT;1;CRITICAL: Unable to establish connection with cfengine master server
[1360275470] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=61 ms
[1360275504] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=62 ms
[1360275538] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=47 ms
[1360275637] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=43 ms
[1360275664] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=52 ms
[1360275688] SERVICE ALERT: viiper-test.apps.elex.be;CONC_REQ_Manager_Lot_Move_Transactions;OK;HARD;3;OK - 'Manager:_Lot_Move_Transactions' Completed Normal      07/02/2013 23:11
[1360275727] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;CRITICAL;SOFT;2;CRITICAL: Unable to establish connection with cfengine master server
[1360275784] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=50 ms
[1360275817] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=49 ms
[1360275854] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=28 ms
[1360275924] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=45 ms
[1360275926] SERVICE ALERT: dvorak-test.colo.elex.be;LOAD;OK;HARD;3;OK - load average: 3.83, 5.62, 10.84
[1360275949] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=55 ms
[1360275978] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=25 ms
[1360276048] SERVICE ALERT: ldap-test.colo.elex.be;CFAGENT;CRITICAL;HARD;3;CRITICAL: Unable to establish connection with cfengine master server
[1360276104] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=72 ms
[1360276135] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=13 ms
[1360276234] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=67 ms
[1360276240] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=58 ms
[1360276265] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=20 ms
[1360276509] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=57 ms
[1360276550] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=78 ms
[1360276556] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;HARD;3;P1.PL WARNING - io write latency=59 ms
[1360276825] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=19 ms
[1360276856] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=38 ms
[1360276869] SERVICE ALERT: jaguar-test.kiev.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=38 ms
[1360277018] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=86 ms
[1360277116] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=41 ms
[1360277187] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO read;WARNING;SOFT;1;P1.PL WARNING - io read latency=48 ms
[1360277309] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;OK;HARD;3;P1.PL OK - io write latency=24 ms
[1360277337] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;WARNING;SOFT;2;P1.PL WARNING - io write latency=72 ms
[1360277395] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;UNKNOWN;SOFT;1;ERROR: No response from remote host '10.35.80.176'
[1360277430] SERVICE ALERT: hammy-test.sofia.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=35 ms
[1360277506] SERVICE ALERT: geofront-test.colo.elex.be;ESXi IO read;OK;SOFT;2;P1.PL OK - io read latency=0 ms
[1360277614] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;CRITICAL;SOFT;1;P1.PL CRITICAL - io write latency=94 ms
[1360277654] SERVICE ALERT: shifu-test.colo.elex.be;ESXi IO write;OK;SOFT;3;P1.PL OK - io write latency=26 ms
[1360277687] SERVICE ALERT: elaine-test.sensors.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=50 ms
[1360277687] SERVICE ALERT: oin-test.colo.elex.be;ESXi IO write;WARNING;SOFT;1;P1.PL WARNING - io write latency=51 ms
[1360277714] SERVICE ALERT: blade1mm1.sofia.elex.be;Bladecenter Power 1 usage;OK;SOFT;2;OK Power1-Allocation (57.555879494655)
[1360277934] SERVICE ALERT: zonda-test.kiev.elex.be;ESXi IO write;OK;SOFT;2;P1.PL OK - io write latency=37 ms